<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

## Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 11.12.2016 um 23:59 Uhr
### Wichtig: Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue06.ipynb ab (nicht umbenennen, keine Kopie erstellen!). Sonst kann eine Bewertung nicht stattfinden.

# Aufgabe 1: King-Kong (? Punkte)
### a) 
Implementieren Sie eine Funktion ```king_kong(x)```, welche für eine gegebene natürliche Zahl ```x```
- 'King' zurückgibt, falls ```x``` durch 3 teilbar ist oder die Ziffer 3 enthält
- 'Kong' zurückgibt, falls ```x``` durch 7 teilbar ist oder die Ziffer 7 enthält
- "KingKong" zurückgibt, falls ```x``` die beiden oberen Bedingungen erfüllt
- ```x``` zurückgibt, falls keine der oberen drei Bedingungen erfüllt wurden.

(? Punkte)

<!-- grade_id: kingkong_sol_1, points: 0, grade: false, solution: true -->
```python
def king_kong(x):
    # Your code here
    ### BEGIN SOLUTION
    r = x
    king_erfuellt = False
    if x % 3 == 0 or '3' in str(x):
        r = 'King'
        king_erfuellt = True
    if x % 7 == 0 or '7' in  str(x):
        if king_erfuellt == True:
            return 'KingKong'
        return 'Kong'
    return r
    ### END SOLUTION
```

<!-- grade_id: kongkong_test_1_0, points: 0, grade: true, solution: false -->
```python
assert(king_kong(1) == 1)
assert(king_kong(2) == 2)
assert(king_kong(3) == 'King')
assert(king_kong(6) == 'King')
assert(king_kong(7) == 'Kong')
assert(king_kong(17) == 'Kong')
assert(king_kong(28) == 'Kong')
assert(king_kong(21) == 'KingKong')
assert(king_kong(37) == 'KingKong')
```

<!-- grade_id: kingkong_test_1_1, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für durch 3 Teilbare Zahlen:
for i in range(1,100):
    if i % 3 == 0 and not (i % 7 == 0 or '7' in str(i)):
        assert(king_kong(i) == 'King')
### HIDEEND
```

<!-- grade_id: kingkong_test_1_2, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für durch 7 Teilbare Zahlen:
for i in range(1,100):
    if i % 7 == 0 and not (i % 3 == 0 or '3' in str(i)):
        assert(king_kong(i) == 'Kong')
### HIDEEND
```

<!-- grade_id: kingkong_test_1_3, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für Zahlen die 3, 7 oder beide Ziffern enthalten:
for i in range(1,100):
    if (i % 3 == 0 or '3' in str(i)) and (i % 7 == 0 or '7' in str(i)):
        assert(king_kong(i) == 'KingKong')
### HIDEEND
```

<!-- grade_id: kingkong_test_1_4, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Zahlen, für welche kein String zurückgegeben wird:
for i in range(1,100):
    if not((i % 3 == 0 or '3' in str(i)) or (i % 7 == 0 or '7' in str(i))):
        assert(king_kong(i) == i)
### HIDEEND
```

### b) 
Implementieren Sie eine Funktion ```king_kong_liste(n)```, welche für eine gegebene natürliche Zahl ```n``` eine Liste der ersten n Zahlen zurückgibt, für welche die Funktion ```king_kong(n)``` aus Teil a) aufgerufen wurde.

<!-- grade_id: kingkong_sol_2, points: 0, grade: false, solution: true -->
```python
def king_kong_liste(n):
    # Your code here
    ### BEGIN SOLUTION
    l = []
    for i in range(1,n+1):
        l.append(king_kong(i))
    return l
    ### END SOLUTION
```

<!-- grade_id: kingkong_test_2_0, points: 0, grade: true, solution: false -->
```python
assert(king_kong_liste(10) == [1,2,'King',4,5,'King','Kong',8,'King',10])
```

<!-- grade_id: kingkong_test_2_1, points: 2, grade: true, solution: false -->
```python
### HIDESTART
def king_kong_liste_sol(n):
    l = []
    for i in range(1,n+1):
        l.append(king_kong(i))
    return l

assert(king_kong_liste(100) == king_kong_liste_sol(100))
### HIDEEND
```