<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
HIER und NUR HIER Matrikelnummer und email eintragen; diese Zeile entfernen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile). Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue05.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen!). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie  die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Bei fehlendem docstring werden automatisch KEINE Punkte vergeben! 


# Abgabetermin: 4.12.2016 um 23:59 Uhr

# Aufgabe 1: King-Kong 
### a) 
Implementieren Sie eine Funktion ```king_kong(x)```, welche für eine gegebene natürliche Zahl ```x```
- 'King' zurückgibt, falls ```x``` durch 3 teilbar ist oder die Ziffer 3 enthält
- 'Kong' zurückgibt, falls ```x``` durch 7 teilbar ist oder die Ziffer 7 enthält
- "KingKong" zurückgibt, falls ```x``` die beiden oberen Bedingungen erfüllt
- ```x``` zurückgibt, falls keine der oberen drei Bedingungen erfüllt wurden.


<!-- grade_id: kingkong_sol_1, points: 0, grade: false, solution: true -->
```python
def king_kong(x):
    # Your code here
    ### BEGIN SOLUTION
    r = x
    king_erfuellt = False
    if x % 3 == 0 or '3' in str(x):
        r = 'King'
        king_erfuellt = True
    if x % 7 == 0 or '7' in  str(x):
        if king_erfuellt == True:
            return 'KingKong'
        return 'Kong'
    return r
    ### END SOLUTION
```

<!-- grade_id: kongkong_test_1_0, points: 0, grade: true, solution: false -->
```python
assert(king_kong(1) == 1)
assert(king_kong(2) == 2)
assert(king_kong(3) == 'King')
assert(king_kong(6) == 'King')
assert(king_kong(7) == 'Kong')
assert(king_kong(17) == 'Kong')
assert(king_kong(28) == 'Kong')
assert(king_kong(21) == 'KingKong')
assert(king_kong(37) == 'KingKong')
```

<!-- grade_id: kingkong_test_1_1, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für durch 3 Teilbare Zahlen:
for i in range(1,30):
    if i % 3 == 0 and not (i % 7 == 0 or '7' in str(i)):
        assert(king_kong(i) == 'King')
        
# Teste Ausgaben für durch 7 Teilbare Zahlen:
for i in range(1,30):
    if i % 7 == 0 and not (i % 3 == 0 or '3' in str(i)):
        assert(king_kong(i) == 'Kong')
        
# Teste Ausgaben für Zahlen die 3, 7 oder beide Ziffern enthalten:
for i in range(1,30):
    if (i % 3 == 0 or '3' in str(i)) and (i % 7 == 0 or '7' in str(i)):
        assert(king_kong(i) == 'KingKong')
        
# Teste Zahlen, für welche kein String zurückgegeben wird:
for i in range(1,30):
    if not((i % 3 == 0 or '3' in str(i)) or (i % 7 == 0 or '7' in str(i))):
        assert(king_kong(i) == i)
### HIDEEND
```

### b) 
Implementieren Sie eine Funktion `king_kong_liste(n)`, welche für eine gegebene natürliche Zahl n eine Liste der Zahlen aus `{1...n}` zurückgibt, für welche die Funktion `king_kong` aus Teil a) eine Zahl zurückgibt.

Hinweis: Achten Sie auf Start- und Endindices. Mit einer List comprehension ist das ganz leicht.

<!-- grade_id: kingkong_sol_2, points: 0, grade: false, solution: true -->
```python
def king_kong_liste(n):
    # Your code here
    ### BEGIN SOLUTION
    return [x for x in range(1, n+1) if str(king_kong(x)).isnumeric()]
    ### END SOLUTION
```

<!-- grade_id: kingkong_test_2_0, points: 0, grade: true, solution: false -->
```python
assert(king_kong_liste(10) == [1,2,4,5,8,10])
```

<!-- grade_id: kingkong_test_2_1, points: 1, grade: true, solution: false -->
```python
### HIDESTART
assert(king_kong_liste(50) == [1, 2, 4, 5, 8, 10, 11, 16, 19, 20, 22, 25, 26, 29, 40, 41, 44, 46, 50])
### HIDEEND
```

# Aufgabe 2: Binomialkoeffizienten


Das Pascalsche Dreieck ist eine grafische Darstellung der Binomialkoeffizienten $\binom{n}{k}$, aus der sich eine einfache Berechnung dieser Koeffizienten ergibt. Dabei kann die Variable $n$ als Zeilenindex und $k$ als Spaltenindex interpretiert werden. Hinweis: https://de.wikipedia.org/wiki/Pascalsches_Dreieck

Leiten Sie aus dem Pascalschen Dreieck ein rekursives Vorgehen ab, um Binominalkoeffizienten zu berechnen.
Schreiben Sie eine Funktion namens ```pascal```, welche zwei ganzzahlige Parameter übergeben bekommt und die entsprechenden Koeffizienten des Pascalschen Dreiecks zurückgibt. Implementieren Sie die Funktion ```pascal``` mittels Rekursion.

<!-- grade_id: 27, points: 0, grade: false, solution: true -->
```python
def pascal(row, col):
### BEGIN SOLUTION        
    """Compute elements of Pascals triangle

    Keyword arguments:
    row -- (int) Row nr. 
    col -- (int) Column nr.
    
    Return:
    int -- pascal coefficient
    """
    if (col == 0 or col == row):
        return 1
    else:
        return pascal(row-1,col-1) + pascal(row-1,col)
### END SOLUTION
```

<!-- grade_id: PD_open_test, points: 1, grade: true, solution: false -->
```python
assert(pascal(0,0)==1)
assert(pascal(1,0)==1)
assert(pascal(1,1)==1)
assert(pascal(10,1)==10)
assert(pascal(10,5)==252)
```

<!-- grade_id: PD_hidden_test, points: 1, grade: true, solution: false -->
```python
### HIDESTART
assert(pascal(5,2)==10)
# TODO: weitere Tests einfügen 
### HIDEND
```

<!-- grade_id: PD_hidden_test_2, points: 0, grade: false, solution: true -->
```python
### HIDESTART
assert(pascal.__doc__!=None)
# TODO: weitere Tests einfügen 
### HIDEND
```

##  c) 
Schreiben Sie eine Funktion ```pd_print```, die die ersten n Reihen des pascalschen Dreiecks in folgenderm Format ausgibt:
```
     1 
    1 1 
   1 2 1 
  1 3 3 1 
 1 4 6 4 1 
 ```

<!-- grade_id: 29, points: 1, grade: true, solution: true -->
```python
def pd_print(n):
### BEGIN SOLUTION
"""Print elements of Pascal-triangle

    Keyword arguments:
    n -- (int) nr. of rows
    
    Return:
    None
    """
    for i in range(0,n):
        print(" "*(n-i-1),end="")
        for j in range(0,i+1):
            print(pascal(i,j),end=" ")
        print("\n",end="")
### END SOLUTION
```

Output:
```
1 
   1 1 
  1 2 1 
 1 3 3 1 
1 4 6 4 1
```

<!-- grade_id: 30, points: 0, grade: true, solution: false -->
```python
### HIDESTART
assert(pd_print.__doc__!=None)
# TODO: weitere Tests einfügen 
### HIDEND
```

# Aufgabe 3: Folgen implementieren

### a)
Schreiben Sie eine Funktion `berechne_folge(n)`, welche für eine gegebene natürliche Zahl `n` eine Liste mit den einzelnen Elemenenten der Folge ausgibt. Die einzelnen Elemente berechnen sich wie folgt:

1. Beginne mit irgendeiner natürlichen Zahl n > 0.
2. Ist n gerade, so nimm als nächstes n/2.
3. Ist n ungerade, so nimm als nächstes 3n + 1.
4. Wiederhole die Vorgehensweise mit der erhaltenen Zahl.

Der Algorithmus soll abbrechen, sobald eine 1 zu den Elementen der Folge hinzugefügt wurde.


<!-- grade_id: collatz_sol_1, points: 0, grade: false, solution: true -->
```python
def berechne_folge(n):
    # Your code here
    ### BEGIN SOLUTION
    L = []
    while(True):
        L.append(n)
        if n == 1:
            return L
        if n % 2 == 0:
            n = n//2
        else:
            n = 3*n+1
    ### END SOLUTION
```

<!-- grade_id: collatz_test_1, points: 2, grade: true, solution: false -->
```python
assert(berechne_folge(19) == [19, 58, 29, 88, 44, 22, 11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1])
assert(berechne_folge(15) == [15, 46, 23, 70, 35, 106, 53, 160, 80, 40, 20, 10, 5, 16, 8, 4, 2, 1])
### HIDESTART
assert(berechne_folge(7) == [7, 22, 11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1])
assert(berechne_folge(6) == [6, 3, 10, 5, 16, 8, 4, 2, 1])
assert(berechne_folge(5) == [5, 16, 8, 4, 2, 1])
### HIDEEND
```

###  b)
Begründen Sie, warum Sie in Ihrer Implementierung eine while- oder eine for-Schleife benutzt haben.

<!-- grade_id: collatz_sol_2, points: 1, grade: true, solution: true -->
- While-Schleife: Anzahl der Elemente der Folge ist nicht von vorneherein klar, daher bietet sich die While-Schleife an
- for-Schleife: Macht hier keinen Sinn

# Aufgabe 4
Schreiben Sie eine Funktion `matrix_mult(X,Y)`, welche zwei gegebene Matrizen `X` und `Y` multipliziert und die Ergebnismatrix zurückgibt.
Falls die Dimesionen der Matzrizen nicht korrekt sind, so soll `None` zurückgegeben werden.
Die Darstellung der Eingabematrizen läßt sich an den Tests ablesen.
Sie dürfen annehmen, dass die Dimension der Matrizen immer größer als 0 ist.

<!-- grade_id: matrixmult_sol, points: 0, grade: false, solution: true -->
```python
def matrix_mult(X,Y):
    # Your code here
    ### BEGIN SOLUTION
    # Teste Dimension:
    if len(X) != len(Y[0]):
        return None
    
    C = [[0 for col in range(len(Y[0]))] for row in range(len(X))]
    for i in range(len(X)):
        for j in range(len(Y[0])):
            for k in range(len(Y)):
                C[i][j] += X[i][k]*Y[k][j]
    return C
    ### END SOLUTION
```

<!-- grade_id: matrixmult_test, points: 3, grade: true, solution: false -->
```python
X = [[12,7,3], [4,5,6], [7,8,9]]
Y = [[5,8,1], [6,7,3], [4,5,9]]
r = matrix_mult(X,Y)
assert(r == [[114, 160, 60], [74, 97, 73], [119, 157, 112]])

# Falsche Dimension:
X = [[12,7,3], [4,5,6], [7,8,9]]
Y = [[5,8], [6,7], [4,5]]
assert(matrix_mult(X,Y) == None)

### HIDESTART
X = [[1,2,3], [4 ,5,3], [1 ,8,6]]
Y = [[5,4,1], [6,2,3], [4,1,3]]
r = matrix_mult(X,Y)
assert(r == [[29, 11, 16], [62, 29, 28], [77, 26, 43]])
### HIDEEND
```

# Aufgabe 5
Der Dozent einer Programmiervorlesung nutzt ein tolles neues Tool, um die Lösungen seiner Studierenden möglichst komfortabel bewerten zu können.
Leider erhält er von diesem Tool jedoch zwei getrennte Listen mit Matrikelnummern und Punkten, wie die folgenden:

```Python
matrikel = [12345, 12531, 18034, 17824]
punkte = [19, 13, 0, 20]
```

Dabei bezieht sich der erste Eintrag in der Liste "Punkte" auf die erste Matrikelnummer, usw.
Praktischer wäre es jedoch, ein ```dict``` zu haben, welches Matrikelnummern auf Puntke abbildet, z. B. so:

```Python
matrikel_punkte = {12345: 19, 12531: 13, 18034: 0, 17824: 20}
```

## a) 

Schreiben Sie eine Funktion namens ```verknuepfe(l1, l2)```, an die zwei Listen als Parameter übergeben werden  und die die Verknüpfung dieser beiden Listen als Liste von Tuplen zurückgibt (so wie ```matrikel_punkte``` eine Verknüpfung der Listen ```matrikel``` und ```punkte``` ist).

Hinweis: Verwenden Sie hierzu u. a. die Funktion zip(), https://docs.python.org/3/library/functions.html#zip

<!-- grade_id: zipa2, points: 0, grade: false, solution: true -->
```python
def verknuepfe(l1, l2):
    return dict(list(zip(l1, l2)))
```

<!-- grade_id: zipa3, points: 1, grade: true, solution: false -->
```python
matrikel = [12345, 12531, 18034, 17824]
punkte = [19, 13, 0, 20]
matrikel_punkte = {12345: 19, 12531: 13, 18034: 0, 17824: 20}
assert(matrikel_punkte == verknuepfe(matrikel,punkte))

### HIDESTART
assert({} == verknuepfe([], []))
### HIDEEND
```

## b) 
Leider stellt sich nach der Bewertung des dritten Übungsblattes heraus, dass das Bewertungs-Tool einen merkwürdigen Fehler hat:
Befindet sich am Ende der Liste ```matrikel``` eine Folge von Matrikelnummern, deren zugehörige Studierende alle null Punkte erreicht haben, so tauchen diese in der Liste ```punkte``` nicht auf.
Die liste ```punkte``` ist dann also kürzer als ```matrikel```.

Beispiel:
```Python
matrikel = [12345, 12531, 18034, 17824, 12635]
punkte = [6, 0, 3]
```

anstelle von:
```Python
punkte = [6, 0, 3, 0, 0]
```

Wie würde Ihre Funktion damit umgehen, dass die Längen der beiden Listen nicht übereinstimmen?

<!-- grade_id: zipb2, points: 1, grade: true, solution: true -->
Die betroffenen Matrikelnummern am Ende der Liste würden einfach nicht auftauchen.

## c) 
Schreiben Sie eine Funktion namens ```verknuepfe_besser(l1, l2)```, welche Ihre Funktion aus a) so erweitert, dass auch bei einer Eingabe wie in Teilaufgabe b) beschrieben, alle Matrikelnummern im zurückgegebenen ```dict``` auftauchen (mit den richtigen erzielten Punkten).

<!-- grade_id: zipb4, points: 0, grade: false, solution: true -->
```python
def verknuepfe_besser(l1, l2):
    return dict(list(zip(l1, l2 + [0] * (len(l1)-len(l2)))))
```

<!-- grade_id: zipb5, points: 1, grade: true, solution: false -->
```python
matrikel = [12345, 12531, 18034, 17824, 12635]
punkte = [6, 0, 3]
matrikel_punkte = {17824: 0, 12345: 6, 18034: 3, 12531: 0, 12635: 0}
assert(matrikel_punkte == verknuepfe_besser(matrikel,punkte))

#TODO: hidden asserts
```

## d) 
Nach Veröffentlichung der Ergebnisse stellte sich heraus, dass einige Nachkorrekturen nötig waren. 
Der Dozent möchte die korrigierten Ergebnisse nun wieder in das Tool geben.
Leider muss man dem Tool dazu wieder zwei Listen von der Form wie ```matrikel``` und ```punkte``` aus Teilaufgabe a) übergeben.
Ihm liegen die korrigierten Ergebnisse jedoch nur in der Form wie ```matrikel_punkte``` vor.

Schreiben Sie eine Funktion namens ```entknuepfe(l)```, welche ein solches ```dict``` übergeben bekommt und ein Paar (also 2-Tupel) der gewünschten Listen zurückgibt.

Verwenden Sie hierzu abermals die Funktion zip().

Hinweis: Sie müssen hier keine Sonderfälle (leeres ```dict``` o.ä.) beachten.

<!-- grade_id: zipc2, points: 0, grade: false, solution: true -->
```python
def entknuepfe(l):
### BEGIN SOLUTION
    entknuepft_als_liste_von_tuples = list(zip(*l.items()))
    l1 = list(entknuepft_als_liste_von_tuples[0])
    l2 = list(entknuepft_als_liste_von_tuples[1])
    
    return (l1, l2)
### END SOLUTION
```

<!-- grade_id: zipc3, points: 1, grade: true, solution: false -->
```python
matrikel_punkte = {12345: 19, 12531: 13, 18034: 0, 17824: 20}

matrikel_punkte_entknuepft = entknuepfe(matrikel_punkte)

matrikel = matrikel_punkte_entknuepft[0]
punkte = matrikel_punkte_entknuepft[1]

for i in range(len(matrikel)):
    assert(punkte[i] == matrikel_punkte[matrikel[i]])

### HIDESTART
 # Hier wuesste ich jetzt keinen guten hidden test...
### HIDEEND
```

# Aufgabe 6: Äquivalenz von Code
### a)
Gegeben sind die folgenden beiden Codefragmente:

```Python
L = [1, 2, 3, 4, 5]
for i in range(len(L)):
    L[i] += 1
print(L)
```

und 
```Python
L = [1, 2, 3, 4, 5]
L = [x+1 for x in L]
print(L)
```
Finden Sie Situationen, wo diese beiden Code-Abschnitte *nicht* den gleichen Effekt haben.
Hinweis: Was ist, wenn es andere Namen für die ursprünliche Liste gibt?

<!-- grade_id: comprehension_sol, points: 1, grade: true, solution: true -->
Die folgenden beiden Codefragmente sind nicht äquivalent:

```Python
L = [1, 2, 3, 4, 5]
T = L
for i in range(len(L)):
    L[i] += 1
print(T) # Gibt [2, 3, 4, 5, 6] aus
```

```Python
L = [1, 2, 3, 4, 5]
T = L
L = [x+1 for x in L]
print(T) # Gibt [1, 2, 3, 4, 5] aus
```


### b) 
Gegeben ist die folgende Funktion `schnitt(s1, s2)`, welche die Schnittmenge zweier Strings `s1` und `s2` berechnen sollen.

```Python
def schnitt(s1,s2):
    return [x for x in s1 if x in s2]
```

Gegeben sind nun zwei beliebige Strings `s1` und `s2`. Überlegen Sie sich, ob die Aufrufe `schnitt(s1,s2)` und `schnitt(s2,s1)` äquivalent sind. Begründen Sie Ihre Antwort.

<!-- grade_id: schnitt_sol, points: 1, grade: true, solution: true -->
Die Aufrufe sind nicht äquivalent, wie das folgende Beispiel zeigt.
```Python
def schnitt(s1,s2):
    return [x for x in s1 if x in s2]

s1 = 'aaabbb'
s2 = 'aabb'

print(schnitt(s1,s2)) # Gibt ['a', 'a', 'a', 'b', 'b', 'b'] aus
print(schnitt(s2,s1)) # Gibt ['a', 'a', 'b', 'b'] aus
```

Im ersten Aufruf wird für jedes `a` in `s1` überprüft, ob sich `a` in `s2` befindet.
Da dies der Fall ist, enthält die resultierende Liste 3 mal `'a'`.
Im zweiten Aufruf wird lediglich für zwei `'a'`s überprüft, ob sich `a` in `s1` befindet, wodurch die Ergebnisliste nur zwei mal `'a'` enthält.

# Aufgabe 7

Neben der *Anweisung*
```Python
if BEDINGUNG:
    Block1
else:
    Block2
```

kennt Python einen *Ausdruck*:
```Python
a if BEDINGUNG else b
```

Sie können Details hier nachlesen: https://docs.python.org/3/reference/expressions.html#conditional-expressions und https://www.python.org/dev/peps/pep-0308/ 
(Tutorial z.B. hier: 
http://pythoncentral.io/one-line-if-statement-in-python-ternary-conditional-operator/
)

## a) 
Setzen Sie ```Block1```, ```Block2``` und ```a```, ```b``` zueinander in Bezug - was entspricht was?

<!-- grade_id: ifa2, points: 1, grade: true, solution: true -->
```python
Block1 entspricht a, Block2 entspricht b.
```

## b) 
Sind ```a``` und  ```b``` Ausdrücke oder Blöcke?

<!-- grade_id: ifb2, points: 1, grade: true, solution: true -->
Es handelt sich hierbei um Ausdrücke (je nachdem, welcher Fall gilt, wird das Ergebnis des Gesamt-Ausdrucks einer dieser Ausdrücke).

# Aufgabe 8

Angenommen, Sie haben zwei geschachtelte Schleifen. 
In der inneren Schleife führen Sie ein ```continue``` durch. 
Geben Sie ein Code-Beispiel an, aus dem ersichtlich wird, ob nur die innere oder auch die äußere Schleife beeinflusst wird.

<!-- grade_id: schleifena2, points: 1, grade: true, solution: true -->
```python
for i in range(2):
    print('i:', i)    
    for j in range(100):
        if j >= 3:
            continue        
        print('j:', j)
    print('i after:', i)
```

Output:
```
i: 0
j: 0
j: 1
j: 2
i after: 0
i: 1
j: 0
j: 1
j: 2
i after: 1
```