<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

## Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 11.12.2016 um 23:59 Uhr
### Wichtig: Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue06.ipynb ab (nicht umbenennen, keine Kopie erstellen!). Sonst kann eine Bewertung nicht stattfinden.

# Aufgabe 1: King-Kong (? Punkte)
### a) 
Implementieren Sie eine Funktion ```king_kong(x)```, welche für eine gegebene natürliche Zahl ```x```
- 'King' zurückgibt, falls ```x``` durch 3 teilbar ist oder die Ziffer 3 enthält
- 'Kong' zurückgibt, falls ```x``` durch 7 teilbar ist oder die Ziffer 7 enthält
- "KingKong" zurückgibt, falls ```x``` die beiden oberen Bedingungen erfüllt
- ```x``` zurückgibt, falls keine der oberen drei Bedingungen erfüllt wurden.

(? Punkte)

<!-- grade_id: kingkong_sol_1, points: 0, grade: false, solution: true -->
```python
def king_kong(x):
    # Your code here
    ### BEGIN SOLUTION
    r = x
    king_erfuellt = False
    if x % 3 == 0 or '3' in str(x):
        r = 'King'
        king_erfuellt = True
    if x % 7 == 0 or '7' in  str(x):
        if king_erfuellt == True:
            return 'KingKong'
        return 'Kong'
    return r
    ### END SOLUTION
```

<!-- grade_id: kongkong_test_1_0, points: 0, grade: true, solution: false -->
```python
assert(king_kong(1) == 1)
assert(king_kong(2) == 2)
assert(king_kong(3) == 'King')
assert(king_kong(6) == 'King')
assert(king_kong(7) == 'Kong')
assert(king_kong(17) == 'Kong')
assert(king_kong(28) == 'Kong')
assert(king_kong(21) == 'KingKong')
assert(king_kong(37) == 'KingKong')
```

<!-- grade_id: kingkong_test_1_1, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für durch 3 Teilbare Zahlen:
for i in range(1,100):
    if i % 3 == 0 and not (i % 7 == 0 or '7' in str(i)):
        assert(king_kong(i) == 'King')
### HIDEEND
```

<!-- grade_id: kingkong_test_1_2, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für durch 7 Teilbare Zahlen:
for i in range(1,100):
    if i % 7 == 0 and not (i % 3 == 0 or '3' in str(i)):
        assert(king_kong(i) == 'Kong')
### HIDEEND
```

<!-- grade_id: kingkong_test_1_3, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für Zahlen die 3, 7 oder beide Ziffern enthalten:
for i in range(1,100):
    if (i % 3 == 0 or '3' in str(i)) and (i % 7 == 0 or '7' in str(i)):
        assert(king_kong(i) == 'KingKong')
### HIDEEND
```

<!-- grade_id: kingkong_test_1_4, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Zahlen, für welche kein String zurückgegeben wird:
for i in range(1,100):
    if not((i % 3 == 0 or '3' in str(i)) or (i % 7 == 0 or '7' in str(i))):
        assert(king_kong(i) == i)
### HIDEEND
```

### b) 
Implementieren Sie eine Funktion ```king_kong_liste(n)```, welche für eine gegebene natürliche Zahl ```n``` eine Liste der ersten n Zahlen zurückgibt, für welche die Funktion ```king_kong(n)``` aus Teil a) aufgerufen wurde.

<!-- grade_id: kingkong_sol_2, points: 0, grade: false, solution: true -->
```python
def king_kong_liste(n):
    # Your code here
    ### BEGIN SOLUTION
    l = []
    for i in range(1,n+1):
        l.append(king_kong(i))
    return l
    ### END SOLUTION
```

<!-- grade_id: kingkong_test_2_0, points: 0, grade: true, solution: false -->
```python
assert(king_kong_liste(10) == [1,2,'King',4,5,'King','Kong',8,'King',10])
```

<!-- grade_id: kingkong_test_2_1, points: 2, grade: true, solution: false -->
```python
### HIDESTART
def king_kong_liste_sol(n):
    l = []
    for i in range(1,n+1):
        l.append(king_kong(i))
    return l

assert(king_kong_liste(100) == king_kong_liste_sol(100))
### HIDEEND
```

## Aufgabe 2: Binomialkoeffizienten


a) Schreiben Sie eine Funktion namens ```fakultaet```, welche ein ganzzahliges Parameter übergeben bekommen und die entsprechenden Fakultät zurückgibt. Implementieren Sie die Funktion mit Rekursion und nicht mit Schleifen.

<!-- grade_id: 21, points: 0, grade: false, solution: true -->
```python
def fakultaet(zahl):
    if (zahl <= 1):
        return 1
    return fakultaet(zahl-1)*zahl

print(fakultaet(10))
```

Output:
```
3628800
```

<!-- grade_id: 22, points: 1, grade: true, solution: false -->
```python
assert(fakultaet(0)==1)
assert(fakultaet(1)==1)
assert(fakultaet(10)==3628800)
```

b) Schreiben Sie eine Funktion namens ```binom```, welche zwei ganzzahlige Parametern ```n, k``` übergeben bekommt und die entsprechende binomische Koeffizient 

$$\binom{n}{k} = \frac{n!}{k!(n-k)!} $$

zurückgibt. 

<!-- grade_id: 24, points: 0, grade: false, solution: true -->
```python
def combination(n,k):
    return int(fakultaet(n)/(fakultaet(k)*fakultaet(n-k)))
```

<!-- grade_id: 25, points: 1, grade: true, solution: false -->
```python
assert(combination(0,0)==1)
assert(combination(1,0)==1)
assert(combination(1,1)==1)
assert(combination(10,1)==10)
assert(combination(10,5)==252)
```

c) Das Pascalsche Dreieck ist eine Form der grafischen Darstellung der Binomialkoeffizienten $\binom{n}{k}$, die eine einfache Berechnung dieser erlaubt. Dabei kann die Variable $n$ als Zeilenindex und $k$ als Spaltenindex interpretiert werden, wobei die Zählung mit Null beginnt. Hinweis: https://de.wikipedia.org/wiki/Pascalsches_Dreieck

Schreiben Sie eine Funktion namens ```pascal```, welche zwei ganzzahlige Parametern übergeben bekommt und die entsprechende Koeffizient vom Pascalschen Dreieck zurückgibt. Benutzen Sie dabei nicht die Funktion ```combination``` und implementieren Sie die Funktion mittels Rekursion.

<!-- grade_id: 27, points: 0, grade: false, solution: true -->
```python
def pascal(row, col):
    if (col == 0 or col == row):
        return 1
    else:
        return pascal(row-1,col-1) + pascal(row-1,col)
```

d) Schreiben Sie ein Programm, dass die 5 ersten Reihen vom pascalschen Dreieck in folgender Format ausgibt:
     1 
    1 1 
   1 2 1 
  1 3 3 1 
 1 4 6 4 1 

<!-- grade_id: 29, points: 1, grade: true, solution: true -->
```python
n=5
for i in range(0,n):
    print(" "*(n-i-1),end="")
    for j in range(0,i+1):
        print(pascal(i,j),end=" ")
    print("\n",end="")
```

Output:
```
1 
   1 1 
  1 2 1 
 1 3 3 1 
1 4 6 4 1
```

## Aufgabe 2: Kalendar

http://www.ntu.edu.sg/home/ehchua/programming/java/j2a_basicsexercises.html
DateUtil function (at the end of the page)



# Aufgabe 2: Folgen implementieren

### a)
Schreiben Sie eine Funktion `berechne_folge(n)`, welche für eine gegebene natürliche Zahl `n` eine Liste mit den einzelnen Elemenenten der Folge ausgibt. Die einzelnen Elemente berechnen sich wiefolgt:

1. Beginne mit irgendeiner natürlichen Zahl n > 0.
2. Ist n gerade, so nimm als nächstes n/2.
3. Ist n ungerade, so nimm als nächstes 3n + 1.
4. Wiederhole die Vorgehensweise mit der erhaltenen Zahl.

Der Algorithmus soll hier abbrechen, sobald eine 1 zu den Elementen der Folge hinzugefügt wurde.


<!-- grade_id: collatz_sol_1, points: 0, grade: false, solution: true -->
```python
def berechne_folge(n):
    # Your code here
    ### BEGIN SOLUTION
    L = []
    while(True):
        L.append(n)
        if n == 1:
            return L
        if n % 2 == 0:
            n = n//2
        else:
            n = 3*n+1
    ### END SOLUTION
```

<!-- grade_id: collatz_test_1, points: 2, grade: true, solution: false -->
```python
assert(berechne_folge(19) == [19, 58, 29, 88, 44, 22, 11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1])
assert(berechne_folge(15) == [15, 46, 23, 70, 35, 106, 53, 160, 80, 40, 20, 10, 5, 16, 8, 4, 2, 1])
```

###  b)
Begründen Sie, warum Sie in Ihrer Implementierung eine while- oder eine for-Schleife benutzt haben.

<!-- grade_id: collatz_sol_2, points: 1, grade: true, solution: true -->
- While-Schleife: Anzahl der Elemente der Folge ist nicht von vorneherein klar, daher bietet sich die While-Schleife an
- for-Schleife: Macht hier keinen Sinn

### c)
Warum ist es nicht einfach die Terminierung der Schleife zu beweisen?

Zusatzaufgabe (gibt keine Punkte): Versuchen Sie, Terminierung zu beweisen. Achtung, SCHWER! 

<!-- grade_id: collatz_sol_3, points: 1, grade: true, solution: true -->
Es ist nicht einfach eine geeignete Schleifeninvariante anzugeben.   
Die zu implementierende Folge ist die Collatz-Folge, welche (vermutlich) immer in dem Zyklus (4,2,1) mündet. Doch der Beweis dieser Vermutung ist bis zum heutigen Zeitpunkt noch nicht gelungen.
Für mehr Informationen siehe https://de.wikipedia.org/wiki/Collatz-Problem.