<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

## Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 11.12.2016 um 23:59 Uhr
### Wichtig: Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue06.ipynb ab (nicht umbenennen, keine Kopie erstellen!). Sonst kann eine Bewertung nicht stattfinden.

# Aufgabe 1: King-Kong (? Punkte)
### a) 
Implementieren Sie eine Funktion ```king_kong(x)```, welche für eine gegebene natürliche Zahl ```x```
- 'King' zurückgibt, falls ```x``` durch 3 teilbar ist oder die Ziffer 3 enthält
- 'Kong' zurückgibt, falls ```x``` durch 7 teilbar ist oder die Ziffer 7 enthält
- "KingKong" zurückgibt, falls ```x``` die beiden oberen Bedingungen erfüllt
- ```x``` zurückgibt, falls keine der oberen drei Bedingungen erfüllt wurden.

(? Punkte)

<!-- grade_id: kingkong_sol_1, points: 0, grade: false, solution: true -->
```python
def king_kong(x):
    # Your code here
    ### BEGIN SOLUTION
    r = x
    king_erfuellt = False
    if x % 3 == 0 or '3' in str(x):
        r = 'King'
        king_erfuellt = True
    if x % 7 == 0 or '7' in  str(x):
        if king_erfuellt == True:
            return 'KingKong'
        return 'Kong'
    return r
    ### END SOLUTION
```

<!-- grade_id: kongkong_test_1_0, points: 0, grade: true, solution: false -->
```python
assert(king_kong(1) == 1)
assert(king_kong(2) == 2)
assert(king_kong(3) == 'King')
assert(king_kong(6) == 'King')
assert(king_kong(7) == 'Kong')
assert(king_kong(17) == 'Kong')
assert(king_kong(28) == 'Kong')
assert(king_kong(21) == 'KingKong')
assert(king_kong(37) == 'KingKong')
```

<!-- grade_id: kingkong_test_1_1, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für durch 3 Teilbare Zahlen:
for i in range(1,100):
    if i % 3 == 0 and not (i % 7 == 0 or '7' in str(i)):
        assert(king_kong(i) == 'King')
### HIDEEND
```

<!-- grade_id: kingkong_test_1_2, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für durch 7 Teilbare Zahlen:
for i in range(1,100):
    if i % 7 == 0 and not (i % 3 == 0 or '3' in str(i)):
        assert(king_kong(i) == 'Kong')
### HIDEEND
```

<!-- grade_id: kingkong_test_1_3, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Ausgaben für Zahlen die 3, 7 oder beide Ziffern enthalten:
for i in range(1,100):
    if (i % 3 == 0 or '3' in str(i)) and (i % 7 == 0 or '7' in str(i)):
        assert(king_kong(i) == 'KingKong')
### HIDEEND
```

<!-- grade_id: kingkong_test_1_4, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste Zahlen, für welche kein String zurückgegeben wird:
for i in range(1,100):
    if not((i % 3 == 0 or '3' in str(i)) or (i % 7 == 0 or '7' in str(i))):
        assert(king_kong(i) == i)
### HIDEEND
```

### b) 
Implementieren Sie eine Funktion ```king_kong_liste(n)```, welche für eine gegebene natürliche Zahl ```n``` eine Liste der ersten n Zahlen zurückgibt, für welche die Funktion ```king_kong(n)``` aus Teil a) aufgerufen wurde.

<!-- grade_id: kingkong_sol_2, points: 0, grade: false, solution: true -->
```python
def king_kong_liste(n):
    # Your code here
    ### BEGIN SOLUTION
    l = []
    for i in range(1,n+1):
        l.append(king_kong(i))
    return l
    ### END SOLUTION
```

<!-- grade_id: kingkong_test_2_0, points: 0, grade: true, solution: false -->
```python
assert(king_kong_liste(10) == [1,2,'King',4,5,'King','Kong',8,'King',10])
```

<!-- grade_id: kingkong_test_2_1, points: 2, grade: true, solution: false -->
```python
### HIDESTART
def king_kong_liste_sol(n):
    l = []
    for i in range(1,n+1):
        l.append(king_kong(i))
    return l

assert(king_kong_liste(100) == king_kong_liste_sol(100))
### HIDEEND
```

## Aufgabe 2: Binomialkoeffizienten


a) Schreiben Sie eine Funktion namens ```fakultaet```, welche ein ganzzahliges Parameter übergeben bekommen und die entsprechenden Fakultät zurückgibt. Implementieren Sie die Funktion mit Rekursion und nicht mit Schleifen.

<!-- grade_id: 21, points: 0, grade: false, solution: true -->
```python
def fakultaet(zahl):
    if (zahl <= 1):
        return 1
    return fakultaet(zahl-1)*zahl

print(fakultaet(10))
```

Output:
```
3628800
```

<!-- grade_id: 22, points: 1, grade: true, solution: false -->
```python
assert(fakultaet(0)==1)
assert(fakultaet(1)==1)
assert(fakultaet(10)==3628800)
```

b) Schreiben Sie eine Funktion namens ```binom```, welche zwei ganzzahlige Parametern ```n, k``` übergeben bekommt und die entsprechende binomische Koeffizient 

$$\binom{n}{k} = \frac{n!}{k!(n-k)!} $$

zurückgibt. 

<!-- grade_id: 24, points: 0, grade: false, solution: true -->
```python
def combination(n,k):
    return int(fakultaet(n)/(fakultaet(k)*fakultaet(n-k)))
```

<!-- grade_id: 25, points: 1, grade: true, solution: false -->
```python
assert(combination(0,0)==1)
assert(combination(1,0)==1)
assert(combination(1,1)==1)
assert(combination(10,1)==10)
assert(combination(10,5)==252)
```

c) Das Pascalsche Dreieck ist eine Form der grafischen Darstellung der Binomialkoeffizienten $\binom{n}{k}$, die eine einfache Berechnung dieser erlaubt. Dabei kann die Variable $n$ als Zeilenindex und $k$ als Spaltenindex interpretiert werden, wobei die Zählung mit Null beginnt. Hinweis: https://de.wikipedia.org/wiki/Pascalsches_Dreieck

Schreiben Sie eine Funktion namens ```pascal```, welche zwei ganzzahlige Parametern übergeben bekommt und die entsprechende Koeffizient vom Pascalschen Dreieck zurückgibt. Benutzen Sie dabei nicht die Funktion ```combination``` und implementieren Sie die Funktion mittels Rekursion.

<!-- grade_id: 27, points: 0, grade: false, solution: true -->
```python
def pascal(row, col):
    if (col == 0 or col == row):
        return 1
    else:
        return pascal(row-1,col-1) + pascal(row-1,col)
```

d) Schreiben Sie ein Programm, dass die 5 ersten Reihen vom pascalschen Dreieck in folgender Format ausgibt:
     1 
    1 1 
   1 2 1 
  1 3 3 1 
 1 4 6 4 1 

<!-- grade_id: 29, points: 1, grade: true, solution: true -->
```python
n=5
for i in range(0,n):
    print(" "*(n-i-1),end="")
    for j in range(0,i+1):
        print(pascal(i,j),end=" ")
    print("\n",end="")
```

Output:
```
1 
   1 1 
  1 2 1 
 1 3 3 1 
1 4 6 4 1
```

## Aufgabe 2: Kalendar

http://www.ntu.edu.sg/home/ehchua/programming/java/j2a_basicsexercises.html
DateUtil function (at the end of the page)



# Aufgabe 2: Folgen implementieren

### a)
Schreiben Sie eine Funktion `berechne_folge(n)`, welche für eine gegebene natürliche Zahl `n` eine Liste mit den einzelnen Elemenenten der Folge ausgibt. Die einzelnen Elemente berechnen sich wiefolgt:

1. Beginne mit irgendeiner natürlichen Zahl n > 0.
2. Ist n gerade, so nimm als nächstes n/2.
3. Ist n ungerade, so nimm als nächstes 3n + 1.
4. Wiederhole die Vorgehensweise mit der erhaltenen Zahl.

Der Algorithmus soll hier abbrechen, sobald eine 1 zu den Elementen der Folge hinzugefügt wurde.


<!-- grade_id: collatz_sol_1, points: 0, grade: false, solution: true -->
```python
def berechne_folge(n):
    # Your code here
    ### BEGIN SOLUTION
    L = []
    while(True):
        L.append(n)
        if n == 1:
            return L
        if n % 2 == 0:
            n = n//2
        else:
            n = 3*n+1
    ### END SOLUTION
```

<!-- grade_id: collatz_test_1, points: 2, grade: true, solution: false -->
```python
assert(berechne_folge(19) == [19, 58, 29, 88, 44, 22, 11, 34, 17, 52, 26, 13, 40, 20, 10, 5, 16, 8, 4, 2, 1])
assert(berechne_folge(15) == [15, 46, 23, 70, 35, 106, 53, 160, 80, 40, 20, 10, 5, 16, 8, 4, 2, 1])
```

###  b)
Begründen Sie, warum Sie in Ihrer Implementierung eine while- oder eine for-Schleife benutzt haben.

<!-- grade_id: collatz_sol_2, points: 1, grade: true, solution: true -->
- While-Schleife: Anzahl der Elemente der Folge ist nicht von vorneherein klar, daher bietet sich die While-Schleife an
- for-Schleife: Macht hier keinen Sinn

### c)
Warum ist es nicht einfach die Terminierung der Schleife zu beweisen?

Zusatzaufgabe (gibt keine Punkte): Versuchen Sie, Terminierung zu beweisen. Achtung, SCHWER! 

<!-- grade_id: collatz_sol_3, points: 1, grade: true, solution: true -->
Es ist nicht einfach eine geeignete Schleifeninvariante anzugeben.   
Die zu implementierende Folge ist die Collatz-Folge, welche (vermutlich) immer in dem Zyklus (4,2,1) mündet. Doch der Beweis dieser Vermutung ist bis zum heutigen Zeitpunkt noch nicht gelungen.
Für mehr Informationen siehe https://de.wikipedia.org/wiki/Collatz-Problem.

# Aufgabe 3
Schreiben Sie eine Funktion `matrix_mult(X,Y)`, welche zwei gegebene $n \times n$ Matrizen `X` und `Y` multipliziert und die Ergebnismatrix zurückgibt.
Die Darstellung der Eingabematrizen lassen sich an den Tests ablesen.

<!-- grade_id: matrixmult_sol, points: 0, grade: false, solution: true -->
```python
def matrix_mult(X,Y):
    # Your code here
    ### BEGIN SOLUTION
    C = [[0 for col in range(len(Y[0]))] for row in range(len(X))]
    for i in range(len(X)):
        for j in range(len(Y[0])):
            for k in range(len(Y)):
                C[i][j] += X[i][k]*Y[k][j]
    return C
    ### END SOLUTION
```

<!-- grade_id: matrixmult_test, points: 4, grade: true, solution: false -->
```python
X = [[12,7,3], [4 ,5,6], [7 ,8,9]]
Y = [[5,8,1], [6,7,3], [4,5,9]]
r = matrix_mult(X,Y)
assert(r == [[114, 160, 60], [74, 97, 73], [119, 157, 112]])

### HIDESTART
X = [[1,2,3], [4 ,5,3], [1 ,8,6]]
Y = [[5,4,1], [6,2,3], [4,1,3]]
r = matrix_mult(X,Y)
assert(r == [[29, 11, 16], [62, 29, 28], [77, 26, 43]])
### HIDEEND
```

# Aufgabe 5
Der Dozent einer Programmiervorlesung nutzt ein tolles neues Tool, um die Lösungen seiner Studierenden möglichst komfortabel bewerten zu können.
Leider erhält er von diesem Tool jedoch zwei getrennte Listen mit Matrikelnummern und Punkten, wie die folgenden:

```Python
matrikel = [12345, 12531, 18034, 17824]
punkte = [19, 13, 0, 20]
```

Dabei bezieht sich der erste Eintrag in der Liste "Punkte" auf die erste Matrikelnummer, usw.
Praktischer wäre es jedoch, eine Liste von Matrikelnummer-Punkte-Paaren zu haben, z. B. so:

```Python
matrikel_punkte = [(12345, 19), (12531, 13), (18034, 0), (17824, 20)]
```

Schreiben Sie eine Funktion namens ```verknuepfe(l1, l2)```, die zwei Listen als Parameter zurückgegeben bekommt und die Verknüpfung dieser beiden Listen als Liste zurückgibt (so wie ```matrikel_punkte``` eine Verknüpfung der Listen ```matrikel``` und ```punkte``` ist).

Hinweis: Verwenden Sie hierzu u. a. die Funktion zip(), https://docs.python.org/3/library/functions.html#zip

<!-- grade_id: zipa2, points: 0, grade: false, solution: true -->
```python
def verknuepfe(l1, l2):
    return list(zip(l1, l2))
```

<!-- grade_id: zipa3, points: 0, grade: true, solution: false -->
```python
matrikel = [12345, 12531, 18034, 17824]
punkte = [19, 13, 0, 20]
matrikel_punkte = [(12345, 19), (12531, 13), (18034, 0), (17824, 20)]

assert(matrikel_punkte == verknuepfe(matrikel,punkte))

#TODO: Hidden asserts
```

b) Leider stellt sich nach der Bewertung des dritten Übungsblattes heraus, dass das Bewertungs-Tool einen merkwürdigen Fehler hat:
Befindet sich am Ende der Liste ```matrikel``` eine Folge von Matrikelnummern, deren zugehörige Studierende alle null Punkte erreicht haben, so tauchen diese in der Liste ```punkte``` nicht auf.
Die liste ```punkte``` ist dann also kürzer als ```matrikel```.

Beispiel:
```Python
matrikel = [12345, 12531, 18034, 17824, 12635]
punkte = [6, 0, 3]
```

anstelle von:
```Python
punkte = [6, 0, 3, 0, 0]
```

Wie würde Ihre Funktion damit umgehen, dass die Längen der beiden Listen nicht übereinstimmen?

<!-- grade_id: zipb2, points: 0, grade: true, solution: true -->
Die betroffenen Matrikelnummern am Ende der Liste würden einfach nicht auftauchen.

Schreiben Sie eine Funktion namens ```verknuepfe_besser(l1, l2)```, welche Ihre Funktion aus a) so erweitert, dass die richtige (eigentlich erwartete) Liste zurückgegeben wird.

<!-- grade_id: zipb4, points: 0, grade: false, solution: true -->
```python
def verknuepfe_besser(l1, l2):
    return list(zip(l1, l2 + [0] * (len(l1)-len(l2))))
```

<!-- grade_id: zipb5, points: 0, grade: true, solution: false -->
```python
matrikel = [12345, 12531, 18034, 17824, 12635]
punkte = [6, 0, 3]
matrikel_punkte = [(12345, 6), (12531, 0), (18034, 3), (17824, 0), (12635, 0)]
assert(matrikel_punkte == verknuepfe_besser(matrikel,punkte))

#TODO: hidden asserts
```

c) Nach Veröffentlichung der Ergebnisse stellte sich heraus, dass einige Nachkorrekturen nötig waren. 
Der Dozent möchte die korrigierten Ergebnisse nun wieder in das Tool geben.
Leider muss man dem Tool dazu wieder zwei Listen von der Form wie ```matrikel``` und ```punkte``` aus Teilaufgabe a) übergeben.
Ihm liegen die korrigierten Ergebnisse jedoch nur in der Form wie ```matrikel_punkte``` vor.

Schreiben Sie eine Funktion namens ```entknuepfe(l)```, welche eine solche Liste übergeben bekommt und ein Paar (also 2-Tupel) der gewünschten Listen zurückgibt.

Verwenden Sie hierzu abermals die Funktion zip().

Hinweis: Sie müssen hier keine Sonderfälle (leere Liste o.ä.) beachten.

<!-- grade_id: zipc2, points: 0, grade: true, solution: true -->
```python
def entknuepfe(l):
    entknuepft_als_liste_von_tuples = list(zip(*l))
    l1 = list(entknuepft_als_liste_von_tuples[0])
    l2 = list(entknuepft_als_liste_von_tuples[1])
    
    return (l1, l2)
    #TODO: das riecht danach, als ginge es auch eleganter...
```

<!-- grade_id: zipc3, points: 0, grade: true, solution: false -->
```python
matrikel_punkte = [(12345, 19), (12531, 13), (18034, 0), (17824, 20)]
matrikel = [12345, 12531, 18034, 17824]
punkte = [19, 13, 0, 20]

assert((matrikel, punkte) == entknuepfe(matrikel_punkte))

#TODO: Hidden asserts
```

# Aufgabe 6

Neben der *Anweisung*
```Python
if BEDINGUNG:
	Block1
else:
	Block2
```

kennt Python folgenden *Ausdruck*:
```Python
a if BEDINGUNG else b
```

Sie können Details hier (URL zu Python doc) nachlesen.

a) Stellen Sie den Bezug zwischen Blcok 1, 2 und den a, b her - was entpsricht was?

b) Sie a, b Ausdrücke oder Blöcke?

c) Was können Sie also in a, b insbesondere NICHT tun?

d) Sie können in a oder b Funktionen mit Seiteneffekten aufrufen, zum Beispiel eine Funktion, die eine globale Variable verändert. Kann ein solcher Seiteneffekt das Result dieses if-else-Ausdrucks beeinflussen?

Beispiel:

x = None
def f():
	global x
	x = True
	return x

def g():
	global x
	x = False
	return x

# Ausdruck:
print( g() if x else f() )

(Antwort: natürlich nicht; die Bedingung wird ja zuerst ausgewertet, auch wenn die nach a steht )