<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->


Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in der *obige* Zelle ein. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

## Aufgabe 1: Boolesche Variablen und Vergleiche


a) (2 Punkte) Die Variablen $a$, $b$ und $c$ sind boolesche Variablen unbekannter Belegung. Verkürzen Sie folgende boolesche Ausdrücke so weit wie möglich und geben Sie die Zwischenschritte an: 
  
```Python
( (True | a ) and ( False | b) )  
( (a | b ) | ( False | b) )  
(a | b) and (a | c)
(a and (a | b)) and True
a and b | b and c and (b | c)
```




<!-- grade_id: 2, points: 2, grade: true, solution: true -->


## Aufgabe 2:  Mathematische Operationen

Mit Python kann man mathematische Ausdrücke auswerten lassen. Das wird durch die folgenden mathematischen Operatoren: `+, -, *, /,` Potenz (`**`) und Modulo (`%`) unterstützt. 

a) (2 Punkte) Berechnen sie das Produkt von `a` und `b` und geben sie das Ergebnis aus.
```Python 
a = 20
b = 66974
```

<!-- grade_id: 14, points: 2, grade: true, solution: true -->
```python

```

# Aufgabe 3
(6 Punkte) Schreiben Sie eine Funktion ```gesamtvermoegen(zinssatz, vermoegen, jahre)```, welches das Gesamtvermögen nach ```jahre``` Jahren bei einem Startkapital von ```vermoegen``` Euro und einem Zinssatz von ```zinssatz``` % berechnet. 
Das Gesamtvermögen kann man mit folgender Formel ausrechnen: 
$$K \left(1 + \frac{p}{100} \right)^n, $$ wobei $K$ das Startkapital, $n$ die Anzahl der Jahre und $p$ den Zinssatz in Prozent bezeichnen. 

<!-- grade_id: 15, points: 2, grade: false, solution: true -->
```python
def gesamtvermoegen (p, k, n):
    return k * (1 + p / 100)**n
```

<!-- grade_id: 15b, points: 1, grade: true, solution: false -->
```python
assert(round(gesamtvermoegen(5, 1000, 3), 3) == 1157.625)
```

Output:
```
1157.6250000000002
```

<!-- grade_id: 15c, points: 5, grade: true, solution: false -->
```python
### HIDESTART
for p,k,n in zip([1, 5,10,15],[500,1000,2000,100000],[1,2,5,100]):
    assert(round(gesamtvermoegen(p, k, n), 5) == round(k * (1 + p / 100)**n, 5))
### HIDEEND
```