<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

Format:

MatrikelNummer E-Mail-Adresse
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
234567 testStudent2@mail.upb.de
345678 vorname.nachname@mail.upb.de

Hinweis: Das sieht nicht nach separaten Zeilen aus, wenn Sie die Markdownzelle anzeigen lassen. Das ist egal; wichtig ist Formatierung der Eingabe.


## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile), für das entsprechende Übungsblatt! (Wenn Sie Blatt 9 abgeben wolllen hilft es nicht, bei Blatt 1 auf Submit zu drücken!) Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue09.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen, auch nicht Dateiname hue9.ipynb !). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 15.1.2017 um 23:59 Uhr

# Aufgabe 2: Schach
In dieser Aufgabe sollen Sie eine einfache Version eines Schachspiels implementieren.
Die Klassen `Feld` und `Figur` sind Ihnen dazu bereits vorgegeben.

## a)
Schreiben Sie eine Klasse `Schachbrett`, welche die folgenden Anforderungen erfüllt:
 - Im Konstruktor soll ein Spielfeld aus 64 Feldern (8x8, 8 Zeilen und 8 Spalten, siehe Bild, siehe Tests) anlegt werden.
 ![text](schachbrett.gif)

 - `get_feld(x,y)`: Liefert das Objekt der Klasse `Feld`, welches die übergebenen Koordinaten hat. Falls `x` oder `y` keine korrekten Koordinaten sind (d.h. `not x in range(1,9)` oder `not y in range(1,9)`), dann soll ein `ValueError` geworfen werden.
 - `setze_figur(figur, x, y)`: Setzt die übergebene Figur auf das Feld mit den übergebenen Koordinaten. Sie können davon ausgehen, dass das Feld nicht besetzt ist und die Figur vorher auf keinem Feld gestanden hat. Beachten Sie, dass die Beziehung zwischen `Feld` und `Figur` bidirektional sein soll (siehe Tests).

```python
class Feld:
    """Ein einzelnes Feld eines Schachbretts"""
    
    def __init__(self, koordinaten):
        """Konstruktor"""
        self.koordinaten = koordinaten
        self.figur = None
    
    def __eq__(self, feld):
        """Zum Vergleich mit anderem Feld"""
        return feld.koordinaten[0] == self.koordinaten[0] and feld.koordinaten[1] == self.koordinaten[1]
    
class Figur:
    """Eine Figur, welche auf einem Feld des Schachbrettes stehen kann."""
    
    def __init__(self):
        """Konstruktor"""
        self.feld = None
        
    def berechne_zuege(self, brett):
        """Liefert eine Liste mit Feldern, welche alle Felder enthält, auf die sich die Figur ausgehend von ihrer 
        aktuellen Position in einem Zug bewegen kann. Muss von erbenden Klassen überschrieben werden!"""
        pass
```

<!-- grade_id: schach_sol_1, points: 0, grade: false, solution: true -->
```python
class Schachbrett:
    ### BEGIN SOLUTION
    """Das Schachbrett, bestehend aus insgesamt 64 Feldern"""
    
    def __init__(self):
        """Konstruktor, in welchem ein 8x8-Schachbrett angelegt wird"""
        self.felder = []
        for i in range(1,9):
            zeile = []
            for j in range(1,9):
                zeile.append(Feld((i,j)))
            self.felder.append(zeile)
        
    def get_feld(self, x, y):
        """Liefert das Objekt der Klasse `Feld`, welches die übergebenen Koordinaten hat. Falls `x` oder `y` keine 
        korrekten Koordinaten sind (d.h. `not x in range(1,9)` oder `not y in range(1,9)`), dann soll ein `ValueError` 
        geworfen werden.
        
        Args:
            x: x-Koordinate des gesuchten Feldes (int)
            y: y-Koordinate des gesuchten Feldes (int)
            
        Returns:
            Feld Das gesuchte Feld
        """
        for zeile in self.felder:
            for feld in zeile:
                if feld.koordinaten[0] == x and feld.koordinaten[1] == y:
                    return feld
        raise ValueError()
        
    def setze_figur(self, figur, x, y):
        """ Setzt die übergebene Figur auf das Feld mit den übergebenen Koordinaten, indem es eine bidirektionale Beziehung
        zwischen der Figur und dem Feld herstellt.
        
        Args:
            figur: Die Figur, welche auf das Feld gesetzt werden soll (Figur)
            x: x-Koordinate des gesuchten Feldes (int)
            y: y-Koordinate des gesuchten Feldes (int)
        """
        try:
            feld = self.get_feld(x,y)
            feld.figur = figur
            figur.feld = feld
        except ValueError:
            pass
    ### END SOLUTION
```

<!-- grade_id: schach_test_1, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
s = Schachbrett()
for i in range(1,9):
    for j in range(1,9):
        feld = s.get_feld(i,j)
        assert(feld.koordinaten[0] == i and feld.koordinaten[1] == j)

# Tests falsche Koordinaten in get_feld:
try:
    s.get_feld(0,1)
    assert(False)
except ValueError:
    assert(True)
    
k = Figur()
s.setze_figur(k, 4, 5)
assert(s.get_feld(4, 5).figur == k and k.feld == s.get_feld(4, 5))

### HIDESTART
s = Schachbrett()
for i in range(1,9):
    for j in range(1,9):
        feld = s.get_feld(i,j)
        assert(feld.koordinaten[0] == i and feld.koordinaten[1] == j)

# Tests falsche Koordinaten in get_feld:
try:
    s.get_feld(1,9)
    assert(False)
except ValueError:
    assert(True)
    
k = Figur()
s.setze_figur(k, 2, 7)
assert(s.get_feld(2, 7).figur == k and k.feld == s.get_feld(2, 7))
### HIDEEND
```

## b)
Schreiben Sie eine Klasse `Koenig`, welche von der Klasse `Figur` erbt.
Diese Klassen sollen die Methode `berechne_zuege` der Klasse `Figur` überschreiben.
Die Methode `berechne_zuege` soll eine Liste zurückgeben, welche alle Felder enthält, auf die sich der König ausgehend von seiner aktuellen Position in einem Zug bewegen kann.
Gehen Sie davon aus, dass alle anderen 63 Felder nicht mit anderen Figuren belegt sind.
Der König kann sich im Schach immer nur um ein Feld weiter bewegen, allerdings in jede beliebige Richtung: Nach Oben, Unten Links, Rechts oder in die Diagonalen (wir ignorieren hier die Möglichkeit der Rochade!).   
Beispiel: Das nachfolgende Bild gibt die möglichen Züge des Königs an, wenn dieser sich auf dem Feld mit den Koordinaten (5,4) befindet.
 ![text](koenig.png)
Beachten Sie Randfälle.

<!-- grade_id: schach_sol_2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Koenig(Figur):
    """Repräsentiert die Schachfigur des Königs"""
    
    def berechne_zuege(self, brett):
        """Berechnet alle möglichen Züge des Königs"""
        l = []
        indices = (self.feld.koordinaten[0], self.feld.koordinaten[1])

        if indices[0] > 1:
            f = brett.get_feld(indices[0]-1, indices[1])
            l.append(f)
        if indices[1] > 1:
            f = brett.get_feld(indices[0], indices[1]-1)
            l.append(f)
        if indices[0] > 1 and indices[1] > 1:
            f = brett.get_feld(indices[0]-1, indices[1]-1)
            l.append(f)
        if indices[0] < 8:
            f = brett.get_feld(indices[0]+1, indices[1])
            l.append(f)
        if indices[1] < 8:
            f = brett.get_feld(indices[0], indices[1]+1)
            l.append(f)
        if indices[0] < 8 and indices[1] < 8:
            f = brett.get_feld(indices[0]+1, indices[1]+1)
            l.append(f)
        if indices[0] > 1 and indices[1] < 8:
            f = brett.get_feld(indices[0]-1, indices[1]+1)
            l.append(f)
        if indices[0] < 8 and indices[1] > 1:
            f = brett.get_feld(indices[0]+1, indices[1]-1)
            l.append(f)
        return l
### END SOLUTION
```

<!-- grade_id: schach_test_2, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
s = Schachbrett()
k = Koenig()
s.setze_figur(k, 5, 4)
felder = k.berechne_zuege(s)
assert(s.get_feld(5, 3) in felder)
assert(s.get_feld(5, 5) in felder)
assert(s.get_feld(6, 3) in felder)
assert(s.get_feld(6, 4) in felder)
assert(s.get_feld(6, 5) in felder)
assert(s.get_feld(4, 4) in felder)
assert(s.get_feld(4, 5) in felder)
assert(s.get_feld(4, 3) in felder)
assert(len(felder) == 8)

### HIDESTART
s = Schachbrett()
k = Koenig()
s.setze_figur(k, 1, 1)
felder = k.berechne_zuege(s)
assert(s.get_feld(2, 1) in felder)
assert(s.get_feld(1, 2) in felder)
assert(s.get_feld(2, 2) in felder)
assert(len(felder) == 3)
### HIDEEND
```

## c)
Analog zu Teilaufgabe b: Schreiben Sie nun eine Klasse `Turm`, die von `Figur` erbt und in `berechne_zuege` jeweils alle möglichen Felder, auf die sich der Turm ausgehend von seiner aktuellen Position in einem Zug bewegen kann, zurückgibt.
Gehen Sie davon aus, dass alle anderen 63 Felder nicht mit anderen Figuren belegt sind.
Ein Turm kann sich im Schach immer nur in waagerechte oder senkrechte Richtung bewegen, allerdings um beliebig viele Felder.   
Beispiel: Das nachfolgende Bild gibt die möglichen Züge des Turms an, wenn dieser sich auf dem Feld mit den Koordinaten (5,4) befindet.
![text](turm.png)

Achten Sie auch hier auf Randfälle.

<!-- grade_id: schach_sol_3, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Turm(Figur):
    """Repräsentiert die Schachfigur des Turms"""
    
    def berechne_zuege(self, brett):
        """Berechnet alle möglichen Züge des Turms"""
        l = []
        indices = (self.feld.koordinaten[0], self.feld.koordinaten[1])
        x = indices[0]
        y = indices[1]
        
        while x > 1:
            x = x - 1
            l.append(brett.get_feld(x,y))
        x = indices[0]
        
        while x < 8:
            x = x + 1
            l.append(brett.get_feld(x,y))
        x = indices[0]
        
        while y > 1:
            y = y - 1
            l.append(brett.get_feld(x,y))
        y = indices[1]
        
        while y < 8:
            y = y + 1
            l.append(brett.get_feld(x,y))
        
        return l
### END SOLUTION
```

<!-- grade_id: schach_test_3, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
s = Schachbrett()
t = Turm()
s.setze_figur(t, 4, 5)
felder = t.berechne_zuege(s)
assert(s.get_feld(1, 5) in felder)
assert(s.get_feld(2, 5) in felder)
assert(s.get_feld(3, 5) in felder)
assert(s.get_feld(5, 5) in felder)
assert(s.get_feld(6, 5) in felder)
assert(s.get_feld(7, 5) in felder)
assert(s.get_feld(8, 5) in felder)
assert(s.get_feld(4, 1) in felder)
assert(s.get_feld(4, 2) in felder)
assert(s.get_feld(4, 3) in felder)
assert(s.get_feld(4, 4) in felder)
assert(s.get_feld(4, 6) in felder)
assert(s.get_feld(4, 7) in felder)
assert(s.get_feld(4, 8) in felder)
assert(len(felder) == 14)

### HIDESTART
s = Schachbrett()
t = Turm()
s.setze_figur(t, 1, 1)
felder = t.berechne_zuege(s)
assert(s.get_feld(2, 1) in felder)
assert(s.get_feld(3, 1) in felder)
assert(s.get_feld(4, 1) in felder)
assert(s.get_feld(5, 1) in felder)
assert(s.get_feld(6, 1) in felder)
assert(s.get_feld(7, 1) in felder)
assert(s.get_feld(8, 1) in felder)
assert(s.get_feld(1, 2) in felder)
assert(s.get_feld(1, 3) in felder)
assert(s.get_feld(1, 4) in felder)
assert(s.get_feld(1, 5) in felder)
assert(s.get_feld(1, 6) in felder)
assert(s.get_feld(1, 7) in felder)
assert(s.get_feld(1, 8) in felder)
assert(len(felder) == 14)
### HIDEEND
```

<!-- grade_id: docu_test, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Alle Dokumentationen vorhanden? 

# Aufgabe 2:
assert(Schachbrett.__doc__ != None)
assert(Schachbrett.__init__.__doc__ != None)
assert(Schachbrett.get_feld.__doc__ != None)
assert(Schachbrett.setze_figur.__doc__ != None)

assert(Koenig.__doc__ != None)
assert(Koenig.berechne_zuege.__doc__ != None)
```