<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->


Tragt bitte eure Matrikel Nummern und Email Adressen *zeilenweise* in das Feld über dieser Zelle ein. Achtet dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. Benutzen Sie bitte keine Leerzeilen, einziges Leerzeichen ist Zwischen Matrikelnummer und Mailadresse.

Format:

MatrikelNr. Emailaddresse

z.B.:

123456 testStudent@mail.upb.de 

## Aufgabe 1: Boolsche Variablen und Vergleichungen


a) (2 Punkte) Die Variablen $a$, $b$ und $c$ sind boolesche Variablen unbekannter Belegung. Verkürzen Sie folgende boolesche Ausdrücke so weit wie möglich und geben Sie die zwishenschritte an: 
  
```Python
( (True | a ) and ( False | b) )  
( (a | b ) | ( False | b) )  
(a | b) and (a | c)
(a and (a | b)) and True
a and b | b and c and (b | c)
```




<!-- grade_id: a, points: 2, grade: true, solution: true -->


b) (2 Punkte) Nehmen Sie an, dass die Variablen $a$, $b$ und $c$ die folgende Werte haben: 
    
```Python
a, b, c = True, True, False
```
   
Welche Werten werden die Variablen $d$,$e$,$f$ und $g$ haben? Versuchen Sie die Aufgabe ohne Hilfe von Python zu lösen.

```Python
d = b and c
e = b or c
f = not a and c
g = (b and c) or not a
```

<!-- grade_id: c, points: 2, grade: true, solution: true -->



c) (2 Punkte) Nehmen Sie die folgende Belegung an: 

```Python
a, b, c, d, foo = 2, -1, "2", 0, "foo"
```

Nennen Sie welche der folgenden Bedingungen als `True` ausgewertet wird 

```Python
a < b
(a < b) and d
(a > b) or d
b <= a
b >= b
b > d
a != b
d
b == -1
a == c
c == "2"
foo == "foo"
foo = "foo"
```

<!-- grade_id: e, points: 2, grade: true, solution: true -->


## Aufgabe 2: Typen von Variablen

a) (2 Punkte) Welche der folgenden Namen sind in Python legal? Begründen Sie Ihre Antwort: 
```Python
or
_or
var
var2
2var
my_var
dein-var
```


<!-- grade_id: f, points: 2, grade: true, solution: true -->


b) (2 Punkte) Benennen Sie die Typen der folgenden Variablen:
```Python
a = 4
b = 4.0
c = '4.0'
d = False
e = 'True'
f = 'Test'
g = int
```

<!-- grade_id: h, points: 2, grade: true, solution: true -->


## Aufgabe 3: Programmieren mit Pseudocode

a) (2 Punkt) Wie würden Sie ein Programm in Pseudocode aufschreiben, das die folgende Ungleichung:

$$1 \leq 5 \cdot x \leq 100 $$

für ganzzahlige $x \in \mathbb{Z}$ löst?  Hinweis:  https://de.wikipedia.org/wiki/Pseudocode

<!-- grade_id: i, points: 1, grade: true, solution: true -->
```python
print('Die Vielfacher von 5 zwischen 1 und 100 sind:')
for x in range(1,101):
    if((x%5)==0):
        print(x/5)
```

b) (2 Punkte) Sie möchten Geld von einem Geldautomat abheben, welcher jede Art von *Euroscheinen* ausgibt. Euromünzen sind ausgeschlossen. Schreiben Sie einen Pseudocode, welcher zu Ihrem eingegebenen Betrag die Mindestanzahl der Scheinen nennt. 

<!-- grade_id: j, points: 2, grade: true, solution: true -->
```python
nb = input('Geben Sie den gewünschten Betrag ein: ')
betrag = int(nb)

euroScheine = [500, 200, 100, 50, 20, 10, 5, 1]
restBetrag = betrag

euroScheine[2]

for i in range(1,len(euroScheine)):
    if ((restBetrag / euroScheine[i])>0):
        faktor = restBetrag // euroScheine[i]
        restBetrag = restBetrag - faktor*euroScheine[i]
        if (faktor > 0):
            print(faktor, 'x', euroScheine[i])
    

```

## Aufgabe 4: Eingabe und Ausgabe von Daten

a) (1 Punkt) Lesen Sie mit **input** ein Wort, ein Zeichen und eine Fließkommazahl ein und geben
Sie die Inhalte mit **print** wieder aus.

<!-- grade_id: k, points: 1, grade: true, solution: true -->
```python
nb = input('Geben Sie ein Wort ein: ')
word = str(nb)
nb = input('Geben Sie ein Zeichen ein: ')
strVal = str(nb)
nb = input('Geben Sie ein Fließkommazahl ein: ')
floatNr = float(nb)

print('Das eingegebene Wort ist: ', word)
print('Das eingegebene Zeichen ist: ', strVal)
print('Die Fließkommazahl ist: ', floatNr)
print('Type ist: ', type(nb))
```

b) (2 Punkte) Erweitern Sie das Programm, indem es den Namen des Benutzers einliest und ihn dann mit seinem
Namen begrüßt. 

<!-- grade_id: m, points: 2, grade: true, solution: true -->
```python
nb = input('Geben Sie Ihren Namen ein: ')
word = str(nb)
print('Hallo ', nb)
```

c) (2 Punkte) Geben Sie mithilfe des Modulo-Operators **%** die letzte Ziffer einer vom Benutzer eingegebenen
Ganzzahl aus.

<!-- grade_id: o, points: 2, grade: true, solution: true -->
```python
nb = input('Geben Sie eine ganzzahlige Ziffer ein: ')
ziffer = int(nb)
letzteZiffer = ziffer % 10
print('Letzte Ziffer ist: ', letzteZiffer)
```

## Aufgabe 5: Datenmanipulation

Mit Python kann man mathematische Ausdrücke auswerten lassen. Das wird durch die folgenden mathematischen Operatoren: `+, -, *, /,` Potenz (`**`) und Modulo (`%`) unterstützt.  

a) (2 Punkte) Lesen Sie zwei ganze Zahlen ein und geben Sie deren Summe und Produkt aus. 



<!-- grade_id: p, points: 4, grade: false, solution: true -->
```python
def calc_summ(zahl1,zahl2):
    """ 1. Lesen Sie zwei ganze Zahlen ein, 
        2. geben Sie deren Summe zu den Variabel summ ein 
        3. drücken Sie das Ergebnis aus.
    """
    ### BEGIN SOLUTION
    nb = input('Geben Sie die erste Zahl ein: ')
    zahl1 = int(nb)
    nb = input('Geben Sie die zweite Zahl ein: ')
    zahl2 = int(nb)

    summ = zahl1 + zahl2

    print('Summe ist: ', summ)
    ### END SOLUTION  
    
    return summ
```

<!-- grade_id: pp, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(calc_summ(1, 2), 3)
assert_equal(calc_summ(4, 5), 9)
```

```python
def calc_prod(zahl1,zahl2):
    """ 1. Lesen Sie zwei ganze Zahlen ein, 
        2. geben Sie deren Produkt zu den Variabel prod ein 
        3. drücken Sie das Ergebnis aus.
    """
    ### BEGIN SOLUTION
    nb = input('Geben Sie die erste Zahl ein: ')
    zahl1 = int(nb)
    nb = input('Geben Sie die zweite Zahl ein: ')
    zahl2 = int(nb)

    prod = zahl1 * zahl2

    print('Produkt ist: ', prod)
    ### END SOLUTION  
    
    return prod
```

<!-- grade_id: ppp, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(calc_prod(1, 2), 2)
assert_equal(calc_prod(4, 4), 16)
```

b) (4 Punkte) Lösen Sie die quadratische Gleichung $ax^2 + bx + c = 0$ für vom Benutzer 
eingegebene Parameter $a$, $b$ und $c$. Dazu können Sie die Lösungsformel für die quadratische 
Gleichung benutzen: $$x_{1,2}=\frac{-b \pm \sqrt{b^2 - 4ac}}{2a}.$$ 

<!-- grade_id: r, points: 0, grade: false, solution: true -->
```python
def calc_roots(a,b,c):
    """ Lesen Sie die Parametern von einer quadratische Gleichung ein.
        Übergeben sie die Lösungen zu die Variablen x_1 und x_2.
    """
    ### BEGIN SOLUTION
        " Schreiben Sie die vom Benutzer eingegebene Parametern auf die Variablen a,b,c ein.
        "
        nb = input('Geben Sie den ersten Parameter a ein: ')
        a = float(nb)
        nb = input('Geben Sie den zweiten Parameter b ein: ')
        b = float(nb)
        nb = input('Geben Sie den dritten Parameter c ein: ')
        c = float(nb)
        
        " Berechnen Sie die Lösungen und geben Sie die Ergebnissen auf die Variablen x_1 und x_2 ein
        "
        D= (b**2 - 4*a*c)**0.5;
        x_1 = (-b + D)/2*a;
        x_2 = (-b - D)/2*a;
        
        " Drücken Sie die Ergebnissen aus:
        "
        print('Die erste Wurzel ist: ', x_1);
        print('Die zweite Wurzel ist: ', x_2);
    ### END SOLUTION  
    
    return (x_1,x_2)
```

<!-- grade_id: rr, points: 4, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(calc_roots(1, -3, 2), [1, 2])
assert_equal(calc_roots(1, -7, 12), [3, 4])
```

c) (2 Punkte) Geben Sie mithilfe des Modulo-Operators % die letzte Ziffer einer vom Benutzer eingegebenen Ganzzahl aus.

<!-- grade_id: ss, points: 0, grade: false, solution: true -->
```python
def print_last(ziffer):
    """ Geben Sie ein ganzzahlige Ziffer ein.
        Drücken Sie die letzte Ziffer aus und übergeben Sie diesen Wert zu den Variabel last ein.
    """
    ### BEGIN SOLUTION
    nb = input('Geben Sie ein ganzzahlige Ziffer ein: ')
    ziffer = int(nb)
    last = ziffer%10
    print('Letzte Ziffer ist: ', last)
    ### END SOLUTION  
    
    return (last)
```

<!-- grade_id: sss, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(print_last(5421002), 2)
assert_equal(print_last(54210), 0)
```