<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

Format:

MatrikelNummer E-Mail-Adresse
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
234567 testStudent2@mail.upb.de
345678 vorname.nachname@mail.upb.de

Hinweis: Das sieht nicht nach separaten Zeilen aus, wenn Sie die Markdownzelle anzeigen lassen. Das ist egal; wichtig ist Formatierung der Eingabe.


## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile), für das entsprechende Übungsblatt! (Wenn Sie Blatt 8 abgeben wolllen hilft es nicht, bei Blatt 1 auf Submit zu drücken!) Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue08.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen, auch nicht Dateiname hue8.ipynb !). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 25.12.2016 um 23:59 Uhr

# Aufgabe 3: Arztpraxis
Gegeben ist die Klasse `Patient` (in komplexerer Form bereits bekannt aus Hausübungsblatt 7).   
Stellen Sie sich vor, Sie arbeiten in einer Arztpraxis und sollen ein System zur Verwaltung der Patienten im Wartezimmer implementieren.
Das System soll Patienten, die in die Praxis kommen, ins Wartezimmer verweisen.
Im Büro des Arztes darf sich maximal ein Patient gleichzeitig befinden (der Patient, der gerade vom Arzt behandelt wird).
Sobald das Büro des Arztes wieder frei wird (d.h. die Behandlung abgeschlossen ist), kann nun ein Patient im Wartezimmer aufgerufen und in das Büro des Arztes gebeten werden.
Für jeden Aufruf wird immer der Patient gewählt, der die längste Zeit im Wartezimmer verbracht hat.

```python
class Patient:
    """Modelliert einen Patienten in einfacher Form"""
    
    def __init__(self, name, alter):
        """Konstruktor
        
        Args:
            name: Name des Patienten (str)
            alter: Alter des patienten (int)
        """
        self.name = name
        self.alter = alter
        
    def get_name(self):
        """Gibt den Namen des Patienten zurück
        
        Returns:
            str Name des Patienten
        """
        return self.name
    
    def get_alter(s):
        """Gibt das Alter des Patienten zurück
        
        Returns:
            str Alter des Patienten
        """
        return s.alter
    
    def __str__(self):
        """Packt die Informationen zu diesem Patient in einen String
        
        
        Returns:
            str Patienteninfos
        """
        return "Name: " + self.name + ", Alter: " + str(self.alter)
```

### a)
Implementieren Sie eine Klasse `Arztpraxis`, welche die folgenden Methoden bereitstellt:
- `betretePraxis(p)`: Der gegebene Patient `p` betritt die Praxis und wird ins Wartezimmer gebeten.
- `naechsterPatientZumArzt()`: Entfernt den Patienten aus dem Wartezimmer, welcher die längste Zeit dort verbracht hat, schickt diesen zum Arzt und gibt diesen zurück. Falls das Wartezimmer leer ist oder sich noch ein Patient in Behandlung befindet, werden keine Aktionen ausgeführt und lediglich `None` zurückgegeben.
- `beimArzt()`: Gibt den Patienten zurück, welcher sich gerade beim Arzt befindet. Falls sich kein Patient gerade beim Arzt befindet, wird `None` zurückgegeben.
- `behandlungAbgeschlossen()`: Der Patient, der sich aktuell in Behandlung befindet verlässt die Praxis, sodass das Zimmer des Arztes nun wieder leer ist.

<!-- grade_id: praxis_1_sol, points: 0, grade: false, solution: true -->
```python
class Arztpraxis:
    ### BEGIN SOLUTION
    def __init__(self):
        self.warteschlange = []
        self.inBehandlung = None
        
    def betretePraxis(self, p):
        self.warteschlange.append(p)
    
    def naechsterPatientZumArzt(self):
        # Büro des Arztes noch besetzt:
        if not self.beimArzt() == None:
            return None
        
        # Kein Patient mehr im Wartezimmer:
        if len(self.warteschlange) == 0:
            return None
        
        p = self.warteschlange.pop(0)
        self.inBehandlung = p
        return p
    
    def beimArzt(self):
        return self.inBehandlung
    
    def behandlungAbgeschlossen(self):
        self.inBehandlung = None
    ### END SOLUTION    
```

<!-- grade_id: praxis_1_test, points: 3, grade: true, solution: false -->
```python
praxis = Arztpraxis()
patienten = [Patient("Max M", 26), Patient("Ulrike Eustermann", 13), Patient("Michaela M", 20)]

# Patienten betreten die Praxis:
praxis.betretePraxis(patienten[0])
praxis.betretePraxis(patienten[1])
praxis.betretePraxis(patienten[2])
assert(praxis.beimArzt() == None)

# Rufe Patienten auf:
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == patienten[0])

praxis.behandlungAbgeschlossen()
assert(praxis.beimArzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == patienten[1])

### HIDESTART
praxis = Arztpraxis()
patienten = [Patient("Ralf B", 26), Patient("Jessica W", 13), Patient("Philipp K", 20)]

# Teste Randfall(1):
assert(praxis.naechsterPatientZumArzt() == None)

# Patienten betreten die Praxis:
praxis.betretePraxis(patienten[0])
praxis.betretePraxis(patienten[1])
praxis.betretePraxis(patienten[2])
assert(praxis.beimArzt() == None)

# Rufe Patienten auf:
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == patienten[0])

# Teste Randfall(2):
assert(praxis.naechsterPatientZumArzt() == None)

praxis.behandlungAbgeschlossen()
assert(praxis.beimArzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == patienten[1])
### HIDEEND
```

Unsere Arztpraxis soll nun zwischen Kassen- und Privatpatienten unterscheiden.
Für Privatpatienten soll zusätzlich zu deren Namen und Alter auch noch die Anschrift (als String) gepspeichert werden.
Privatpatienten werden von der Praxis bevorzugt behandelt, d.h. für die Auswahl des als nächstes zu behandelden Patienten wird immer erst geprüft, ob sich ein Privatpatient im Wartezimmer befindet.
Ist dies der Fall so wird dieser den Kassenpatienten vorgezogen, unabhängig von der Zeit, die er im Wartezimmer verbracht hat.
Falls sich mehrere Privatpatienten im Wartezimmer befinden, so wird der sich am längsten im Wartezimmer befindliche Privatpatient als nächstes aufgerufen.   

## b) 
Erweitern Sie zunächst die Klasse Patient zu einer Klasse `Privatpatient`, welche über ein Attribut `anschrift` sowie eine Methode `get_anschrift()` verfügt.
Überschreiben Sie auch die Methode `__str__()`, sodass die Tests korrekt laufen.

<!-- grade_id: praxis_2_sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Privatpatient(Patient):
    
    def __init__(self, anschrift, *args):
        super().__init__(*args)
        self.anschrift = anschrift
        
    def get_anschrift(self):
        return self.anschrift
    
    def __str__(self):
        return "<Privatpatient> " + super().__str__() + " Anschrift: " + self.anschrift
### END SOLUTION
```

<!-- grade_id: praxis_2_test, points: 1, grade: true, solution: false -->
```python
p = Privatpatient("Flotowstr. 52", "Daniela L", 28)
assert(isinstance(p, Patient))
assert(isinstance(p, Privatpatient))
assert(p.__str__() == "<Privatpatient> Name: Daniela L, Alter: 28 Anschrift: Flotowstr. 52")
### HIDESTART
p2 = Privatpatient("Konstanzer Strasse 73", "Uwe  P", 25)
assert(isinstance(p2, Patient))
assert(isinstance(p2, Privatpatient))
assert(p2.__str__() == "<Privatpatient> Name: Uwe  P, Alter: 25 Anschrift: Konstanzer Strasse 73")
### HIDEEND
```

## c)
Erweitern Sie nun Ihre Klasse `Arztpraxis` aus Aufgabenteil a) zu einer Klasse `Arztpraxisunfair`, sodass diese nun Privatpatienten wie oben beschrieben bevorzugt behandelt.

<!-- grade_id: praxis_3_sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Arztpraxisunfair(Arztpraxis):

    def __init__(self):
        super().__init__()
        self.privat_schlange = []
        
    def betretePraxis(self, p):
        if isinstance(p, Privatpatient):
            self.privat_schlange.append(p)
        else:
            super().betretePraxis(p)
    
    def naechsterPatientZumArzt(self):
        # Büro des Arztes noch besetzt:
        if not self.beimArzt() == None:
            return None
        
        # Kein Patient mehr im Wartezimmer:
        if len(self.privat_schlange) and len(self.warteschlange) == 0:
            return None
        
        # Ziehe Privatpatienten vor:
        if len(self.privat_schlange) > 0:
            p = self.privat_schlange.pop(0)
            self.inBehandlung = p
            return p
        else:
            super().naechsterPatientZumArzt()
### END SOLUTION
```

<!-- grade_id: praxis_4_test, points: 2, grade: true, solution: false -->
```python
praxis = Arztpraxisunfair()
patienten = [Patient("Max M", 26), Patient("Ulrike Eustermann", 13), Patient("Michaela M", 20), 
             Privatpatient("Flotowstr. 52", "Daniela L", 28), Privatpatient("Brandenburgische Str. 51", "Markus S", 56)]

# Patienten betreten die Praxis:
praxis.betretePraxis(patienten[0])
praxis.betretePraxis(patienten[1])
praxis.betretePraxis(patienten[2])
praxis.betretePraxis(patienten[3])
praxis.betretePraxis(patienten[4])
assert(praxis.beimArzt() == None)

# Rufe Patienten auf:
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == patienten[3])

praxis.behandlungAbgeschlossen()
assert(praxis.beimArzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == patienten[4])

praxis.behandlungAbgeschlossen()
assert(praxis.beimArzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == patienten[0])

### HIDESTART
praxis = Arztpraxisunfair()
patienten = [Patient("Ralf B", 26), Privatpatient("Brandenburgische Str. 51", "Markus S", 56)]

# Teste Randfall(1):
assert(praxis.naechsterPatientZumArzt() == None)

# Patienten betreten die Praxis:
praxis.betretePraxis(patienten[0])
praxis.betretePraxis(patienten[1])
assert(praxis.beimArzt() == None)

# Rufe Patienten auf:
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == patienten[1])

# Teste Randfall(2):
assert(praxis.naechsterPatientZumArzt() == None)

praxis.behandlungAbgeschlossen()
assert(praxis.beimArzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == patienten[0])
### HIDEEND
```

<!-- grade_id: docu_test, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# TODO!!
### HIDEEND
```