<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

Format:

MatrikelNummer E-Mail-Adresse
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
234567 testStudent2@mail.upb.de
345678 vorname.nachname@mail.upb.de

Hinweis: Das sieht nicht nach separaten Zeilen aus, wenn Sie die Markdownzelle anzeigen lassen. Das ist egal; wichtig ist Formatierung der Eingabe.


## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile), für das entsprechende Übungsblatt! (Wenn Sie Blatt 8 abgeben wolllen hilft es nicht, bei Blatt 1 auf Submit zu drücken!) Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue08.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen, auch nicht Dateiname hue8.ipynb !). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 25.12.2016 um 23:59 Uhr

# Aufgabe 1: Vererbung
Für diese Aufgabe betrachten Sie den Code in der nachfolgenden Code-Zelle.

```python
class A:
    
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def to_string(self):
        return "x, y = ({},{})".format(self.x, self.y)
    
    def __str__(self):
        return self.to_string()
    
class B(A):
    
    def __init__(self, z, *args):
        A.__init__(self, *args)
        self.z = z
        
    def __str__(self):
        return super().__str__()
    
class C(A):
    
    def __init__(self, *args):
        super().__init__(*args)
        
    def to_string(self):
        return "Klasse C, {}".format(super().to_string())
```

## a)
Nennen Sie die beiden Möglichkeiten, den Konstruktor einer Oberklasse aufzurufen und argumentieren Sie anhand von *DRY*, welche der beiden Möglichkeiten sinnvoller erscheint.

<!-- grade_id: vererbung_1_sol, points: 2, grade: true, solution: true -->
Der Konstruktor einer Oberklasse (hier die Klasse A) lässt sich entweder via 
```Python
A.__init__()
```
oder via 
```Python
super().__init__()
```
aufrufen.
Bei der ersten Variante ist allerdings unschön, dass der Klassenname `A` dadurch mehrfach aufgeschrieben wird und somit DRY (Don't repeat yourself) verletzt wird.
Weiterhin ist der erste Aufruf fehleranfällig, falls man die Klasse `A` umbenennt oder die Oberklasse in der Unterklasse ändern möchte.   
Die Methode via `super()` verwendet ist da allgemeiner und eignet sich nicht nur für Aufrufe von Konstruktoren der Oberklasse.

## b)
Warum kann man im Konstruktor der Klasse `B` den Aufruf
```Python
    A.__init__(self, *args)
```
nicht zu 
```Python
    A.__init__(*args)
```
vereinfachen? In der Klasse `C` kann man schließlich auch den Konstruktor von `A` ohne die Angabe des Parameters `self` fehlerfrei durchführen.

<!-- grade_id: vererbung_2_sol, points: 1, grade: true, solution: true -->
Ein Aufruf des Oberklassenkonstruktors via `Oberklasse.__init__()` benötigt den Parameter `self`, da sonst nicht klar ist für welches Objekt der Konstruktor aufgerufen werden soll.
Durch solch einen Aufruf wird die Methode `__init__` nämlich als eine Klassenmethode gesehen und nicht als eine Objektspezifische Methode.

## c)
Verwenden Sie in der nachfolgenden Code-Zelle die Build-In Funktion `isinstance`, um zu überprüfen, ob die Objekte, welche von den Variablen `a`, `b` und `c` referenziert werden Instanzen von den obigen Klassen `A`, `B` und `C` sind.
Geben Sie das Ergebnis (`True` oder `False`) jedes einzelnen Aufrufs mittles `print` aus.

<!-- grade_id: vererbung_3_sol, points: 1, grade: true, solution: true -->
```python
a = A(1,2)
b = B(1,2,3)
c = C(4,5)

### BEGIN SOLUTION
print(isinstance(a,A))
print(isinstance(a,B))
print(isinstance(a,C))
print(isinstance(b,A))
print(isinstance(b,B))
print(isinstance(b,C))
print(isinstance(c,A))
print(isinstance(c,B))
print(isinstance(c,C))
### END SOLUTION
```

## d)
Geben Sie für die beiden folgenden Befehle an, welche Methoden in welcher Reihenfolge aufgerufen werden.

1. `print(b)`
2. `print(c)`

Beispielsweise soll dies für die Anweisung `b = B(1,2,3)` in dem folgenden Format geschehen:

`b = B(1,2,3)` $\rightarrow$ `B.__init__()` $\rightarrow$ `A.__init__()`


<!-- grade_id: vererbung_sol_4, points: 2, grade: true, solution: true -->
1. `print(b)` $\rightarrow$ `B.__str__()` $\rightarrow$ `A.__str__()` $\rightarrow$ `A.to_string()`   
2. `print(c)` $\rightarrow$ `A.__str__()` $\rightarrow$ `C.to_string()` $\rightarrow$ `A.to_string()`

# Aufgabe 2: Löschen im Binärbaum
In dieser Aufgabe sollen Sie die Operation `delete()` der Klasse `BinTree` aus der Vorlesung implementieren.

## a)
Beantworten Sie zunächst die folgenden Fragen:
1. Nehmen Sie an, Sie kennen den Knoten, welcher den zu löschenden Schlüssel enthält bereits. Zwischen welchen drei Fällen müssen Sie beim Löschen unterscheiden?
2. Welche Zustandsinformationen müssen Sie sich in den rekursiven Aufrufen merken?

<!-- grade_id: bintree_sol_1, points: 2, grade: true, solution: true -->
1. - Der zu löschende Knoten ist ein Blattknoten
   - Der zu löschende Knoten hat genau einen Kindknoten
   - Der zu löschende Knoten hat zwei Kindknoten
2. Das kann je nach Art der Implementierung variieren: Was benötigt wird sind Referenzen zu dem zu löschenden Knoten und dessen Vaterknoten. Der Vaterknoten kann aber auch in der Klasse TreeNode eingerichtet werden. Dadurch würden die Methoden `add` und `delete` aber etwas komplizierter. In der folgenden Lösung wird der Vaterknoten daher immer in jedem rekursiven Aufruf übergeben.

## b)
Vervollständigen Sie nun die Methode `delete()` in der unten angegebenen Klasse `BinTree`, sodass die angegebenen Tests funktionieren.
Sie dürfen dabei annehmen, dass der zu löschende Schlüssel im Baum vorhanden ist.   

*Tipp*: Überlegen Sie sich eine geschickte Variante um Schlüssel, deren Knoten zwei Kindknoten haben, ohne unnötiges verschieben von Referenzen zu löschen.

<!-- grade_id: bintree_sol_2, points: 0, grade: false, solution: true -->
```python
class TreeNode:
    def __init__(self, key, val):
        self.key = key
        self.value = val
        self.leftChild = None
        self.rightChild = None

    def isLeaf(self):
        return ((self.leftChild is None) and (self.rightChild is None))

class BinaryTree:
    
    def __init__(self):
        self.root = None
        self.size = 0
    
    def add(self, key, value):
        """Füge einen Eintrag unter key an die richtige Stelle ein"""
        if self.root is None:
            # Baum ist leer, wir erzeugen den ersten Eintrag
            self.root = TreeNode(key, value)
        else:
            # Versuche, den Eintrag unter dem Knoten self.root einzufuegen das wird laenger; also eigene Methode
            self._add(key, value, self.root)

        self.size += 1

    def _add(self, key, value, currentNode):
        if key < currentNode.key:
            # fuege im linken Teil ein
            if currentNode.leftChild is None:
                # hier muss der neue Knoten hin!
                currentNode.leftChild = TreeNode(key, value)
            else:
                # es gibt schon Kinder, suche dort weiter:
                self._add(key, value, currentNode.leftChild)
        elif key > currentNode.key:
            # der Fall is ganz symmetrisch
            if currentNode.rightChild is None:
                currentNode.rightChild = TreeNode(key, value)
            else:
                self._add(key, value, currentNode.rightChild)
        else:
            # wenn key schon existiert: aktualisiere Wert
            currentNode.value = value
    
    def iterate(self):
        """Traversiere den Baum 'preorder'."""
        return self._iterate(self.root, [])
    
    def _iterate(self, t, s):
        s.append(t.key)
        if t.leftChild != None:
            self._iterate(t.leftChild, s)
        if t.rightChild != None:
            self._iterate(t.rightChild, s)
        return s
        
    ### BEGIN SOLUTION
    def remove(self, key):
        """Entferne Eintrag unter key falls vorhanden; erhalte Baumstruktur!"""
        if self.root.key == key and self.size == 1:
            # wir entfernen den letzten Knoten:
            self.root = None
        else:
            self._remove(key, self.root, None)

        self.size -= 1

    def _remove(self, key, currentNode, parent):            
        if key == currentNode.key:
            # Zu löschender Knoten gefunden -- hat er Kinder?
            if currentNode.isLeaf():
                # Zu löschender Knoten hat keine Kinder - Lösche Knotenreferenz vom Parent-Node:
                if parent.leftChild != None and parent.leftChild.key == key:
                    parent.leftChild = None
                else:
                    parent.rightChild = None
            elif currentNode.leftChild == None and currentNode.rightChild != None:
                # Zu löschender Knoten hat einen Kindknoten (rechts):
                if currentNode == self.root:
                    self.root = currentNode.rightChild
                else:
                    if parent.leftChild != None and parent.leftChild.key == key:
                        parent.leftChild = currentNode.rightChild
                    else:
                        parent.rightChild = currentNode.rightChild
            elif currentNode.leftChild != None and currentNode.rightChild == None: # analog zum vorherigen Fall!
                # Zu löschender Knoten hat einen Kindknoten (links):
                if currentNode == self.root:
                    self.root = currentNode.leftChild
                else:
                    if parent.leftChild != None and parent.leftChild.key == key:
                        parent.leftChild = currentNode.leftChild
                    else:
                        parent.rightChild = currentNode.leftChild
            else:
                # Zu löschender Knoten hat zwei Kindknoten - Suche nach in-order-vorgänger:
                vorgaenger = currentNode.leftChild
                vorgaengerparent = currentNode
                while vorgaenger.rightChild != None:
                    vorgaengerparent = vorgaenger
                    vorgaenger = vorgaenger.rightChild
                    
                # Tausche key-value Paare des zu löschenden Knotens mit dem in-order-vorgänger:
                k,v = currentNode.key, currentNode.value
                currentNode.key, currentNode.value = vorgaenger.key, vorgaenger.value
                vorgaenger.key, vorgaenger.value = k,v
                
                # Nun braucht nurnoch der vorgaengerknoten gelöscht werden, dieser hat weniger als 2 Kindknoten!
                self._remove(key, vorgaenger, vorgaengerparent)
        elif key < currentNode.key:
            self._remove(key, currentNode.leftChild, currentNode)
        else:
            self._remove(key, currentNode.rightChild, currentNode)
    ### END SOLUTION
```

<!-- grade_id: bintree_test_2, points: 5, grade: true, solution: false -->
```python
t = BinaryTree()
t.add(2, "value")
t.add(7, "value")
t.add(5, "value")
t.add(1, "value")
t.add(6, "value")
t.add(9, "value")
t.add(3, "value")
t.add(11,"value")
t.add(4, "value")

assert(t.iterate() == [2, 1, 7, 5, 3, 4, 6, 9, 11])
t.remove(1)
assert(t.iterate() == [2, 7, 5, 3, 4, 6, 9, 11])
t.remove(3)
assert(t.iterate() == [2, 7, 5, 4, 6, 9, 11])
t.remove(7)
assert(t.iterate() == [2, 6, 5, 4, 9, 11])
t.remove(2)
assert(t.iterate() == [6, 5, 4, 9, 11])
```

# Aufgabe 3: Arztpraxis
Gegeben ist die Klasse `Patient` (in komplexerer Form bereits bekannt aus Hausübungsblatt 7).   
Stellen Sie sich vor, Sie arbeiten in einer Arztpraxis und sollen ein System zur Verwaltung der Patienten im Wartezimmer implementieren.
Das System soll Patienten, die in die Praxis kommen, ins Wartezimmer verweisen.
Im Büro des Arztes darf sich maximal ein Patient gleichzeitig befinden (der Patient, der gerade vom Arzt behandelt wird).
Sobald das Büro des Arztes wieder frei wird (d.h. die Behandlung abgeschlossen ist), kann nun ein Patient im Wartezimmer aufgerufen und in das Büro des Arztes gebeten werden.
Für jeden Aufruf wird immer der Patient gewählt, der die längste Zeit im Wartezimmer verbracht hat.

```python
class Patient:
    """Modelliert einen Patienten in einfacher Form"""
    
    def __init__(self, name, alter):
        """Konstruktor
        
        Args:
            name: Name des Patienten (str)
            alter: Alter des patienten (int)
        """
        self.name = name
        self.alter = alter
        
    def get_name(self):
        """Gibt den Namen des Patienten zurück
        
        Returns:
            str Name des Patienten
        """
        return self.name
    
    def get_alter(s):
        """Gibt das Alter des Patienten zurück
        
        Returns:
            str Alter des Patienten
        """
        return s.alter
    
    def __str__(self):
        """Packt die Informationen zu diesem Patient in einen String
        
        
        Returns:
            str Patienteninfos
        """
        return "Name: " + self.name + ", Alter: " + str(self.alter)
```

### a)
Implementieren Sie eine Klasse `Arztpraxis`, welche die folgenden Methoden bereitstellt:
- `betrete_praxis(p)`: Der gegebene Patient `p` betritt die Praxis und wird in das Wartezimmer gebeten.
- `naechster_patient_zum_arzt()`: Entfernt den Patienten aus dem Wartezimmer, welcher die längste Zeit dort verbracht hat, schickt diesen zum Arzt und gibt diesen zurück. Falls das Wartezimmer leer ist oder sich noch ein Patient in Behandlung befindet, werden keine Aktionen ausgeführt und lediglich `None` zurückgegeben.
- `beim_arzt()`: Gibt den Patienten zurück, welcher sich gerade beim Arzt befindet. Falls sich kein Patient gerade beim Arzt befindet, wird `None` zurückgegeben.
- `behandlung_abgeschlossen()`: Der Patient, der sich aktuell in Behandlung befindet verlässt die Praxis, sodass das Zimmer des Arztes nun wieder leer ist.

<!-- grade_id: praxis_1_sol, points: 0, grade: false, solution: true -->
```python
class Arztpraxis:
    ### BEGIN SOLUTION
    """Klasse für die Verwaltung von Patienten in einer Arztpraxis"""
    
    def __init__(self):
        """Konstruktor"""
        self.warteschlange = []
        self.inBehandlung = None
        
    def betrete_praxis(self, p):
        """Der gegebene Patient p betritt die Praxis und wird in das Wartezimmer gebeten.
        
        Args:
            p: Der hinzugekommene Patient (Patient)
        """
        self.warteschlange.append(p)
    
    def naechster_patient_zum_arzt(self):
        """Entfernt den Patienten aus dem Wartezimmer, welcher die längste Zeit dort verbracht hat und 
        schickt diesen zum Arzt. Falls das Wartezimmer leer ist oder sich noch ein Patient in Behandlung befindet, 
        werden keine Aktionen ausgeführt.
        
        Returns:
            Patient Den Patienten der zum Arzt geschickt wurde, oder None, falls keine Aktionen ausgeführt wurden.
        """
        # Büro des Arztes noch besetzt:
        if not self.beim_arzt() == None:
            return None
        
        # Kein Patient mehr im Wartezimmer:
        if len(self.warteschlange) == 0:
            return None
        
        p = self.warteschlange.pop(0)
        self.inBehandlung = p
        return p
    
    def beim_arzt(self):
        """Gibt den Patienten zurück, welcher sich gerade beim Arzt befindet.
        
        Returns:
            Patient Den Patienten, welcher aktuell in Behandlung ist oder None, sich kein Patient gerade beim Arzt befindet.
        """
        return self.inBehandlung
    
    def behandlung_abgeschlossen(self):
        """Der Patient, der sich aktuell in Behandlung befindet verlässt die Praxis. 
        Nach Aufruf dieser Methode gilt beim_arzt() == None.
        """
        self.inBehandlung = None
    ### END SOLUTION    
```

<!-- grade_id: praxis_1_test, points: 3, grade: true, solution: false -->
```python
praxis = Arztpraxis()
patienten = [Patient("Max M", 26), Patient("Ulrike Eustermann", 13), Patient("Michaela M", 20)]

# Patienten betreten die Praxis:
praxis.betrete_praxis(patienten[0])
praxis.betrete_praxis(patienten[1])
praxis.betrete_praxis(patienten[2])
assert(praxis.beim_arzt() == None)

# Rufe Patienten auf:
praxis.naechster_patient_zum_arzt()
assert(praxis.beim_arzt() == patienten[0])

praxis.behandlung_abgeschlossen()
assert(praxis.beim_arzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechster_patient_zum_arzt()
assert(praxis.beim_arzt() == patienten[1])

### HIDESTART
praxis = Arztpraxis()
patienten = [Patient("Ralf B", 26), Patient("Jessica W", 13), Patient("Philipp K", 20)]

# Teste Randfall(1):
assert(praxis.naechster_patient_zum_arzt() == None)

# Patienten betreten die Praxis:
praxis.betrete_praxis(patienten[0])
praxis.betrete_praxis(patienten[1])
praxis.betrete_praxis(patienten[2])
assert(praxis.beim_arzt() == None)

# Rufe Patienten auf:
praxis.naechster_patient_zum_arzt()
assert(praxis.beim_arzt() == patienten[0])

# Teste Randfall(2):
assert(praxis.naechster_patient_zum_arzt() == None)

praxis.behandlung_abgeschlossen()
assert(praxis.beim_arzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechster_patient_zum_arzt()
assert(praxis.beim_arzt() == patienten[1])
### HIDEEND
```

Unsere Arztpraxis soll nun zwischen Kassen- und Privatpatienten unterscheiden.
Für Privatpatienten soll zusätzlich zu deren Namen und Alter auch noch die Anschrift (als String) gepspeichert werden.
Privatpatienten werden von der Praxis bevorzugt behandelt, d.h. für die Auswahl des als nächstes zu behandelden Patienten wird immer erst geprüft, ob sich ein Privatpatient im Wartezimmer befindet.
Ist dies der Fall so wird dieser den Kassenpatienten vorgezogen, unabhängig von der Zeit, die er im Wartezimmer verbracht hat.
Falls sich mehrere Privatpatienten im Wartezimmer befinden, so wird der sich am längsten im Wartezimmer befindliche Privatpatient als nächstes aufgerufen.   

## b) 
Erweitern Sie zunächst die Klasse Patient zu einer Klasse `Privatpatient`, welche über ein Attribut `anschrift` sowie eine Methode `get_anschrift()` verfügt.
Überschreiben Sie auch die Methode `__str__()`, sodass die Tests korrekt laufen.

<!-- grade_id: praxis_2_sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Privatpatient(Patient):
    """Realisiert einen von der Arztpraxis bevorzugten Patienten"""
    
    def __init__(self, anschrift, *args):
        """ Konstruktor
        
        Args:
            anschrift: Anschrift des Privatpatienten (str)
        """
        super().__init__(*args)
        self.anschrift = anschrift
        
    def get_anschrift(self):
        """Gibt die Anschrift des Privatpatienten zurück
        
        Returns:
            str Anschrift des Privatpatienten
        """
        return self.anschrift
    
    def __str__(self):
        """Überschreibt die Standardausgabe
        
        Returns:
            Informationen über den Privatpatienten (str)
        """
        return "<Privatpatient> " + super().__str__() + " Anschrift: " + self.anschrift
### END SOLUTION
```

<!-- grade_id: praxis_2_test, points: 1, grade: true, solution: false -->
```python
p = Privatpatient("Flotowstr. 52", "Daniela L", 28)
assert(isinstance(p, Patient))
assert(isinstance(p, Privatpatient))
assert(p.__str__() == "<Privatpatient> Name: Daniela L, Alter: 28 Anschrift: Flotowstr. 52")
### HIDESTART
p2 = Privatpatient("Konstanzer Strasse 73", "Uwe P", 25)
assert(isinstance(p2, Patient))
assert(isinstance(p2, Privatpatient))
assert(p2.__str__() == "<Privatpatient> Name: Uwe P, Alter: 25 Anschrift: Konstanzer Strasse 73")
### HIDEEND
```

## c)
Erweitern Sie nun Ihre Klasse `Arztpraxis` aus Aufgabenteil a) zu einer Klasse `Arztpraxisunfair`, sodass diese nun Privatpatienten wie oben beschrieben bevorzugt behandelt.

<!-- grade_id: praxis_3_sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Arztpraxisunfair(Arztpraxis):
    """Erweiterung der Klasse Arztpraxis, sodass Privatpatienten bevorzugt werden."""
    
    def __init__(self):
        """Konstruktor"""
        super().__init__()
        self.privat_schlange = []
        
    def betrete_praxis(self, p):
        """Der gegebene Patient p betritt die Praxis und wird in das Wartezimmer gebeten. 
        Privatpatienten werden separat gespeichert.
        
        Args:
            p: Der hinzugekommene Patient (Patient)
        """
        if isinstance(p, Privatpatient):
            self.privat_schlange.append(p)
        else:
            super().betrete_praxis(p)
    
    def naechster_patient_zum_arzt(self):
        """Entfernt den Patienten aus dem Wartezimmer, welcher die längste Zeit dort verbracht hat und 
        schickt diesen zum Arzt. Falls das Wartezimmer leer ist oder sich noch ein Patient in Behandlung befindet, 
        werden keine Aktionen ausgeführt. Privatpatienten werden bei der Auswahl den normalen Patienten vorgezogen.
        
        Returns:
            Patient Den Patienten der zum Arzt geschickt wurde, oder None, falls keine Aktionen ausgeführt wurden.
        """        
        # Büro des Arztes noch besetzt:
        if not self.beim_arzt() == None:
            return None
        
        # Kein Patient mehr im Wartezimmer:
        if len(self.privat_schlange) and len(self.warteschlange) == 0:
            return None
        
        p = None
        if len(self.privat_schlange) > 0:
            # Ziehe Privatpatienten vor:
            p = self.privat_schlange.pop(0)
            self.inBehandlung = p
            return p
        else:
            # Normaler Patient kann zum Arzt:
            super().naechster_patient_zum_arzt()
### END SOLUTION
```

<!-- grade_id: praxis_4_test, points: 2, grade: true, solution: false -->
```python
praxis = Arztpraxisunfair()
patienten = [Patient("Max M", 26), Patient("Ulrike E", 13), Patient("Michaela M", 20), 
             Privatpatient("Flotowstr. 52", "Daniela L", 28), Privatpatient("Brandenburgische Str. 51", "Markus S", 56)]

# Patienten betreten die Praxis:
praxis.betrete_praxis(patienten[0])
praxis.betrete_praxis(patienten[1])
praxis.betrete_praxis(patienten[2])
praxis.betrete_praxis(patienten[3])
praxis.betrete_praxis(patienten[4])
assert(praxis.beim_arzt() == None)

# Rufe Patienten auf:
praxis.naechster_patient_zum_arzt()
assert(praxis.beim_arzt() == patienten[3])

praxis.behandlung_abgeschlossen()
assert(praxis.beim_arzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechster_patient_zum_arzt()
assert(praxis.beim_arzt() == patienten[4])

praxis.behandlung_abgeschlossen()
assert(praxis.beim_arzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechster_patient_zum_arzt()
assert(praxis.beim_arzt() == patienten[0])

### HIDESTART
praxis = Arztpraxisunfair()
patienten = [Patient("Ralf B", 26), Privatpatient("Brandenburgische Str. 51", "Markus S", 56)]

# Teste Randfall(1):
assert(praxis.naechster_patient_zum_arzt() == None)

# Patienten betreten die Praxis:
praxis.betrete_praxis(patienten[0])
praxis.betrete_praxis(patienten[1])
assert(praxis.beim_arzt() == None)

# Rufe Patienten auf:
praxis.naechster_patient_zum_arzt()
assert(praxis.beim_arzt() == patienten[1])

# Teste Randfall(2):
assert(praxis.naechster_patient_zum_arzt() == None)

praxis.behandlung_abgeschlossen()
assert(praxis.beim_arzt() == None)

# Rufe nächsten Patienten auf:
praxis.naechster_patient_zum_arzt()
assert(praxis.beim_arzt() == patienten[0])
### HIDEEND
```

<!-- grade_id: docu_test, points: 1, grade: true, solution: false -->
```python
# Tests für das Vorhandensein der Dokumentation

# Aufgabe 2:
assert(BinaryTree.remove.__doc__ != None)

# Aufgabe 3:
assert(Arztpraxis.betrete_praxis.__doc__ != None)
assert(Arztpraxis.naechster_patient_zum_arzt.__doc__ != None)
assert(Arztpraxis.beim_arzt.__doc__ != None)
assert(Arztpraxis.behandlung_abgeschlossen.__doc__ != None)

assert(Privatpatient.__init__.__doc__ != None)
assert(Privatpatient.get_anschrift.__doc__ != None)
assert(Privatpatient.__str__.__doc__ != None)

assert(Arztpraxisunfair.betrete_praxis.__doc__ != None)
assert(Arztpraxisunfair.naechster_patient_zum_arzt.__doc__ != None)
assert(Arztpraxisunfair.beim_arzt.__doc__ != None)
assert(Arztpraxisunfair.behandlung_abgeschlossen.__doc__ != None)
```