<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->


Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 20.11.2016 um 23:59 Uhr

## Aufgabe 1: Strings

a) (1 Punkt) Schreiben Sie eine Funktion namens ```firstLast```, welche einen ```string```-Parameter übergeben bekommt und das erste und letzte Zeichen des übergebenen Strings ausgibt. Dokumentieren Sie die Funktion, indem Sie Ihr einen ```docstring``` zuweisen und die Eingabe und Rückgabe parametern beschreiben. Drucken Sie anschließend das Attribut Ihres Funktions durch ```__doc__``` und ```help()``` aus. Führen Sie dieselbe Art von Dokumentation für **alle** Funktionen in diesen Übungsblatt. 

<!-- grade_id: 2, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def firstLast(val):
    """Returns the first and last characters of a string
   
    Args: 
        val (str): Input string
    Returns: 
        str: First and last character of the input
    """
    return val[0] + val[-1]

print(firstLast.__doc__)
help(firstLast)

word='fox'
print('First and last characters of', word,'are', firstLast(word))
### END SOLUTION
```

Output:
```
Returns the first and last characters of a string
   
    Args: 
        val (str): Input string
    Returns: 
        str: First and last character of the input
    
Help on function firstLast in module __main__:

firstLast(val)
    Returns the first and last characters of a string
    
    Args: 
        val (str): Input string
    Returns: 
        str: First and last character of the input

First and last characters of fox are fx
```

<!-- grade_id: 3, points: 1, grade: true, solution: false -->
```python
assert(firstLast('fox')=='fx')
```

b) (1 Punkt) Schreiben Sie eine Funktion namens ```reverse```, welche einen ```string```-Parameter übergeben bekommt und die Zeichenfolge in umgekehrter Reihenfolge zurückgibt. 

<!-- grade_id: 11, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def reverse(val):
    """Returns the characters of a string in reverse order
   
    Args: 
        val (str): Input string
    Returns: 
        str: reverse order of the input
    """
    return  val[::-1]

word='fox'
print('Reverse order of', word,'is', reverse(word))
### END SOLUTION
```

Output:
```
Reverse order of fox is xof
```

<!-- grade_id: 12, points: 1, grade: true, solution: false -->
assert(reverse('fox')=='xof')

c) (1 Punkt) Schreiben Sie eine Funktion namens ```concat```, welche zwei ```string```-Parameter übergeben bekommt und die Verkettung dieser beiden Strings zurückgibt.

<!-- grade_id: 14, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def concat(str1, str2):
    """Concatenates the input strings
   
    Args: 
        str1 (str): first input string
        str2 (str): second input string
    Returns: 
        str: concatenation of the input
    """
    return  str1 + str2

word1 = 'Kunst'
word2 = 'stoff'
print('Concatenation of', word1, 'and', word2, 'is', concat(word1, word2))
### END SOLUTION
```

Output:
```
Concatenation of Kunst and stoff is Kunststoff
```

<!-- grade_id: 15, points: 1, grade: true, solution: false -->
```python
assert(concat('Kunst','stoff' )=='Kunststoff')
```

d) (1 Punkt) Schreiben Sie eine Funktion namens ```sumLength```, welche zwei ```string```-Parameter übergeben bekommt und die Gesamtanzahl der Zeichen in beiden Strings zurückgibt.  

<!-- grade_id: 17, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def sumLength(str1, str2):
    """Returns the total size of input characters
   
    Args: 
        str1 (str): first input string
        str2 (str): second input string
    Returns: 
        int: Gesamtanzahl der Zeichen in beiden Strings
    """
    return  len(str1) + len(str2)

word1 = 'Kunst'
word2 = 'stoff'
print('Totalsize of characters of', word1, 'and', word2, 'is', sumLength(word1, word2))
### END SOLUTION
```

Output:
```
Totalsize of characters of Kunst and stoff is 10
```

<!-- grade_id: 18, points: 1, grade: true, solution: false -->
```python
assert(sumLength('Kunst','stoff' )==10)
```

e) (1 Punkt) Wie würden Sie den Inhalt eines ```string``` verändern? Nennen Sie mindestens zwei verschiedene Möglichkeiten und erklären Sie, warum diese funktionieren.

<!-- grade_id: 20, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION

# Solution A: list conversion
a = 'string'
b=list(a)
b[0]='d'
c = ''.join(b)
#try to work also with str(b) and show that it does not work
print(a,b,c)

# Solution B: direct modification
a = 'string'
b = 'd' + a[2:]
print(a,b)

### END SOLUTION
    
```

Output:
```
string ['d', 't', 'r', 'i', 'n', 'g'] dtring
string dring
```

## Aufgabe 2: Tuples

a) (1 Punkt) Schreiben Sie eine Funktion namens ```tupleConv```, welche zwei Parametern von beliebigen Datentypen übergeben bekommt und diese zusammen als ```tuple``` zurückgibt.

<!-- grade_id: 22, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def tupleConv(par1, par2):
    """Returns the input parameters as tuple
   
    Args: 
        par1 : first input parameter
        par2 : second input parameter
    Returns: 
        int: Gesamtanzahl der Zeichen in beiden Strings
    """
    return  (par1,par2)

word1 = 'Kunst'
word2 = 2
print('A tuple of parameters', word1, 'and', word2, 'is', tupleConv(word1, word2))
### END SOLUTION
```

Output:
```
A tuple of parameters Kunst and 2 is ('Kunst', 2)
```

<!-- grade_id: 23, points: 1, grade: true, solution: false -->
```python
assert(tupleConv('Kunst',2)==('Kunst',2))
```

b) (1 Punkt) Schreiben Sie eine Funktion namens ```tupleCheck```, welche ein ```tuple``` und ein beliebiges Parameter übergebenen bekommt. Die Funktion soll checken ob der zweite Parameter in das erste ```tuple``` sich befindet und das Ergebniss zurückgeben.

<!-- grade_id: 25, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def tupleCheck(par1, par2):
    """Checks whether an element is in the tuple
   
    Args: 
        par1 (tuple): first input parameter
        par2 : second input parameter
    Returns: 
        bool: search result
    """
    return  par2 in par1

word1 = 'Kunst'
word2 = 2
tupleCheck((word1,word2),4)
### END SOLUTION
```

Output:
```
False
```

<!-- grade_id: 26, points: 1, grade: true, solution: false -->
```python
word1 = 'Kunst'
word2 = 2
assert(tupleCheck((word1,word2),2)==True)
```

c) (1 Punkt) Schreiben Sie eine Funktion namens ```tupleFilter```, welche ein ```tuple``` als Parameter übergeben bekommt und dessen erstes und vorletztes Element als Tuple zurückgibt.

<!-- grade_id: 28, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def tupleFilter(par):
    """Returns first element and second element from the last of a tuple
   
    Args: 
        par1 (tuple): input parameter
    Returns: 
        tuple: first and second element 
    """
    return  (par[0], par[-2])

a = ("1","2","3","4")
print(tupleFilter(a))
### END SOLUTION
```

Output:
```
('1', '3')
```

<!-- grade_id: 29, points: 1, grade: true, solution: false -->
```python
assert(tupleFilter(a)==("1","3"))
```

## Aufgabe 3: Listen und Mengen

a) (1 Punkt) Schreiben Sie eine Funktion namens ```singularList```, welche eine ```list```-Parameter übergeben bekommt und eine Liste mit Werten zurückgibt die nur einmal vorkommen. 

<!-- grade_id: 31, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def singularList(par):
    """Returns values that appear only once
   
    Args: 
        par1 (list): input parameter
    Returns: 
        list: singular values from the list
    """
    return  list(set(par))

a = [1,2,3,3,3,4,5,5,6]
print(singularList(a))
### END SOLUTION
```

Output:
```
[1, 2, 3, 4, 5, 6]
```

<!-- grade_id: 32, points: 1, grade: true, solution: false -->
```python
assert(singularList(a)==[1, 2, 3, 4, 5, 6])
```

b) (1 Punkt) Schreiben Sie eine Funktion namens ```setDifference```, welche zwei ```list```-Parameter übergeben bekommt und deren Mengendifferenz als ```list``` zurückgibt.

<!-- grade_id: 34, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def setDifference(list1, list2):
    """Returns the set difference between two lists
   
    Args: 
        list1 (list): first input parameter
        list2 (list): second input parameter
    Returns: 
        list: set difference between two lists
    """
    return  list(set(list1) - set(list2))

a = [1,2,3,3,3,4,5,5,6]
b = [1,3,3,3,4,5,6]
print(setDifference(a,b))
### END SOLUTION
```

Output:
```
[2]
```

<!-- grade_id: 35, points: 1, grade: true, solution: false -->
```python
assert(setDifference(a,b)==[2])
```

c) (2 Punkte) Schreiben Sie eine Funktion namens ```concLists```, welche zwei ```list```-Parametern übergeben bekommt und jede zweite Element von deren Hintereinanderreihung als ```string```-Datentyp zurückgibt.

<!-- grade_id: 40, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def concLists(list1, list2):
    """Returns the concatenation between two lists
   
    Args: 
        list1 (list): first input parameter
        list2 (list): second input parameter
    Returns: 
        str: concatenation of two lists
    """
    val = list1 + list2
    
    return  str(val[::2])

a = [1,2,3,3,3,4,5,5,6]
b = [1,3,3,3,4,5,6]
concLists(a,b)
### END SOLUTION

```

Output:
```
'[1, 3, 3, 5, 6, 3, 3, 5]'
```

<!-- grade_id: 41, points: 2, grade: true, solution: false -->
```python
assert(concLists(a,b)=='[1, 3, 3, 5, 6, 3, 3, 5]')
```

## Aufgabe 4: Dictionaries

a) (4 Punkte) Schreiben Sie zwei ```dict```-Variablen namens ```color_en_de``` und ```color_de_fr```, welche die Farben rot, grün, blau und gelb auf deutsch-englisch und auf deutsch-französisch enthält. Als Referenz benutzen Sie die folgende Tabelle: 

|DE   | EN   | FR|
|----|------|---|
|rot |red | rouge |   
|gelb |yellow  | jaune|
|grün  |green  | vert|
|blau  |blue | bleu |

Schreiben Sie dann, anhand der Daten eine Funktion namens ```translator```, welche die Farben von Englisch auf Französisch übersetzt. 

<!-- grade_id: 51, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
color_en_de = {"red" : "rot", "green" : "grün", "blue" : "blau", "yellow":"gelb"}
color_de_fr = {"rot" : "rouge", "grün" : "vert", "blau" : "bleu", "gelb":"jaune"}

def translator(word, color_en_de, color_de_fr):
    """Translates colors from english to french
   
    Args: 
        word (str): english word
        color_en_de (dict): english german dict
        color_de_fr (dict): german french dict
    Returns: 
        str: translation in french
    """    
    return color_de_fr[color_en_de[word]]

word="yellow"
print("The french word for", word, "is", translator(word, color_en_de, color_de_fr))
### END SOLUTION
```

Output:
```
The french word for yellow is jaune
```

<!-- grade_id: 52, points: 4, grade: true, solution: false -->
```python
color_en_de = {"red" : "rot", "green" : "grün", "blue" : "blau", "yellow":"gelb"}
color_de_fr = {"rot" : "rouge", "grün" : "vert", "blau" : "bleu", "gelb":"jaune"}
assert(translator("yellow", color_en_de, color_de_fr)=="jaune")
```

## Aufgabe 5: Funktionen

a) (1 Punkt) Gegeben ist die folgende Funktion: 

```Python
def f(a):
    return a+3
```

Welche Werte werden bei folgendem Aufruf ausgegeben:

```Python
a=2
print(f(a+3), f(f(a+3)))
``` 
Begründen Sie kurz.

<!-- grade_id: 61, points: 1, grade: true, solution: true -->
### BEGIN SOLUTION
Ausgabe: 8, 11
    Grund für 8: 
    Erste Aufruf: 2+3=5
    Zweite Aufruf: 5+3 = 8
        
    Grund für 11
    Erste Aufruf: 2 + 3 = 5
    Zweite Aufruf: 5 + 3 = 8
    Dritte Aufruf: 8 + 3 = 11
### END SOLUTION

b) (1 Punkt) Was ist der Rückgabewert des folgenden Funktionsaufrufs? Begründen Sie kurz.

```Python
def ausgabe(x):
    print(x,2*x)
    return 2*x
    print(x, 3*x)
    return 3*x
```

<!-- grade_id: 63, points: 1, grade: true, solution: true -->
### BEGIN SOLUTION
Ausgabe: 2*x
    Grund: nach dem ersten ```return``` wird die Funktion nicht mehr weiter Ausgeführt.
### END SOLUTION

c) (1 Punkt) Gegeben ist folgendes Programm. Welche Werte haben am Ende die in der drittletzten Zeile deklarierten Variablen $a,b,x,y$? Begründen Sie kurz.

```Python
def foo(x, y):
    a = 24
    x,y = y,x
    b = 4
    b = 16
    c = 10
    print(a,b,x,y)

a,b,x,y = 1,15,3,4
foo(17,4)
print(a,b,x,y)
```

<!-- grade_id: 65, points: 1, grade: true, solution: true -->
### BEGIN SOLUTION
Die Variablen $a,b,x,y$ haben folgende Werte: 1, 15, 3, 4. Grund: sie werden nicht von ```foo``` beeinflusst. 
### END SOLUTION        

d) Gegeben sind die folgenden beiden Funktionen:

```Python
def foo1(a):
     print(a + 1)
def foo2(a):
     return a + 1
```     
Einer der beiden folgenden Aufrufe ist fehlerhaft. Finden Sie, welcher es ist und begründen Sie kurz.
```Python
foo1(3) + 1
foo2(3) + 1
```

<!-- grade_id: 67, points: 1, grade: true, solution: true -->
### BEGIN SOLUTION
Die erste Aufruf ist fehlerhaft, da ```foo1``` kein Wert zurückgibt (hat None als Rückgabewert.)
### END SOLUTION