<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->


Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 27.11.2016 um 23:59 Uhr
### Wichtig: Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue04.ipynb ab (nicht umbenennen, keine Kopie erstellen!). Sonst kann eine Bewertung nicht stattfinden.

Bitte entfernen Sie auch die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben.

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab (Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht).

## Aufgabe 1: Round-off

(2 Punkte) Erklären Sie kurz was das folgende Programm tut und führen Sie es aus. 
Eigentlich sollte ```r``` nach der Programmablauf immer noch gleich 4.0 sein, ist es aber nicht. 
Ab welchem Wert von ```n``` wird ```r=1.0```?
Geben Sie ```r``` in beiden Schleifen aus. 
Erklären Sie anhand dieser Ausgaben, warum ```r``` am Ende gleich 1.0 ist und woran dies liegt.


```Python
for n in range(1, 60):
    r = 4.0
    for i in range(n):
        r = r**0.5
    for i in range(n):
        r = r**2
    print(n,r)
```


<!-- grade_id: 2, points: 0, grade: true, solution: true -->
### BEGIN SOLUTION 
Ab n=53 wird r = 1.0. Das passiert wegen Abrundungsfehlern.

Aufgrund der endlichen Mantisse in der Gleitkommadarstellung lassen sich Zahlen auf einem Computer nicht beliebig genau darstellen. Es muss gerundet werden und die Rundungsfehler vergrössern sich nach jeder Wurzel- und Potenzberechnung.
### END SOLUTION 

## Aufgabe 2: Listen

a) (1 Punkt) Führen Sie das folgende Programm aus: 

```Python
a = [1,2,3,4]
c = a
a[0] = 0
print(c)
```
Warum hat die Liste ```c``` am Ende den gleichen Inhalt wie ```a```? Begründen Sie kurz.

<!-- grade_id: 30001, points: 1, grade: true, solution: true -->


b) (1 Punkt) Führen Sie das folgende Programm aus: 

```Python
def list_return(lst):
    return lst  

a = [1, 2, 3, 4]
b = list_return(a)
a[0] = 0
print(a,b)
```

Obwohl die Variable ```b``` von der Funktion ```list_return``` initialisiert wurde, hat sie den gleichen Inhalt wie die modifizierte Liste ```a```. Begründen Sie kurz.

<!-- grade_id: 30011, points: 1, grade: true, solution: true -->


c) (1 Punkt) Führen Sie das folgende Programm aus:

```Python
def list_return(lst):
    return lst  

a = [1, 2, 3, 4]
b = list_return(a[:])
a[0] = 0
print(a,b)
```
Warum wird jetzt der Inhalt von ```b``` nicht durch die Modifizierung der Liste ```a``` geändert? 

<!-- grade_id: 30021, points: 1, grade: true, solution: true -->


d) (1 Punkt) Schreiben Sie eine neue Funktion namens ```list_copy_shallow```, welche eine Liste übergeben bekommt und eine *flache* Kopie dieser Liste zurückgibt (siehe zum Unterschied von flacher und tiefer Kopie einer Liste Kapitel 5 der VL-Unterlagen).

Nutzen Sie dazu keine vordefinierte Copy-Funktionalität von Listen.

<!-- grade_id: 30031, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def list_copy_shallow(lst):
    return lst[:]  
### END SOLUTION 
```

<!-- grade_id: 30032, points: 1, grade: true, solution: false -->
```python
a = [1, 2, 3, 4]
b = list_copy_shallow(a)
b[0] = 2
assert(a != b)

### HIDESTART
c = [ [1, 2], 3 ]
d = list_copy_shallow(c)
c[0][0] = 2
assert(c == d)
### HIDEEND
```

e) (2 Punkte) Schreiben Sie eine weitere Funktion namens ```list_copy_deep```, welche eine Liste übergeben bekommt und eine *tiefe* Kopie dieser LIste zurückgibt.

Nutzen Sie dazu keine vordefinierte Copy Funktionalität von Listen. 

<!-- grade_id: 30041, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
import copy
def list_copy_deep(lst):
    return copy.deepcopy(lst) #of course, this is only a temporary solution!
### END SOLUTION 
```

<!-- grade_id: 30042, points: 2, grade: true, solution: false -->
```python
a = [1, 2, 3, 4]
b = [5, 6, 7, 8]
c = 9
d = [a, b, c]

e = list_copy_deep(d)
assert(d == e)
d[0][0] = 2
assert(d != e)
```

f) (1 Punkt) List modify Funktion. Er modifiziert eine Liste.

```python
g) (1 Punkt) Eine 2x2 Matrix als Liste von Listen erzeugen. Eine Funktion schreiben, welche die Transposition macht.
```

## Aufgabe 3: Strings

Zeichenketten in Python können Zeichen aus einem großen Zeichenvorrat annehmen; typischerweise dem sog. Universal Coded Character Set Transformation Format mit 8 bits (UTF-8). Damit stehen zum Beispiel Zeichen asiatischer und arabischer Sprachen bereit. 

a) (1 Punkt) Finden Sie heraus, wie viele Bytes benötigt werden, um diese Zeichenkette intern darzustellen. Wievielen Bytes pro Zeichen entspricht das (achten Sie auf die beiden Leerzeichen!) Schreiben Sie eine Funktion der Sie diesen String übergeben und die die Anzahl der Bytes zurückgibt. 

```Python
# Beispiel-Zeichenkette in UTF-8
s = '我们 爱 编程'
s = "äöü"
s = "éàß"

# Lösung ()
def laenge(s):
    return len(bytes(s, "utf-8"))


# Tests
assert(laenge('我们 爱 编程') == 17)
assert(laenge('我们爱') == 9)
assert(laenge('äöü') == 6)
assert(laenge('éàß') == 6)
```

b) (1 Punkt) Untersuchen Sie, wie sich slicing für solche Zeichenketten benimmt. Welchen Slicing-Ausdruck müssen Sie anwenden, um die ersten beiden Zeichen zu erhalten?

c) (1 Punkt) Wäre es sinnvoll, bei solchen Zeichenketten stets die gleiche Anzahl an Bytes pro Zeichen zu benutzen? Begründen Sie! Welchen Verhalten scheint bei Python vorzuliegen?

## Aufgabe 4: Funktionen

Warum muss man die Reihenfolge der Auswertung von Funktionsparametern spezifieren? Wie ist das spezifiziert? Woran würde man das merken? Gibt es ein Beispiel, außer Funktionen mit Seiteneffekt im Aufruf einer Funktion bei dem das auftritt?

Siehe: https://www.securecoding.cert.org/confluence/display/c/EXP30-C.+Do+not+depend+on+the+order+of+evaluation+for+side+effects

 https://docs.python.org/2/reference/expressions.html#evaluation-order)
 
 *** UEBUNG Warum ist die Ferstlegung *vor* Position einfügen sinnvoll?

## Aufgabe 5: Tuples

http://openbookproject.net/thinkcs/python/english3e/tuples.html

Tuples als Buch von Douglas Adams, per Anhalter durch den Galaxis
Einpacken: name, Autor, Veröffentlichung
Auspacken: Möchte gerne Wissen wie der Autor und Veröffentlichung lauten


http://stackoverflow.com/questions/6967632/unpacking-extended-unpacking-and-nested-extended-unpacking

der PEP3132 ist daszu ja ganz witzig.

Tuple ist nicht das gleiche wie mehrere Parameter:

Ein Funktion mit return tuple, und möglicherweise kombiniere das mit die Fragen vom Stackoverflow.
ein Funktion

* Tuple kann man nicht ändern, auch wenn man das vielleicht denkt , siehe:
*** UEBUNG Tuple ändern durch slicing und ändern


* Tuple ist nicht das gleiche wie mehrere Parameter:

*** UEBUNG Falscher Aufruf  von quersumme

```python

```