<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->


Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 27.11.2016 um 23:59 Uhr
### Wichtig: Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue04.ipynb ab (nicht umbenennen, keine Kopie erstellen!). Sonst kann eine Bewertung nicht stattfinden.

Bitte entfernen Sie auch die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben.

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab (Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht).

## Aufgabe 1: Round-off

(2 Punkte) Erklären Sie kurz, was das folgende Programm tut und führen Sie es aus. 
Eigentlich sollte ```r``` nach der Programmablauf immer noch gleich 4.0 sein, ist es aber nicht. 
Ab welchem Wert von ```n``` wird ```r=1.0```?
Geben Sie ```r``` in beiden Schleifen aus. 
Erklären Sie anhand dieser Ausgaben, warum ```r``` am Ende gleich 1.0 ist und woran dies liegt.


```Python
for n in range(1, 60):
    r = 4.0
    for i in range(n):
        r = r**0.5
    for i in range(n):
        r = r**2
    print(n,r)
```


<!-- grade_id: 2, points: 0, grade: true, solution: true -->
### BEGIN SOLUTION 
Ab n=53 wird r = 1.0. Das passiert wegen Abrundungsfehlern.

Aufgrund der endlichen Mantisse in der Gleitkommadarstellung lassen sich Zahlen auf einem Computer nicht beliebig genau darstellen. Es muss gerundet werden und die Rundungsfehler vergrössern sich nach jeder Wurzel- und Potenzberechnung.
### END SOLUTION 

## Aufgabe 2: Listen

a) (1 Punkt) Führen Sie das folgende Programm aus: 

```Python
a = [1,2,3,4]
c = a
a[0] = 0
print(c)
```
Warum hat die Liste ```c``` am Ende den gleichen Inhalt wie ```a```? Begründen Sie kurz.

<!-- grade_id: 30001, points: 1, grade: true, solution: true -->


b) (1 Punkt) Führen Sie das folgende Programm aus: 

```Python
def list_return(lst):
    return lst  

a = [1, 2, 3, 4]
b = list_return(a)
a[0] = 0
print(a,b)
```

Obwohl die Variable ```b``` von der Funktion ```list_return``` initialisiert wurde, hat sie den gleichen Inhalt wie die modifizierte Liste ```a```. Begründen Sie kurz.

<!-- grade_id: 30011, points: 1, grade: true, solution: true -->


c) (1 Punkt) Führen Sie das folgende Programm aus:

```Python
def list_return(lst):
    return lst  
    
a = [1, 2, 3, 4]
b = list_return(a[:])
a[0] = 0
print(a,b)
print(a is b)
print(a == b) 
```
Warum wird jetzt der Inhalt von ```b``` nicht durch die Modifizierung der Liste ```a``` geändert? 

<!-- grade_id: 30021, points: 1, grade: true, solution: true -->


d) Führen Sie das folgende Programm aus: 

```Python
def list_return(lst):
    return lst 
    
def list_modify(lst):
    lst = lst*2
    return lst
    
a = [1, 2, 3]
b = list_return(a)
print(a,b)
print(b is a) 
print(b == a)

c = list_return(a[:])
print(a,c)
print(c is a)
print(c == a)

d = list_modify(a)
print(a is d)
```

Obwohl die Variablen ```c``` und ```a``` den gleichen Inhalt haben, ist die Bewertung ```c is a``` gleich Falsch. Woran liegt das? Begründen Sie kurz. Was liefern die Operatoren ```is``` und ```==``` ?

<!-- grade_id: 300009999, points: 1, grade: true, solution: true -->


e) (1 Punkt) Schreiben Sie eine neue Funktion namens ```list_copy_shallow```, welche eine Liste übergeben bekommt und eine *flache* Kopie dieser Liste zurückgibt (siehe zum Unterschied von flacher und tiefer Kopie einer Liste Kapitel 5 der VL-Unterlagen).

Nutzen Sie dazu keine vordefinierte Copy-Funktionalität von Listen.

<!-- grade_id: 30031, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def list_copy_shallow(lst):
    return lst[:]  
### END SOLUTION 
```

<!-- grade_id: 30032, points: 1, grade: true, solution: false -->
```python
a = [1, 2, 3, 4]
b = list_copy_shallow(a)
b[0] = 2
assert(a != b)

### HIDESTART
c = [ [1, 2], 3 ]
d = list_copy_shallow(c)
c[0][0] = 2
assert(c == d)
### HIDEEND
```

f) (1 Punkt) Führen Sie das folgende Programm aus: 

```Python
def f(a, b):
   print("a: ", a)
   print("b: ", b)

def g(L):
   L[0] = L[0] - 1
   print (L)
   return L

def h(L):
   L[0] = L[0] + 2
   print (L)
   return L

L = [5, 17]

f(g(L), h(L))
```
Warum haben die Variablen ```a``` und ```b``` dieselbe Belegung? Begründen Sie kurz. 

<!-- grade_id: 3000021, points: 1, grade: true, solution: true -->
###BEGIN SOLUTION

Die beiden Variablen haben dieselbe Belegung, weil es insgesamt nur eine einzige Liste L gibt.
Zwar wird deren erstes Element in beiden Funktionen verändert, aber beide Funktionen geben dieselbe Liste zurück, somit ist L der aktuelle Parameter sowohl für den formalen Parameter a als auch für den formalen Parameter b.

###END SOLUTION

g) (1 Punkt) Wie muss die Funktion ```f()``` aufgerufen werden, damit die folgende Ausgabe:

    a:  [4, 17]
    b:  [7, 17]

ausgegeben wird? 
Nutzen Sie dazu die bereits vorhandene Liste
```Python
    L = [5, 17]
```
und verwenden Sie Funktionsaufrufe von ```g()``` und ```h()``` als Parameter eines Aufrufs von ```f()```.

<!-- grade_id: 3000134, points: 1, grade: true, solution: true -->
```python
def f(a, b):
   print("a: ", a)
   print("b: ", b)

def g(L):
   L[0] = L[0] - 1
   print (L)
   return L

def h(L):
   L[0] = L[0] + 2
   print (L)
   return L

L = [5, 17]

###BEGIN SOLUTION
f(g(L[:]), h(L[:]))
#END SOLUTION
```

Output:
```
[4, 17]
[7, 17]
a:  [4, 17]
b:  [7, 17]
```

<!-- grade_id: 30042, points: 2, grade: true, solution: false -->
```python
a = [1, 2, 3, 4]
b = [5, 6, 7, 8]
c = 9
d = [a, b, c]

e = list_copy_deep(d)
assert(d == e)
d[0][0] = 2
assert(d != e)
```

h) (1 Punkt) Gegegeben sei die folgende 2x2-Matrix:

```
1 2
3 4
```

Betrachten Sie folgende Umsetzung einer solchen Matrix in Python:
```Python
    matrix = [[1, 2], [3, 4]]
```

Schreiben Sie eine Funktion namens ```transpose_matrix```, welche eine in dieser Art implementierte 2x2-Matrix als Eingabe erhält und die transponierte Matrix zurückgibt.
Die Matrix selbst soll dabei nicht verändert werden!

Zum Thema Transponieren von Matrizen lesen Sie nötigenfalls hier nach: https://de.wikipedia.org/wiki/Transponierte_Matrix.

<!-- grade_id: 30061, points: 0, grade: true, solution: true -->
```python
###BEGIN SOLUTION
def transpose_matrix(matrix):
    return [[matrix[0][0], matrix[1][0]], [matrix[0][1], matrix[1][1]]]                
### END SOLUTION
```

<!-- grade_id: 30062, points: 1, grade: true, solution: false -->
```python
matrix = [[1, 2], [3, 4]]
matrix_transposed = transpose_matrix(matrix)

assert(matrix_transposed[0][0] == 1)
assert(matrix_transposed[0][1] == 3)
assert(matrix_transposed[1][0] == 2)
assert(matrix_transposed[1][1] == 4)

### HIDESTART
matrix2 = [[5, 6], [7, 8]]
matrix2_transposed = transpose_matrix(matrix2)
assert(matrix2_transposed == [[matrix2[0][0], matrix2[1][0]], [matrix2[0][1], matrix2[1][1]]])
### HIDEEND
```

i) (2 Punkte) Schreiben Sie eine weitere Funktion namens ```matrix_copy_deep```, welche eine Liste übergeben bekommt, die eine 2x2-Matrix repräsentiert, und eine *tiefe* Kopie dieser Liste zurückgibt.

Nutzen Sie dazu keine vordefinierte Copy-Funktionalität von Listen. 

<!-- grade_id: 300000041, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def matrix_copy_deep(matrix):
    return transpose_matrix(transpose_matrix(matrix))
### END SOLUTION 
```

<!-- grade_id: 30041, points: 0, grade: true, solution: false -->
```python
matrix = [[1, 2], [3, 4]]
matrix_deep_copied = matrix_copy_deep(matrix)
matrix[0][0] = 5
assert(matrix != matrix_deep_copied)
```

## Aufgabe 3: Strings

Zeichenketten in Python können Zeichen aus einem großen Zeichenvorrat annehmen; typischerweise dem sog. Universal Coded Character Set Transformation Format mit 8 bits (UTF-8). 
Damit stehen zum Beispiel Zeichen asiatischer und arabischer Sprachen bereit. 

a) (1 Punkt) Finden Sie heraus, wie viele Bytes benötigt werden, um die folgenden Zeichenketten intern darzustellen.
Wie vielen Bytes pro Zeichen entspricht das (achten Sie auf die beiden Leerzeichen!)? (Schreiben Sie Ihre Antwort als Kommentar)


```Python
# Beispiel-Zeichenketten in UTF-8
s1 = '我们 爱 编程'
s2 = "äöü"
s3 = "éàß"
s4 = "Мы любим GP1"
s5 = "ما عاشق برنامه نویسی"
```

Schreiben Sie eine Funktion namens ```utf8_len```, die einen String übergeben bekommt und die die Anzahl der Bytes dieses Strings zurückgibt. 

Hinweis: Sie dürfen hierzu Funktionen der Python-Standardbibliothek nutzen, die nicht in der Vorlesung behandelt wurden.

<!-- grade_id: 3032, points: 0, grade: false, solution: true -->
```python
s1 = '我们 爱 编程'
s2 = "äöü"
s3 = "éàß"
s4 = "Мы любим GP1"
s5 = "ما عاشق برنامه نویسی"

###BEGIN SOLUTION
# Leerzeichen benötigen 1 Byte, Umlaute, das ß und Zeichen mit Akzent 2 Bytes und die asiatischen Bildzeichen 3 Bytes
def laenge(s):
    return len(bytes(s, "utf-8"))
###END SOLUTION
```

<!-- grade_id: 3033, points: 1, grade: true, solution: false -->
```python
s1 = '我们 爱 编程'
s2 = "äöü"
s3 = "éàß"
s4 = "Мы любим GP1"
s5 = "ما عاشق برنامه نویسی"

assert(laenge(s1) == 17)
assert(laenge(s2) == 6)
assert(laenge(s3) == 6)
assert(laenge(s4) == 19)
assert(laenge(s5) == 37)
```

b) (1 Punkt) Wäre es sinnvoll, bei solchen Zeichenketten stets die gleiche Anzahl an Bytes pro Zeichen zu benutzen? Begründen Sie! 

Welches Verhalten scheint bei Python vorzuliegen?

<!-- grade_id: 3037, points: 1, grade: true, solution: true -->
###BEGIN SOLUTION
Argumente für die Verwendung einer stets gleichen Anzahl an Bytes pro Zeichen:
- Leichtes Errechnen der Länge des Strings möglich (evtl. für Pointer-Arithmetik nützlich, für Python egal) 
- Verändern eines Zeichens ändert nicht die Länge des Strings (in Python ja nicht möglich)
Argumente für die Verwendung einer variablen Anzahl an Bytes pro Zeichen:
- Geringerer Speicherplatzbedarf bei Verwendung von Sprachen, die nur oder größtenteils aus 1-Byte-Zeichen bestehen
- Texte in Englisch lassen sich wie ASCII kodieren (in manchen Fällen höhere Kompatibilität)
- Theoretisch unbegrenzte Erweiterbarkeit des Zeichenvorrats

Python nutzt, wie in a) gesehen, eine variable Anzahl an Bytes pro Zeichen.
###END SOLUTION

## Aufgabe 4: Tuples

a) (2 Punkte) Schreiben Sie eine Funktion namens ```order_check```, welche die verschachtelte Tupel ```bestellung``` übergeben bekommt und die gewünschte Bestellung ausgibt. Z.B.: 

```text
Thomas möchte 3 Äpfel
Thomas möchte 1 Birnen
Anna möchte 6 Birnen
```

<!-- grade_id: 40002, points: 0, grade: false, solution: true -->
```python
bestellung = (('Thomas', (('Äpfel', 3), ('Birnen', 1))),
          ('Anna', (('Birnen', 6),)))

def order_check(orders):
###BEGIN SOLUTION    
    dict_fruits = {}
    for user, order in orders:
        for fruit, quantity in order:
            print("%s möchte %d %s" % (user, quantity, fruit))
###END SOLUTION
```

<!-- grade_id: 40003, points: 1, grade: true, solution: false -->
b) (2 Punkte) Erweitern Sie die Funktion ```order_check``` indem sie eine ```dict```-Variable mit der Frucht als Schlüsselwort und die entsprechende Gesamtmenge als Wert zurückgibt. 
Z.B. die Belegung von ```fruechten_menge```

```Python
fruechten_menge = order_check(bestellung)
```
soll so sein:
```Python    
{'Äpfel': 3, 'Birnen': 7}
```

<!-- grade_id: 400004, points: 0, grade: false, solution: true -->
```python
bestellung = (('Thomas', (('Äpfel', 3), ('Birnen', 1))),
          ('Anna', (('Birnen', 6),)))

def order_check(orders):
###BEGIN SOLUTION    
    """Print order list and return total fruit quantity

    Keyword arguments:
    orders -- (tuple) order list 
    
    Return:
    dict_fruits -- (dict) fruit quantity
    """
    dict_fruits = {}
    for user, order in orders:
        for fruit, quantity in order:
            print("%s möchte %d %s" % (user, quantity, fruit))
            
            if fruit in dict_fruits:
                dict_fruits.update({fruit: dict_fruits[fruit] + quantity})                 
            else:
                dict_fruits.update({fruit:quantity}) 
         
    return dict_fruits

fruechten_menge = order_check(bestellung)
###END SOLUTION
```

Output:
```
Thomas möchte 3 Äpfel
Thomas möchte 1 Birnen
Anna möchte 6 Birnen
```

<!-- grade_id: 400005, points: 2, grade: true, solution: false -->
```python
assert(fruechten_menge=={'Äpfel': 3, 'Birnen': 7})
assert(order_check.__doc__ != None)
```

# Aufgabe 5: Docstring

(3 Punkte) Dokumentieren Sie die Funktion ```multiples``` in Docstring gemäß PEP 257. 
Beschreiben Sie die Parameter und Rückgabewerte und was die Funktion leistet.

```Python
def multiples(f1, f2, max_val):
    a = []
    for number in range(1, max_val+1):
        if not number % f1 or not number % f2:
            a.append(number)
    return a
    
b= multiples(3,5,1000)
print(b)
```

<!-- grade_id: 500002, points: 0, grade: false, solution: true -->
```python
def multiples(f1, f2, max_val):
###BEGIN SOLUTION
    """Compute multiples of f1 and f2 up to max_val

    Keyword arguments:
    f1 -- first factor
    f2 -- second factor
    
    Return:
    a -- list containin all the multiples of f1 and f2
    """
###END SOLUTION
    a = []
    for number in range(1, max_val+1):
        if not number % f1 or not number % f2:
            a.append(number)
    return a
```

<!-- grade_id: 500003, points: 3, grade: true, solution: false -->
```python
assert(multiples.__doc__ != None)
```