<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->


# Aufgaben 

## Aufgabe 1: Arithmetik in Rechnern 

(2 Punkte) Wir betrachten das unten angegebene Programm. 

Beantworten Sie folgende Fragen: 

1. Erklären Sie kurz in eigenen Worten, was das Programm tut. 
1. Führen Sie es aus und beschreiben Sie kurz das Ergebnis. 
1. Wenn Sie sich die mathematischen Operationen in den Schleifen betrachten, welchen Wert für ```r``` müssten Sie dann eigentlich erwarten? 
1. Ab welchem Wert von ```n``` wird ```r=1.0``` am Ende des Programms?
1. Erklären Sie, warum ```r``` am Ende gleich 1.0 ist und woran dies liegt. Hinweis: Dazu können Sie beispielsweise das Programm verändern, um auch in den inneren Schleifen ```r``` auszugeben (Sie brauchen diese Ausgaben und diese Programmveränderungen nicht abzugeben!). 


```python
for n in range(1, 60):
    r = 4.0
    for i in range(n):
        r = r**0.5
    for i in range(n):
        r = r**2
    print(n,r)
```

Output:
```
1 4.0
2 4.000000000000002
3 3.9999999999999982
4 3.9999999999999982
5 3.999999999999986
6 3.999999999999986
7 3.999999999999986
8 3.9999999999998854
9 4.000000000000094
10 4.000000000000094
11 4.000000000000094
12 4.000000000000094
13 3.9999999999965343
14 3.999999999989317
15 3.999999999989317
16 3.999999999989317
17 4.000000000047098
18 4.000000000163432
19 4.000000000163432
20 4.000000000629434
21 4.000000000629434
22 4.000000000629434
23 4.000000004354343
24 4.000000011804699
25 4.000000026708689
26 4.000000026708689
27 3.9999999671099733
28 3.9999999671099733
29 3.9999999671099733
30 3.9999999671099733
31 3.9999999671099733
32 3.9999999671099733
33 3.9999961523092837
34 3.9999885230258903
35 3.9999885230258903
36 3.9999885230258903
37 3.9999885230258903
38 3.9999885230258903
39 3.9997443906581256
40 3.9997443906581256
41 3.9997443906581256
42 3.9977918677565043
43 3.9938896813415576
44 3.9860967274650654
45 3.9860967274650654
46 3.9550766750613335
47 3.9550766750613335
48 3.955076
```

<!-- grade_id: 2, points: 2, grade: true, solution: true -->
1. Es zieht wiederholt die Wurzel aus einer Zahl und quadriert danach wiedrer entsprechend oft. 
1. Für eine wachsende Anzahl an Wurzel- und Quadrieroperaionen bleibt das Ergebnis zunächst ungefähr 4. Danach fällt es und pendelt sich auf 1.0 ein.
1. Hätten wir es mit Daten unendlicher Genauigkeit zu tun, sollten sich die Wurzel- und Quadrierfunktionen aufheben (Umkehrfunktion). Es sollte stets wieder 4 herauskommen. 
1. Ab n=53 wird r = 1.0. Das passiert wegen Abrundungsfehlern.
1. Aufgrund der endlichen Mantisse in der Gleitkommadarstellung lassen sich Zahlen auf einem Computer nicht beliebig genau darstellen. Es muss gerundet werden und die Rundungsfehler vergrössern sich nach jeder Wurzel- und Potenzberechnung.

## Aufgabe 2: Listen

### a) (1 Punkt) 

Führen Sie das folgende Programm aus: 

```Python
a = [1,2,3,4]
b = a
a[0] = 0
print(b)
print(a is b)
print(a == b) 
```

Beantworten Sie die folgenden Fragen: 

1. Warum hat die Liste ```b``` am Ende den gleichen Inhalt wie ```a```? 
1. Was ist allgemein die Semantik der Operatoren ```is``` und ```==``` ?
1. Erklären Sie auf dieser Grundlage, wie die Werte von ```a is b``` und ```a == b``` zustande kommen. 

<!-- grade_id: 30001, points: 1, grade: true, solution: true -->
1. Weil beide Variablen a und b Namen für dieselbe Liste (genauer: Namen für dasselbe Listenobjekt) sind
2. ```is``` vergleicht ob zwei Variablen Namen für dasselbe Objekt (hier die in der ersten Zeile eingeführte Liste) sind und ```==``` vergleicht ob die Objekte (hier Listen) die gleichen Werte haben.
3. ```a is b``` ist wahr, weil beide Variablen Namen für dieselbe Liste sind. ```a == b``` ist wahr, weil a und b dann natürlich auch die gleichen Werte enthalten.

### b) (1 Punkt) 

Führen Sie das folgende Programm aus: 

```Python
def list_return(lst):
    return lst  

a = [1, 2, 3, 4]
b = list_return(a)
a[0] = 0
print(a,b)
print(a is b)
print(a == b) 
```

Beantworten Sie die folgenden Fragen: 

1. Obwohl die Variable ```b``` von der Funktion ```list_return``` initialisiert wurde, hat sie den gleichen Inhalt wie die modifizierte Liste ```a```. Erläutern Sie, warum das der Fall ist.
1. Erklären Sie auf dieser Grundlage, wie die Werte von ```a is b``` und ```a == b``` zustandekommen. 

<!-- grade_id: 30011, points: 1, grade: true, solution: true -->
1. Weil die Funktion ```list_return``` genau den Wert zurückgibt, der ihr als Parameter übergeben wurde und bei Aufruf von ```list_return``` der Wert von a übergeben wurde. b ist damit nach der Zuweisung nur ein anderer Name für dieselbe Liste wie die, welche den Namen a hat.
2. Der Grund ist der gleiche wie bei 2a) 3.

### c) (1 Punkt) 

Führen Sie das folgende Programm aus:

```Python
def list_return(lst):
    return lst  
    
a = [1, 2, 3, 4]
b = list_return(a[:])
a[0] = 0
print(a,b)
print(a is b)
print(a == b) 
```

Beantworten Sie folgende Fragen: 

1. Warum wird jetzt der Inhalt von ```b``` nicht durch die Modifizierung der Liste, die durch  ```a```  referenziert wird, geändert? 
1. Erklären Sie auf dieser Grundlage wie die Werte von ```a is b``` und ```a == b``` zustande kommen. 

<!-- grade_id: 30021, points: 1, grade: true, solution: true -->
1. Weil durch den Slicing-Befehl [:] eine (flache) Kopie der vorher definierten Liste erstellt wird und ein Verweis auf diese Kopie an die Funktion ```list_return``` übergeben wird. Daher verweist ```b``` anschließend auf eine andere Liste als ```a```.
2. Die Variablen ```a``` und ```b``` zeigen weder auf dieselbe Liste, noch haben sie die gleichen Werte. Deswegen sind die beide Ausgaben entsprechend nicht wahr.

### d) (1 Punkt) 

Führen Sie das folgende Programm aus: 

```Python
def list_modify(lst):
    lst[2] = lst[0] + lst[1]
    return lst

def list_modify2(lst):
    lst = lst*2
    return lst
    
a = [1, 2, 3, 4]  

c = list_modify(a)
print(a is c)
print(a == c)

d = list_modify2(a)
print(a is d)
print(a == d)
```

Beantworten Sie folgende Fragen: 

1. Erklären Sie die Ergebnisse von ```a is c``` und ```a == c```. 
1. Erklären Sie die Ergebnisse von ```a is d``` und ```a == d```. 

<!-- grade_id: 300009999, points: 1, grade: true, solution: true -->
1. Beide Vergleiche sind wahr, weil a und c auf dieselbe Liste zeigen. Die Funktion ```list_modify``` gibt nämlich einen Verweis auf die als Parameter übergebene Liste zurück, es wird keine Kopie erstellt.
2. Der Befehl ```lst*2``` erzeugt eine neue Liste, die doppelt so lang ist wie die ursprünglich definierte (den Inhalt der ersten Liste zweimal hintereinander enthält) und welche nach Beendigung des Funktionsaufrufs den Namen ```d``` erhält. Daher haben ```a``` und ```d``` weder den gleichen Inhalt, noch zeigen sie auf dieselbe Liste. Daher sind die beiden Vergleiche unwahr.

### e) (1 Punkt) 

Schreiben Sie eine neue Funktion namens ```list_copy_shallow```, welche eine Liste übergeben bekommt und eine *flache* Kopie dieser Liste zurückgibt (zum Unterschied von flacher und tiefer Kopie einer Liste siehe Kapitel 5 der VL-Unterlagen).

Nutzen Sie dazu keine vordefinierte Copy-Funktionalität von Listen.

<!-- grade_id: 30031, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION 
def list_copy_shallow(lst):
    """Create a shallow copy of a list.

    Args:
        lst: the list to be copied
    
    Returns:
        a shallow copy of the input list
    """
    new_list = []
    for item in lst:
        new_list.append(item)
    return new_list
### END SOLUTION 
```

<!-- grade_id: 30032, points: 1, grade: true, solution: false -->
```python
a = [1, 2, 3, 4]
b = list_copy_shallow(a)
b[0] = 2
assert(a != b)
```

### f) (1 Punkt) 

Führen Sie das folgende Programm aus: 

```Python
def f(a, b):
   print("a: ", a)
   print("b: ", b)

def g(L):
   L[0] = L[0] - 1
   return L

def h(L):
   L[0] = L[0] + 2
   return L

L = [5, 17]

f(g(L), h(L))
```
Warum referenzieren die Variablen ```a``` und ```b``` in der Funktion ```f``` dieselben  Werte? Begründen Sie! 

<!-- grade_id: 3000021, points: 1, grade: true, solution: true -->
Die beiden Variablen haben dieselbe Belegung, weil es insgesamt nur eine einzige Liste L gibt.
Zwar wird deren erstes Element in beiden Funktionen verändert, aber beide Funktionen geben dieselbe Liste zurück, somit ist L der aktuelle Parameter sowohl für den formalen Parameter a als auch für den formalen Parameter b.

### g) (1 Punkt) 

Wie muss die Funktion ```f()``` (aus Teilaufgabe f) aufgerufen werden, damit die folgende Ausgabe:

    a:  [4, 17]
    b:  [7, 17]

ausgegeben wird? 
Nutzen Sie dazu die bereits vorhandene Liste
```Python
    L = [5, 17]
```
und verwenden Sie Funktionsaufrufe von ```g()``` und ```h()``` als Parameter für einen Aufruf von ```f()```.

<!-- grade_id: 3000134, points: 1, grade: true, solution: true -->
```python
def f(a, b):
   print("a: ", a)
   print("b: ", b)

def g(L):
   L[0] = L[0] - 1
   print (L)
   return L

def h(L):
   L[0] = L[0] + 2
   print (L)
   return L

L = [5, 17]

### BEGIN SOLUTION
f(g(L[:]), h(L[:]))
### END SOLUTION
```

Output:
```
[4, 17]
[7, 17]
a:  [4, 17]
b:  [7, 17]
```

### h) (1 Punkt) 

Gegegeben sei die folgende 2x2-Matrix:

```
1 2
3 4
```

Betrachten Sie folgende Repräsentation einer solchen Matrix in Python:
```Python
    matrix = [[1, 2], [3, 4]]
```

Schreiben Sie eine Funktion namens ```transpose_matrix```, welche eine in dieser Art implementierte 2x2-Matrix als Eingabe erhält und die transponierte Matrix zurückgibt.
Die Matrix selbst soll dabei nicht verändert werden!

Zum Thema Transponieren von Matrizen lesen Sie nötigenfalls hier nach: https://de.wikipedia.org/wiki/Transponierte_Matrix.

<!-- grade_id: 30061, points: 0, grade: false, solution: true -->
```python
###BEGIN SOLUTION
def transpose_matrix(matrix):
    """Transposes a 2x2 matrix (without changing the original matrix)

    Args:
        matrix: The matrix to be transposed, represented as a list of lists.
            The first list represents the lines and the secont list represents
            the elements of the lines.
    
    Returns:
        The transposed matrix.
    """    
    return [[matrix[0][0], matrix[1][0]], [matrix[0][1], matrix[1][1]]]                
### END SOLUTION
```

<!-- grade_id: 30062, points: 1, grade: true, solution: false -->
```python
matrix = [[1, 2], [3, 4]]
matrix_transposed = transpose_matrix(matrix)

assert(matrix_transposed[0][0] == 1)
assert(matrix_transposed[0][1] == 3)
assert(matrix_transposed[1][0] == 2)
assert(matrix_transposed[1][1] == 4)


```

## Aufgabe 3: Strings

Zeichenketten in Python können Zeichen aus einem großen Zeichenvorrat annehmen; typischerweise dem sog. Universal Coded Character Set Transformation Format mit 8 bits (UTF-8). 
Damit stehen zum Beispiel Zeichen asiatischer und arabischer Sprachen bereit. 

### a) (1 Punkt) 

Finden Sie heraus, wie viele Bytes benötigt werden, um die folgenden Zeichenketten intern darzustellen.
Wie vielen Bytes pro Zeichen entspricht das (achten Sie ggf. auf Leerzeichen!)? Schreiben Sie Ihre Antwort als Kommentar hinter die jeweilige Zeile in der unteren Codezelle.


```Python
# Beispiel-Zeichenketten in UTF-8
s1 = '我们 爱 编程'
s2 = "äöü"
s3 = "éàß"
s4 = "Мы любим GP1"
s5 = "ما عاشق برنامه نویسی"
```

Schreiben Sie eine Funktion namens ```utf8_len```, die einen String übergeben bekommt und  die Anzahl der Bytes dieses Strings zurückgibt. 

Hinweis: Sie dürfen hierzu Funktionen der Python-Standardbibliothek nutzen, die nicht in der Vorlesung behandelt wurden. Dazu müssen Sie etwas recherchieren! 

<!-- grade_id: 3032, points: 0, grade: false, solution: true -->
```python
s1 = '我们 爱 编程' # Bytes: 17
s2 = "äöü" # Bytes: 6
s3 = "éàß" # Bytes: 6
s4 = "Мы любим GP1" # Bytes: 19
s5 = "ما عاشق برنامه نویسی" # Bytes: 37

### BEGIN SOLUTION
# Leerzeichen benötigen 1 Byte, Umlaute, das ß und Zeichen mit Akzent 2 Bytes und die asiatischen Bildzeichen 3 Bytes
def utf8_len(s):
    """Returns the size (in bytes) of a string.

    Args:
        s: A string.
    
    Returns:
        The size of s (in bytes)
    """    
    return len(bytes(s, "utf-8"))
### END SOLUTION
```

<!-- grade_id: 3033, points: 1, grade: true, solution: false -->
```python
s1 = '我们 爱 编程'
s2 = "äöü"
s3 = "éàß"
s4 = "Мы любим GP1"
s5 = "ما عاشق برنامه نویسی"

assert(utf8_len(s1) == 17)
assert(utf8_len(s2) == 6)
assert(utf8_len(s3) == 6)
assert(utf8_len(s4) == 19)
assert(utf8_len(s5) == 37)
```

### b) (1 Punkt) 

Wäre es sinnvoll, bei solchen Zeichenketten stets die gleiche Anzahl an Bytes pro Zeichen zu benutzen? Begründen Sie! 

Welches Verhalten scheint bei Python vorzuliegen?

<!-- grade_id: 3037, points: 1, grade: true, solution: true -->
Argumente für die Verwendung einer stets gleichen Anzahl an Bytes pro Zeichen:
- Leichtes Errechnen der Länge des Strings möglich (evtl. für Pointer-Arithmetik nützlich, für Python egal) 
- Verändern eines Zeichens ändert nicht die Länge des Strings (in Python ja nicht möglich)

Argumente für die Verwendung einer variablen Anzahl an Bytes pro Zeichen:
- Geringerer Speicherplatzbedarf bei Verwendung von Sprachen, die nur oder größtenteils aus 1-Byte-Zeichen bestehen
- Texte in Englisch lassen sich wie ASCII kodieren (in manchen Fällen höhere Kompatibilität)
- Theoretisch unbegrenzte Erweiterbarkeit des Zeichenvorrats

Python nutzt, wie in a) gesehen, eine variable Anzahl an Bytes pro Zeichen.

## Aufgabe 4: Tuples

### a) (2 Punkte) 

Schreiben Sie eine Funktion namens ```order_check```, welche das verschachtelte Tupel ```bestellung``` (siehe unten) übergeben bekommt und die gewünschte Bestellung ausgibt. Z.B.: 

```text
Thomas möchte 3 Äpfel
Thomas möchte 1 Birnen
Anna möchte 6 Birnen
```

(Auf sprachliche korrekte Unterscheidung zwischen Singular und Plural brauchen Sie hier, wie schon im Beispiel, nicht einzugehen.)

Rufen Sie die Funktion einmal mit dem Tupel ```bestellung``` als Parameter auf.

<!-- grade_id: 40002, points: 2, grade: true, solution: true -->
```python
bestellung = (('Thomas', (('Äpfel', 3), ('Birnen', 1))),
          ('Anna', (('Birnen', 6),)))


def order_check(orders):
### BEGIN SOLUTION    
    """Prints order list.

    Args:
        orders: The list of orders (according to the task description).
    
    Returns:
        None
    """
    for user, order in orders:
        for fruit, quantity in order:
            print("%s möchte %d %s" % (user, quantity, fruit))

order_check(bestellung)            
### END SOLUTION
```

Output:
```
Thomas möchte 3 Äpfel
Thomas möchte 1 Birnen
Anna möchte 6 Birnen
```

<!-- grade_id: 400002b, points: 1, grade: true, solution: false -->
```python
assert(order_check.__doc__ != None)
```

### b) (3 Punkte) 

Erweitern Sie die Funktion ```order_check```, so dass sie eine ```dict```-Variable mit der Frucht als Schlüsselwort und der entsprechenden Gesamtmenge (aufaddiert über alle Besteller) als Wert zurückgibt. 
Beispielsweise soll die Belegung von ```fruechte_menge``` nach Ausführung der folgenden Zuweisung:
```Python
fruechte_menge = order_check(bestellung)
```
so sein:
```Python    
{'Äpfel': 3, 'Birnen': 7}
```

(Die Reihenfolge der Elemente in dem ```dict``` ist dabei nicht entscheidend; wie Sie wissen, sind ```dict```-Strukturen nicht reihenfolgetreu.)

<!-- grade_id: 400004, points: 0, grade: false, solution: true -->
```python
bestellung = (('Thomas', (('Äpfel', 3), ('Birnen', 1))),
          ('Anna', (('Birnen', 6),)))

def order_check(orders):
### BEGIN SOLUTION
    """Print order list and return total fruit quantity

    Args:
        orders: The orders (represented as a tuple).
    
    Returns:
        A dict containing the fruit quantities.
    """
    dict_fruits = {}
    for user, order in orders:
        for fruit, quantity in order:            
            if fruit in dict_fruits:
                dict_fruits.update({fruit: dict_fruits[fruit] + quantity})   
                # Alternativ:  dict_fruits[fruit] += quantity
            else:
                dict_fruits.update({fruit:quantity}) 
                # Alternativ: dict_fruits[fruit] = quantity
         
    return dict_fruits

fruechte_menge = order_check(bestellung)
print(fruechte_menge)
### END SOLUTION
```

Output:
```
{'Birnen': 7, 'Äpfel': 3}
```

<!-- grade_id: 400005, points: 2, grade: true, solution: false -->
```python
assert(fruechte_menge=={'Äpfel': 3, 'Birnen': 7})
```

<!-- grade_id: 400006, points: 1, grade: true, solution: false -->
```python
assert(order_check.__doc__ != None)
```

# Aufgabe 5: Docstring

(2 Punkte) Dokumentieren Sie die Funktion ```multiples``` in Docstring gemäß Google-Konvention. 
Beschreiben Sie die Parameter und Rückgabewerte und was die Funktion leistet.

```Python
def multiples(f1, f2, max_val):
    a = []
    for number in range(1, max_val+1):
        if not number % f1 or not number % f2:
            a.append(number)
    return a
    
b= multiples(3,5,1000)
print(b)
```

<!-- grade_id: 500002, points: 0, grade: false, solution: true -->
```python
def multiples(f1, f2, max_val):
### BEGIN SOLUTION
    """Computes all multiples of f1 and f2 up to max_val and returns them as a list

    Args:
        f1: the first factor (int)
        f2: the second factor (int)
        max_val: the maximum value
    
    Returns:
        list containing all the multiples of f1 and f2
    """
### END SOLUTION
    a = []
    for number in range(1, max_val+1):
        if not number % f1 or not number % f2:
            a.append(number)
    return a
```

<!-- grade_id: 500003, points: 2, grade: true, solution: false -->
```python
assert(multiples.__doc__ != None)
```