<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile). Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue07.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen!). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 18.12.2016 um 23:59 Uhr

# Aufgabe 1: String-API
Diese Aufgabe lässt sich eventuell einfacher mithilfe der Python-Dokumentation lösen:
https://docs.python.org/3/library/stdtypes.html.

## a) 
Schreiben Sie eine Funktion `s_verbinder(l,z)`, welches eine Liste mit Strings `l` zu einem einzigen String verbindet, wobei zwischen die einzelnen Listenelemente jeweils das Zeichen `z` eingefügt werden soll.

<!-- grade_id: string_sol_1, points: 0, grade: false, solution: true -->
```python
def s_verbinder(l,z):
    # Your Code here
    ### BEGIN SOLUTION
    """Methode, welche eine Liste mit Strings zu einem einzigen String verbindet."""
    return z.join(l)
    ### END SOLUTION
```

<!-- grade_id: string_test_1, points: 1, grade: true, solution: false -->
```python
assert(s_verbinder(['1','2','3','4'], '-') == '1-2-3-4')
### HIDESTART
assert(s_verbinder(['a','b','c','d'], '-') == 'a-b-c-d')
### HIDEEND
```

## b) 
Schreiben Sie eine Funktion `s_suche(s1,s2)`, welche den niedrigsten Index ausgibt, an welcher der String `s2` in `s1` auftaucht.

<!-- grade_id: string_sol_2, points: 0, grade: false, solution: true -->
```python
def s_suche(s1,s2):
    # Your Code here
    ### BEGIN SOLUTION
    """Methode, welche den Niedrigsten Index ausgibt, an welchem der String s2 in s1 auftaucht."""
    return s1.find(s2)
    ### END SOLUTION
```

<!-- grade_id: string_test_2, points: 1, grade: true, solution: false -->
```python
assert(s_suche('Grundlagen der Programmierung', 'der') == 11)
assert(s_suche('Grundlagen der Programmierung', 'P') == 15)
### HIDESTART
assert(s_suche('Test', 'T') == 0)
assert(s_suche('Test', 'e') == 1)
assert(s_suche('Test', 's') == 2)
assert(s_suche('Test', 't') == 3)
assert(s_suche('Test', 'es') == 1)
assert(s_suche('TestT', 'T') == 0)
### HIDEEND
```

## c) 
Schreiben Sie eine Funktion `split_3(s)`, welche einen string `s` an dem ersten Index `i` aufteilt, an welcher ein Punkt ('.') vorkommt. Es soll ein Triple `(s_links, i, s_rechts)` zurückgegeben werden, wobei `s_links` der Teilstring von `s` ist, welcher alle Zeichen links von `s[i]` enthält und `s_rechts`der Teilstring von `s` ist, welcher alle Zeichen rechts von `s[i]` enthält.

<!-- grade_id: string_sol_3, points: 0, grade: false, solution: true -->
```python
def split_3(s):
    # Your Code here
    ### BEGIN SOLUTION
    """ Methode welche welche einen string s an dem ersten Index i aufteilt, an welcher ein Punkt ('.') vorkommt."""
    l = s.partition('.')
    sol = (l[0], s.index('.'), l[2])
    return sol
    ### END SOLUTION
```

<!-- grade_id: string_test_3, points: 1, grade: true, solution: false -->
```python
assert(split_3('Dies ist ein Satz. Dies ist noch ein Satz.') == ('Dies ist ein Satz', 17, ' Dies ist noch ein Satz.'))
### HIDESTART
assert(split_3('Dies ist ein geheimer Satz. Dies ist noch ein geheimer Satz.') == ('Dies ist ein geheimer Satz', 26, ' Dies ist noch ein geheimer Satz.'))
### HIDEEND
```

# Aufgabe 2: Objektorientierung

Betrachten Sie für diese Aufgabe die gegebene Klasse `Patient`.

```python
class Patient():
    anz_patienten = 0
    
    def __init__(self, name, alter, privat_patient):
        self.name = name
        self.alter = alter
        self.privat_patient = privat_patient
        
    def get_name(self):
        return self.name
    
    def get_alter(s):
        return s.alter
    
    def is_privat_patient(self):
        return self.privat_patient
    
    def __eq__(self, p):
        return self.name == p.name and self.alter == p.alter
```

## a)
Erläutern Sie zunächst den Unterschied der Operationen `==` und `is`. 
Weisen Sie dann den beiden gegebenen Variablen `p1` und `p2` Objekte der Klasse `Patient` zu, sodass `p1 == p2` gilt, jedoch nicht `p1 is p2`.

<!-- grade_id: patient_sol_1_1, points: 1, grade: true, solution: true -->
Der Operator `==` führt die Methode `__eq__` des Objektes aus, welches auf der linken Seite des Operators steht.
Der Operator `is` hingegen überprüft, ob die referenzierten Objekte auf beiden Seiten dieselben sind.

<!-- grade_id: patient_sol_1_2, points: 0, grade: false, solution: true -->
```python
p1 = None # None durch entsprechende Belegung ersetzen!
p2 = None # None durch entsprechende Belegung ersetzen!
### BEGIN SOLUTION
p1 = Patient("Max M", 26, False)
p2 = Patient ("Max M", 26, False)
### END SOLUTION
```

<!-- grade_id: patient_test_1, points: 1, grade: true, solution: false -->
```python
### HIDESTART
assert(isinstance(p1, Patient))
assert(isinstance(p2, Patient))
assert(p1 == p2)
assert(not p2 is p1)
### HIDEEND
```

## b)
Wie Sie sehen, heißt der erste Parameter der Methode `get_alter()` nicht `self`.
Würde das folgende Programm einen Fehler erzeugen?
Wenn ja, an welcher Stelle und warum, wenn nein, erklären Sie warum nicht?

```Python
p1 = Patient("Max Meier", 26, False)
print(p1.get_alter())
```

*Hinweis*: Ein simples Copy-Paste der Fehlermeldung reicht für eine Beschreibung des Fehlers nicht aus.

<!-- grade_id: patient_sol_3, points: 1, grade: true, solution: true -->
Das Programm funktioniert, da der Name für den ersten Parameter einer Objektmethode nur per Konvention `self` lautet.

## c)
Angenommen, die obige Klasse `Patient` würde über keinen Konstruktor verfügen.
Würde das folgende Programm einen Fehler erzeugen?
Wenn ja, an welcher Stelle und warum, wenn nein, erklären Sie warum nicht?   
```Python
p1 = Patient()
print(p1.get_name())
```
*Hinweis*: Ein simples Copy-Paste der Fehlermeldung reicht für eine Beschreibung des Fehlers nicht aus.

<!-- grade_id: patient_sol_4, points: 1, grade: true, solution: true -->
Zunächst einmal wird der fehlende Konstruktor von Python durch einen leere Standardkonstruktor ersetzt.
Das Programm wirft einen Fehler wenn versucht wird die Methode p1.get_name() aufzurufen, da der Patient über keine Objektvariable `name` verfügt (diese wurde im Standardkonstruktor nicht definiert).

# Aufgabe 3: Vektoren
In dieser Aufgabe sollen Sie eine Klasse `Vector` implementieren.
Diese soll mathematische Vektoren repräsentieren.
Ein Vektor $v$ besteht aus $n$ einzelnen Zahlen, also $v = (v_0,\ldots,v_{n-1})$.
Die Dimension $n$ kann hier eine beliebige Zahl aus den natürlichen Zahlen $\mathbb N = \{1,2,3,\ldots\}$ sein.
Die Klasse `Vector` soll für ein Objekt $v = (v_0,\ldots,v_{n-1})$ die folgenden Methoden bereitstellen:
- Einen Konstruktor, welcher über optionale Parameter die Dimension `n` und eine Liste von Integers `vec` verfügt. Falls beide Parameter angegeben sind, können Sie davon ausgehen, dass `n == len(vec)` gilt.
- `__setitem__(i, val)`: Ersetzt das Element $v_i$ des Vektors $v$ mit dem Wert `val`.
- `__getitem__(i)`: Gibt das Element $v_i$ des Vektors $v$ zurück. Falls `i` größer oder gleich der Dimension `n` des Vektors ist, so soll `None` zurückgegeben werden.

Sobald Sie die Methoden `__setitem__(i, val)` und `__getitem__(i)` implementiert haben, sollten Sie auf die einzelnen Vektorelemente mittels Indizierung zugreifen können (siehe Tests).

- `add(a)`: Addiert den Vektor $v = (v_0,\ldots,v_{n-1})$ mit dem übergebenen Vektor $a = (a_0,\ldots,a_{n-1})$ und liefert das Ergebnis als neuen Vektor $e = (v_0+a_0,\ldots,v_{n-1}+a_{n-1})$ zurück. Die beiden Vektroren `v` und `a` sollen dabei nicht verändert werden.
- `subtract(a)`: Subtrahiert den Vektor $a = (a_0,\ldots,a_{n-1})$ von dem Vektor $v = (v_0,\ldots,v_{n-1})$ und liefert das Ergebnis als neuen Vektor $s = (v_0-a_0,\ldots,v_{n-1}-a_{n-1})$ zurück. Die beiden Vektroren `v` und `a` sollen dabei nicht verändert werden.

Für die Methoden `add`, `subtract` und `dot` können Sie davon ausgehen, dass beide Vektoren immer dieselbe Länge haben.

- `dot(w)`: Berechnet das Vektorprodukt für den Vektor $v = (v_0,\ldots,v_{n-1})$ und dem übergebenen Vektor $w = (w_0,\ldots,w_{n-1})$ und liefert das Ergebnis ($\sum_{i=0}^{n-1} v_i\cdot w_i$) zurück.
- `norm()`: Berechnet die Norm ($\sqrt[]{\sum_{i=0}^{n-1} v_i^2}$) des Vektors $v = (v_0,\ldots,v_{n-1})$ und liefert diese zurück.

Es ist ihnen selbstverständlich erlaubt zusätzliche Hilfsmethoden wie z.B. `__str__` zu implementieren, um Ihren Code besser auf Fehler überprüfen zu können.

<!-- grade_id: vector_sol, points: 0, grade: false, solution: true -->
```python
class Vector:
    # Your code here
    ### BEGIN SOLUTION
    """Die Klasse Vektor realisiert einen mathematischen Vektor, auf welchem Vektoroperationen ausgeführt werden können."""
    
    def __init__(self, n=None, vec=None):
        """Erstelle leeren Vektor der Länge n. Falls Initiatiliserungsvektor oder -liste
        gegeben, kopiere dessen Inhalt."""

        self.vec = []
        if n:
            self.vec = [0] * n
        if vec:
            self.vec = [0] * len(vec)
            for i in range(len(vec)):
                self.vec[i] = vec[i]

    def add(self, v):
        """Addiere mit self, aber self nicht verändern"""
        n = len(self.vec)
        result = Vector(n)
        for i in range(n):
            result[i] = self[i] + v[i]

        return result     
    
    def subtract(self, v):
        """Subtrahiere von self, aber self nicht verändern"""
        erg = [0]*len(self)
        for i in range(0, len(v)):
            erg[i] = self.vec[i] - v[i]
        return Vector(vec = erg)
    
    def dot(self, v):
        """Berechnet das Vektorprodukt mit dem übergebenen Vektor"""
        return sum([x*o for x, o  in zip(self.vec, v.vec)])
    
    def norm(self):
        """Berechnet die Norm für den Vektor"""
        return sum([x**2 for x in self.vec])**0.5
    
    def __len__(self):
        return len(self.vec)
    
    def __setitem__(self, i, val):
        """Ersetzt das i-te Element des Vektors"""
        self.vec[i] = val
        
    def __getitem__(self, i):
        """Liefert das i-te Element des Vektors"""
        if i >= len(self.vec):
            return None
        else:
            return self.vec[i]
    ### END SOLUTION
```

<!-- grade_id: vector_test_0, points: 2, grade: true, solution: false -->
```python
# Konstruktor-Test:
v0 = Vector(n=2)
assert(v0[0] == 0) # Zugriff hier mittels Indizierung!
assert(v0[1] == 0)

v1 = Vector(n=2, vec=[1,2])
assert(v1[0] == 1)
assert(v1[1] == 2)

v2 = Vector(vec=[1,2])
assert(v2[0] == 1)
assert(v2[1] == 2)

# Ersetze ein Element von v2:
v2[0] = 2
assert(v2[0] == 2)

### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    for i in range(0,n):
        assert(a[i] == i+1)
    a[0] = 3
    assert(a[0] == 3)
### HIDEEND
```

<!-- grade_id: vector_test_1, points: 1, grade: true, solution: false -->
```python
a = Vector(vec = [1, 2, 3])
b = Vector(vec = [3, 4, 5])

e = a.add(b)
assert(e[0] == 4)
assert(e[1] == 6)
assert(e[2] == 8)

### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    b = Vector(vec = range(0,n))

    e = a.add(b)

    for i in range(0,n):
        assert(e[i] == (a[i] + b[i]))
### HIDEEND
```

<!-- grade_id: vector_test_2, points: 1, grade: true, solution: false -->
```python
a = Vector(vec = [1, 2, 3])
b = Vector(vec = [3, 4, 5])

s = a.subtract(b)
assert(s[0] == -2)
assert(s[1] == -2)
assert(s[2] == -2)
### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    b = Vector(vec = range(0,n))

    s = a.subtract(b)

    for i in range(0,n):
        assert(s[i] == (a[i] - b[i]))
### HIDEEND
```

<!-- grade_id: vector_test_3, points: 1, grade: true, solution: false -->
```python
a = Vector(vec = [1, 2, 3])
b = Vector(vec = [3, 4, 5])
dot = a.dot(b)
assert(dot == 26)
### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    b = Vector(vec = range(0,n))
    dot = 0
    for i in range(0,n):
        dot = dot + a[i] * b[i]
    assert(a.dot(b) == dot)
### HIDEEND
```

<!-- grade_id: vector_test_4, points: 1, grade: true, solution: false -->
```python
a = Vector(vec  = [1, 2, 3])
norm = a.norm()
assert(norm == (14**0.5))
### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    norm = 0
    for i in range(0, n):
        norm = norm + a[i] ** 2
    norm = norm ** 0.5
    assert(a.norm() == norm)    
### HIDEEND
```

# Aufgabe 4: PAUL
In dieser Aufgabe sollen Sie eine einfache Version von PAUL implementieren.
Das System soll aus 3 Klassen bestehen: `Student`, `Vorlesung` und `Paul`.
Die Klasse `Student` ist Ihnen bereits gegeben (siehe untere Codezelle).
Ihre Aufgabe ist nun die Klassen `Vorlesung` und `Paul` zu vervollständigen.
Die Anforderungen an die Klassen, sowie Tests sind Ihnen in den jeweiligen Teilaufgaben gegeben.

```python
class Student:
    """Diese Klasse modelliert einen Studenten mit Attributen matrikelnummer und name."""

    def __init__(self, mat_nr, name):
        """Konstruktor

        Args:
            mat_nr: Matrikelnummer des Studenten (int)
            name: Name des Studenten (str)
        """
        self.mat_nr = mat_nr
        self.name = name
    
    def get_mat_nr(self):
        """Get-Methode für das Matrikelnummer-Attribut

        Returns:
            int Matrikelnummer des Studenten
        """
        return self.mat_nr
    
    def get_name(self):
        """Get-Methode für das Name-Attribut

        Returns:
            str Namen des Studenten
        """
        return self.name
    
    def __str__(self):
        """String-Repräsentation eines Studenten

        Returns:
            str Matrikelnummer und name des Studenten
        """
        return str(self.mat_nr) + " " + self.name
```

## a)
Erstellen Sie eine Klasse `Vorlesung`, welche über die folgenden Methoden verfügt:

- Einen Kontruktor, welcher die ID der Vorlesung (int), sowie den Namen der Vorlesung (str) übergeben bekommt.
- `get_id`: Gibt die ID der Vorlesung zurück
- `get_name`: Gibt den Namen der Vorlesung zurück
- `add_student`: Registriert den Studenten mit der übergebenen Matrikelnummer in der Vorlesung
- `has_student`: Liefert `True`, falls sich ein Student mit der übergebenen Matrikelnummer in der Vorlesung registriert hat, ansonsten `False`.
- `add_grade`: Setzt die Note des Studenten mit der übergebenen Matrikelnummer auf die übergebene Note. Sollte der Student bereits eine Note haben, so wird diese überschrieben. Wenn der Student nicht in der Vorlesung registriert ist, soll keine Aktion ausgeführt werden.
- `get_grade`: Liefert die Note für diese Vorlesung, welche der Student mit der übergebenen Matrikelnummer erhalten hat. Falls die Matrikelnummer nicht in der Vorlesung registriert ist, soll `-1` zurückgegeben werden. Ist der Student zwar registriert, hat aber noch keine Note erhalten, wird `None` zurückgegeben.


<!-- grade_id: paul_sol_1, points: 0, grade: false, solution: true -->
```python
class Vorlesung:
    # Your code here
    ### BEGIN SOLUTION
    """Eine Klasse mit Daten für eine Vorlesung"""
    
    def __init__(self, v_id, name):
        """Konstruktor

        Args:
            v_id: ID der Vorlesung (int)
            name: Name der Vorlesung (str)
        """
        self.v_id = v_id
        self.name = name
        self.studenten = {}
        
    def get_id(self):
        """Get-Methode für das v_id-Attribut

        Returns:
            int ID der Vorlesung
        """
        return self.v_id
    
    def get_name(self):
        """Get-Methode für das name-Attribut

        Returns:
            str Name der Vorlesung
        """
        return self.name
    
    def add_student(self, mat_nr):
        """Registriert den Studenten mit der übergebenen Matrikelnummer in der Vorlesung
        
        Args:
            mat_nr: Matrikelnummer des zu registrierenden Studenten (int)
        """
        self.studenten[mat_nr] = None
        
    def has_student(self, mat_nr):
        """Liefert True, falls sich ein Student mit der übergebenen Matrikelnummer in der Vorlesung registriert hat, 
        ansonsten False.
        
        Args:
            mat_nr: Matrikelnummer des zu suchenden Studenten (int)
            
        Returns:
            bool Suchergebnis als Wahrheitswert
        """
        if mat_nr in self.studenten:
            return True
        return False
    
    def add_grade(self, mat_nr, grade):
        """Setzt die Note des Studenten mit der übergebenen Matrikelnummer auf die übergebene Note. 
        Sollte der Student bereits eine Note haben, so wird diese überschrieben. Wenn der Student nicht in der Vorlesung 
        registriert ist, soll keine Aktion ausgeführt werden.
        
        Args:
            mat_nr: Matrikelnummer des zu registrierenden Studenten (int)
            grade: Note, welche für den Studenten eingetragen werden soll (float)
        """
        if mat_nr in self.studenten:
            self.studenten[mat_nr] = grade
    
    def get_grade(self, mat_nr):
        """Liefert die Note für diese Vorlesung, welche der Student mit der übergebenen Matrikelnummer erhalten hat. 
        Falls die Matrikelnummer nicht in der Vorlesung registriert ist, soll -1 zurückgegeben werden. Ist der Student 
        zwar registriert, hat aber noch keine Note erhalten, wird None zurückgegeben.
        
        Returns:
            float Note des Studenten, None oder -1
        """
        if mat_nr in self.studenten:
            return self.studenten[mat_nr]
        return -1
    ### END SOLUTION
```

<!-- grade_id: paul_test_1_1, points: 1, grade: true, solution: false -->
```python
# Erstelle Studenten:
s1 = Student(9863451, "Ralf B")
s2 = Student(9863452, "Jessica W")
s3 = Student(9863453, "Philipp K")
s4 = Student(9863454, "Andrea B")
s5 = Student(9863455, "Sebastian A")
s6 = Student(9863456, "Luca B")
studenten = [s1,s2,s3,s4,s5,s6]

# Erstelle Vorlesung:
gp1 = Vorlesung(1, "Grundlagen der Programmierung 1")
assert(gp1.get_id() == 1)
assert(gp1.get_name() == "Grundlagen der Programmierung 1")

# Füge Studenten zur Vorlesung hinzu:
for s in studenten:
    gp1.add_student(s.get_mat_nr())
for s in studenten:
    assert(gp1.has_student(s.get_mat_nr()))

### HIDESTART
# Erstelle Vorlesung:
mod = Vorlesung(2, "Modellierung")
assert(mod.get_id() == 2)
assert(mod.get_name() == "Modellierung")

# Füge Studenten zur Vorlesung hinzu:
mod.add_student(s1.get_mat_nr())
mod.add_student(s2.get_mat_nr())
mod.add_student(s3.get_mat_nr())

assert(mod.has_student(s1.get_mat_nr()))
assert(mod.has_student(s2.get_mat_nr()))
assert(mod.has_student(s3.get_mat_nr()))
assert(not mod.has_student(s4.get_mat_nr()))
assert(not mod.has_student(s5.get_mat_nr()))
assert(not mod.has_student(s6.get_mat_nr()))
### HIDEEND
```

<!-- grade_id: paul_test_1_2, points: 2, grade: true, solution: false -->
```python
# Verteile Noten:
gp1.add_grade(s1.get_mat_nr(), 1.0)
gp1.add_grade(s2.get_mat_nr(), 1.3)
gp1.add_grade(s3.get_mat_nr(), 3.0)
gp1.add_grade(s4.get_mat_nr(), 2.7)
gp1.add_grade(s5.get_mat_nr(), 1.3)
gp1.add_grade(s6.get_mat_nr(), 5.0)

assert(gp1.get_grade(s1.get_mat_nr()) == 1.0)
assert(gp1.get_grade(s2.get_mat_nr()) == 1.3)
assert(gp1.get_grade(s3.get_mat_nr()) == 3.0)
assert(gp1.get_grade(s4.get_mat_nr()) == 2.7)
assert(gp1.get_grade(s5.get_mat_nr()) == 1.3)
assert(gp1.get_grade(s6.get_mat_nr()) == 5.0)

# Ein nicht in der Vorlesung registrierter Student will eine Note haben:
gp1.add_grade(1234567, 1.0) 
assert(gp1.get_grade(1234567) == -1) # Matrikelnummer nicht im Kurs vorhanden

### HIDESTART
assert(mod.get_grade(s4.get_mat_nr()) == -1)
assert(mod.get_grade(s5.get_mat_nr()) == -1)
assert(mod.get_grade(s6.get_mat_nr()) == -1)
### HIDEEND
```

## b)
Ergänzen Sie nun die Klasse `PAUL`, welche über die folgenden Methoden verfügt:
- `add_vorlesung`: Erstellt eine neue Vorlesung mit übergebener ID und Namen im System. Sie können davon ausgehen, dass `add_vorlesung` niemals mehrfach für dieselbe Vorlesungs-ID aufgerufen wird.
- `add_student`: Erstellt einen neuen Studenten mit übergebener Matrikelnummer und Namen im System. Sie können davon ausgehen, dass `add_student` niemals mehrfach für dieselbe Matrikelnummer aufgerufen wird.
- `registriere_in_vorlesung`: Fügt den Studenten mit der übergebenen Matrikelnummer der Vorlesung mit der übergebenen ID hinzu. Sie können davon ausgehen, dass `registriere_student_in_vorlesung` niemals mehrfach für ein Paar (`matr_nr`, `v_id`) aufgerufen wird.
- `get_studenten_in_vorlesung`: Liefert eine Liste, welche alle Matrikelnummern enthält, die in der Vorlesung mit der übergebenen ID registriert sind. Falls das System keine Vorlesung mit der übergebenen ID enthält, soll `None` zurückgegeben werden.
- `add_grade`: Fügt dem Studenten mit der übergebenen Matrikelnummer eine Note für die Vorlesung mit ID `v_id` hinzu.
- `get_grades_for_student`: Liefert ein Dictionary mit allen Vorlesungen, in welcher der übergebene Student registriert ist und die entsprechenden Noten, die der Student für die Vorlesung erhalten hat.

Eine beispielhafte Anwendung können Sie an den gegebenen Tests ablesen.

<!-- grade_id: paul_sol_2, points: 0, grade: false, solution: true -->
```python
class Paul:
    # Your code here
    ### BEGIN SOLUTION
    """ Modelliert eine Art Benutzerschnittstelle zur Erstellung bzw. Verwaltung von Studenten und Vorlesungen """
    
    def __init__(self):
        """ Konstruktor """
        self.studenten = []
        self.vorlesungen = {}
    
    def add_vorlesung(self, v_id, name):
        """Erstellt eine neue Vorlesung mit übergebener ID und Namen im System. Sie können davon ausgehen, dass 
        add_vorlesung niemals mehrfach für dieselbe Vorlesungs-ID aufgerufen wird.

        Args:
            v_id: ID der Vorlesung (int)
            name: Name der Vorlesung (str)
        """
        v = Vorlesung(v_id, name)
        self.vorlesungen[v_id] = v
        
    def add_student(self, mat_nr, name):
        """Erstellt einen neuen Studenten mit übergebener Matrikelnummer und Namen im System. Sie können davon ausgehen, dass 
        add_student niemals mehrfach für dieselbe Matrikelnummer aufgerufen wird.

        Args:
            mat_nr: Matrikelnummer des Studenten (int)
            name: Name des Studenten (str)
        """
        s = Student(mat_nr, name)
        self.studenten.append(s)
        
    def registriere_in_vorlesung(self, mat_nr, v_id):
        """Fügt den Studenten mit der übergebenen Matrikelnummer der Vorlesung mit der übergebenen ID hinzu. 
        Sie können davon ausgehen, dass registriere_student_in_vorlesung niemals mehrfach für ein Paar (matr_nr, v_id) 
        aufgerufen wird.

        Args:
            mat_nr: Matrikelnummer des zu registrierenden Studenten (int)
            name: ID der Vorlesung, in welcher der Student registriert werden soll (str)
        """
        self.vorlesungen[v_id].add_student(mat_nr)
        
    def get_studenten_in_vorlesung(self, v_id):
        """Liefert eine Liste, welche alle Matrikelnummern enthält, die in der Vorlesung mit der übergebenen ID 
        registriert sind. Falls das System keine Vorlesung mit der übergebenen ID enthält, soll None zurückgegeben werden.

        Args:
            v_id: ID der Vorlesung (int)
        
        Returns:
            list Matrikelnummern aller Studenten, die in der Vorlesung registriert sind
        """
        if v_id in self.vorlesungen:
            L = []
            for i in self.vorlesungen[v_id].studenten.keys():
                L.append(i)
            return L
        return None
    
    def add_grade(self, mat_nr, grade, v_id):
        """ Fügt dem Studenten mit der übergebenen Matrikelnummer eine Note für die Vorlesung mit ID v_id hinzu.
        
        Args:
            mat_nr: Matrikelnummer des Studenten (int)
            grade: Note des Studenten für die Vorlesung (float)
            v_id: ID der Vorlesung, für welche eine Note verteilt werden soll
        """
        self.vorlesungen[v_id].add_grade(mat_nr, grade)
        
    def get_grades_for_student(self, mat_nr):
        """Liefert ein Dictionary mit allen Vorlesungen, in welcher der übergebene Student registriert ist und 
        die entsprechenden Noten, die der Student für die Vorlesung erhalten hat.
        
        Args:
            mat_nr: Matrikelnummer des Studenten (int)
        
        Returns:
            dict Dictionary mit allen Noten für den gegebenen Studenten
        """
        d = {}
        for v_id in self.vorlesungen:
            v = self.vorlesungen[v_id]
            if v.has_student(mat_nr):
                d[v.get_name()] = v.get_grade(mat_nr)
        return d
    ### END SOLUTION
```

<!-- grade_id: paul_test_2_1, points: 1, grade: true, solution: false -->
```python
# Erstelle System:
paul = Paul()

# Füge Vorlesungen und Studenten hinzu:
paul.add_vorlesung(1, "Grundlagen der Programmierung 1")
paul.add_vorlesung(2, "Modellierung")
paul.add_student(9863451, "Ralf B")
paul.add_student(9863452, "Jessica W")
paul.add_student(9863453, "Philipp K")

# Registriere Studenten in Vorlesung:
paul.registriere_in_vorlesung(9863451, 1)
paul.registriere_in_vorlesung(9863452, 1)
paul.registriere_in_vorlesung(9863453, 1)

paul.registriere_in_vorlesung(9863451, 2)
paul.registriere_in_vorlesung(9863452, 2)
paul.registriere_in_vorlesung(9863453, 2)

assert(paul.get_studenten_in_vorlesung(1) == [9863451, 9863452, 9863453])
assert(paul.get_studenten_in_vorlesung(2) == [9863451, 9863452, 9863453])

### HIDESTART
paul.add_vorlesung(3, "Datenstrukturen und Algorithmen")
paul.add_student(9863454, "Andrea B")
paul.registriere_in_vorlesung(9863454, 3)
assert(paul.get_studenten_in_vorlesung(3) == [9863454])
### HIDEEND
```

<!-- grade_id: paul_test_2_2, points: 2, grade: true, solution: false -->
```python
# Füge Noten hinzu:
paul.add_grade(9863451, 3.0, 1)
paul.add_grade(9863452, 1.3, 1)
paul.add_grade(9863453, 2.0, 1)

paul.add_grade(9863451, 5.0, 2)
paul.add_grade(9863452, 2.3, 2)
assert(paul.get_grades_for_student(9863451) == {'Grundlagen der Programmierung 1': 3.0, 'Modellierung': 5.0})
assert(paul.get_grades_for_student(9863452) == {'Grundlagen der Programmierung 1': 1.3, 'Modellierung': 2.3})
assert(paul.get_grades_for_student(9863453) == {'Grundlagen der Programmierung 1': 2.0, 'Modellierung': None})

### HIDESTART
paul.add_grade(9863454, 2.0, 3)
assert(paul.get_grades_for_student(9863454) == {'Datenstrukturen und Algorithmen': 2.0})
### HIDEEND
```

<!-- grade_id: doc_test, points: 1, grade: true, solution: false -->
```python
# Tests für das Vorhandensein der Dokumentation

# Aufgabe 1
assert(s_verbinder.__doc__ != None)
assert(s_suche.__doc__ != None)
assert(split_3.__doc__ != None)

# Aufgabe 3
assert(Vector.__doc__ != None)
assert(Vector.__setitem__.__doc__ != None)
assert(Vector.__getitem__.__doc__ != None)
assert(Vector.add.__doc__ != None)
assert(Vector.subtract.__doc__ != None)
assert(Vector.dot.__doc__ != None)
assert(Vector.norm.__doc__ != None)

# Aufgabe 4
assert(Vorlesung.__doc__ != None)
assert(Vorlesung.get_id.__doc__ != None)
assert(Vorlesung.get_name.__doc__ != None)
assert(Vorlesung.add_student.__doc__ != None)
assert(Vorlesung.has_student.__doc__ != None)
assert(Vorlesung.add_grade.__doc__ != None)
assert(Vorlesung.get_grade.__doc__ != None)
assert(Paul.__doc__ != None)
assert(Paul.__init__.__doc__ != None)
assert(Paul.add_vorlesung.__doc__ != None)
assert(Paul.add_student.__doc__ != None)
assert(Paul.registriere_in_vorlesung.__doc__ != None)
assert(Paul.get_studenten_in_vorlesung.__doc__ != None)
assert(Paul.add_grade.__doc__ != None)
assert(Paul.get_grades_for_student.__doc__ != None)
```