<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

## Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 11.12.2016 um 23:59 Uhr
### Wichtig: Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue06.ipynb ab (nicht umbenennen, keine Kopie erstellen!). Sonst kann eine Bewertung nicht stattfinden.

# Aufgabe 1: String-API (6 Punkte)
Diese Aufgabe lässt sich eventuell einfacher mithilfe der Python-Dokumentation lösen:
https://docs.python.org/3/library/stdtypes.html.

a) Schreiben Sie eine Funktion `s_verbinder(l,z)`, welches eine Liste mit Strings `l` zu einem einzigen String verbindet, wobei zwischen die einzelnen Listenelemente jeweils das Zeichen `z` eingefügt werden soll.

<!-- grade_id: string_sol_1, points: 0, grade: false, solution: true -->
```python
def s_verbinder(l,z):
    # Your Code here
    ### BEGIN SOLUTION
    return z.join(l)
    ### END SOLUTION
```

<!-- grade_id: string_test_1, points: 1, grade: true, solution: false -->
```python
assert(s_verbinder(['1','2','3','4'], '-') == '1-2-3-4')
### HIDESTART
assert(s_verbinder(['a','b','c','d'], '-') == 'a-b-c-d')
### HIDEEND
```

b) Schreiben Sie eine Funktion `s_suche(s1,s2)`, welche den niedrigsten Index ausgibt, an welcher der String `s2` in `s1` auftaucht.

<!-- grade_id: string_sol_2, points: 0, grade: false, solution: true -->
```python
def s_suche(s1,s2):
    # Your Code here
    ### BEGIN SOLUTION
    return s1.find(s2)
    ### END SOLUTION
```

<!-- grade_id: string_test_2, points: 1, grade: true, solution: false -->
```python
assert(s_suche('Test', 'T') == 0)
assert(s_suche('Test', 'e') == 1)
assert(s_suche('Test', 's') == 2)
assert(s_suche('Test', 't') == 3)
assert(s_suche('Test', 'es') == 1)
### HIDESTART
assert(s_suche('Grundlagen der Programmierung', 'der') == 11)
assert(s_suche('Grundlagen der Programmierung', 'P') == 15)
### HIDEEND
```

c) Schreiben Sie eine Funktion `s_klein(s)`, welche alle Großbuchstaben im string `s` durch entsprechende Kleinbuchstaben ersetzt.

<!-- grade_id: string_sol_3, points: 0, grade: false, solution: true -->
```python
def s_klein(s):
    # Your Code here
    ### BEGIN SOLUTION
    return s.lower()
    ### END SOLUTION
```

<!-- grade_id: string_test_3, points: 1, grade: true, solution: false -->
```python
assert(s_klein('Hallo Welt!') == 'hallo welt!')
### HIDESTART
assert(s_klein('Dies ist ein geHEimer Satz.') == 'dies ist ein geheimer satz.')
### HIDEEND
```

d) Schreiben Sie eine Funktion `s_special_substitution(s)`, welche für einen string `s` sowohl den ersten Buchstaben, als auch alle Buchstaben, welche auf ein Leerzeichen folgen, groß schreibt.

<!-- grade_id: string_sol_4, points: 0, grade: false, solution: true -->
```python
def s_special_substitution(s):
    # Your Code here
    ### BEGIN SOLUTION
    return s.title()
    ### END SOLUTION
```

<!-- grade_id: string_test_4, points: 0, grade: true, solution: false -->
```python
assert(s_special_substitution('Der Typ string für Zeichenketten ist als eine Klasse definiert.') == 
       'Der Typ String Für Zeichenketten Ist Als Eine Klasse Definiert.')
### HIDESTART
assert(s_special_substitution('Dies ist ein geheimer Satz.') == 
       'Dies Ist Ein Geheimer Satz.')
### HIDEEND
```

e) Schreiben Sie eine Funktion `get_words(s)`, welche für einen string `s` eine Liste mit den einzelnen Wörtern von `s` zurückgibt.

<!-- grade_id: string_sol_5, points: 0, grade: false, solution: true -->
```python
def get_words(s):
    # Your Code here
    ### BEGIN SOLUTION
    return s.split(" ")
    ### END SOLUTION
```

<!-- grade_id: string_test_5, points: 0, grade: true, solution: false -->
```python
assert(get_words('Grundlagen der Programmierung') == ['Grundlagen', 'der', 'Programmierung'])
assert(get_words('Der Typ string für Zeichenketten ist als eine Klasse definiert.') ==
       ['Der', 'Typ', 'string', 'für', 'Zeichenketten', 'ist', 'als', 'eine', 'Klasse', 'definiert.'])
### HIDESTART
assert(get_words('Dies ist ein geheimer Satz.') == ['Dies', 'ist', 'ein', 'geheimer', 'Satz.'])
### HIDEEND
```

f) Schreiben Sie eine Funktion `split_3(s)`, welche einen string `s` an dem ersten Index `i` aufteilt, an welcher ein Punkt ('.') vorkommt. Es soll ein Triple `(s_links, i, s_rechts)` zurückgegeben werden, wobei `s_links` der Teilstring von `s` ist, welcher alle Zeichen links von `s[i]` enthält und `s_rechts`der Teilstring von `s` ist, welcher alle Zeichen rechts von `s[i]` enthält.

<!-- grade_id: string_sol_6, points: 0, grade: false, solution: true -->
```python
def split_3(s):
    # Your Code here
    ### BEGIN SOLUTION
    l = s.partition('.')
    sol = (l[0], s.index('.'), l[2])
    return sol
    ### END SOLUTION
```

<!-- grade_id: string_test_6, points: 0, grade: true, solution: false -->
```python
assert(split_3('Dies ist ein Satz. Dies ist noch ein Satz.') == ('Dies ist ein Satz', 17, ' Dies ist noch ein Satz.'))
### HIDESTART
assert(split_3('Dies ist ein geheimer Satz. Dies ist noch ein geheimer Satz.') == ('Dies ist ein geheimer Satz', 26, ' Dies ist noch ein geheimer Satz.'))
### HIDEEND
```

<!-- grade_id: tests, points: 0, grade: false, solution: true -->
```python
class Ellipse:
    def __init__(s, x, y, a1, a2):
        s.x = x
        s.y = y
        s.achse1 = a1
        s.achse2 = a2
        
    def __eq__(s, otherEllipse):
        return (s.x == otherEllipse.x and s.y == otherEllipse.y and s.achse1 == otherEllipse.achse1 and 
                s.achse2 == otherEllipse.achse2)

e1 = Ellipse(0, 1, 2, 3)
e2 = e1
e3 = Ellipse(0, 1, 2, 3)

print("==:")
print(e1 == e2)
print(e1 == e3)
print(e2 == e3)

print("is:")
print(e1 is e2)
print(e1 is e3)
print(e2 is e3)
```

Output:
```
==:
True
True
True
is:
True
False
False
```

<!-- grade_id: ascsac, points: 0, grade: false, solution: true -->
```python
class Patient(object):
    def __init__(self, name, alter):
        self.name = name
        self.alter = alter
    
    def getName(self):
        return self.name
    
    def getAlter(s):
        return s.alter

patient = Patient("Michael", 26)
print(patient.getAlter())
```

Output:
```
26
```