<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile). Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue07.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen!). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 18.12.2016 um 23:59 Uhr

# Aufgabe 1: String-API
Diese Aufgabe lässt sich eventuell einfacher mithilfe der Python-Dokumentation lösen:
https://docs.python.org/3/library/stdtypes.html.

### a) 
Schreiben Sie eine Funktion `s_verbinder(l,z)`, welches eine Liste mit Strings `l` zu einem einzigen String verbindet, wobei zwischen die einzelnen Listenelemente jeweils das Zeichen `z` eingefügt werden soll.

<!-- grade_id: string_sol_1, points: 0, grade: false, solution: true -->
```python
def s_verbinder(l,z):
    # Your Code here
    ### BEGIN SOLUTION
    """Methode, welche eine Liste mit Strings zu einem einzigen STring verbindet."""
    return z.join(l)
    ### END SOLUTION
```

<!-- grade_id: string_test_1, points: 1, grade: true, solution: false -->
```python
assert(s_verbinder(['1','2','3','4'], '-') == '1-2-3-4')
### HIDESTART
assert(s_verbinder(['a','b','c','d'], '-') == 'a-b-c-d')
### HIDEEND
```

### b) 
Schreiben Sie eine Funktion `s_suche(s1,s2)`, welche den niedrigsten Index ausgibt, an welcher der String `s2` in `s1` auftaucht.

<!-- grade_id: string_sol_2, points: 0, grade: false, solution: true -->
```python
def s_suche(s1,s2):
    # Your Code here
    ### BEGIN SOLUTION
    """Methode, welche den Niedrigsten Index ausgibt, an welchem der String s2 in s1 auftaucht."""
    return s1.find(s2)
    ### END SOLUTION
```

<!-- grade_id: string_test_2, points: 1, grade: true, solution: false -->
```python
assert(s_suche('Grundlagen der Programmierung', 'der') == 11)
assert(s_suche('Grundlagen der Programmierung', 'P') == 15)
### HIDESTART
assert(s_suche('Test', 'T') == 0)
assert(s_suche('Test', 'e') == 1)
assert(s_suche('Test', 's') == 2)
assert(s_suche('Test', 't') == 3)
assert(s_suche('Test', 'es') == 1)
assert(s_suche('TestT', 'T') == 0)
### HIDEEND
```

### c) 
Schreiben Sie eine Funktion `s_klein(s)`, welche alle Großbuchstaben im string `s` durch entsprechende Kleinbuchstaben ersetzt.

<!-- grade_id: string_sol_3, points: 0, grade: false, solution: true -->
```python
def s_klein(s):
    # Your Code here
    ### BEGIN SOLUTION
    """Methode, welche alle Großbuchstaben des übergebenen Strings durch entsprechende Kleinbuchstaben ersetzt."""
    return s.lower()
    ### END SOLUTION
```

<!-- grade_id: string_test_3, points: 1, grade: true, solution: false -->
```python
assert(s_klein('Hallo Welt!') == 'hallo welt!')
### HIDESTART
assert(s_klein('Dies ist ein geHEimer Satz.') == 'dies ist ein geheimer satz.')
### HIDEEND
```

# Aufgabe 2: Vektorklasse
In dieser Aufgabe sollen Sie eine Klasse ```Vektor``` implementieren.
Ein Vektor $v$ besteht aus $n$ einzelnen Zahlen, also $v = (v_0,\ldots,v_{n-1})$.
Das $n$ kann hier eine beliebige Zahl aus den natürlichen Zahlen $\mathbb N = \{1,2,3,\ldots\}$ sein.
Die Klasse ```Vektor``` soll für ein Objekt $v = (v_0,\ldots,v_{n-1})$ die folgenden Methoden bereitstellen:
- ```index(i)```: Liefert die Zahl an der $i$-ten Stelle des Vektors, d.h. für $v = (v_0,\ldots,v_{n-1})$ wird $v_i$ zurückgegeben.
- ```add(a)```: Addiert den Vektor $v = (v_0,\ldots,v_{n-1})$ mit dem übergebenen Vektor $a = (a_0,\ldots,a_{n-1})$ und liefert das Ergebnis als neuen Vektor $e = (v_0+a_0,\ldots,v_{n-1}+a_{n-1})$ zurück.
- ```substract(a)```: Subtrahiert den Vektor $a = (a_0,\ldots,a_{n-1})$ von dem Vektor $v = (v_0,\ldots,v_{n-1})$ und liefert das Ergebnis als neuen Vektor $s = (v_0-a_0,\ldots,v_{n-1}-a_{n-1})$ zurück.
- ```dot(w)```: Berechnet das Vektorprodukt für den Vektor $v = (v_0,\ldots,v_{n-1})$ und dem übergebenen Vektor $w = (w_0,\ldots,w_{n-1})$ und liefert das Ergebnis ($\sum_{i=0}^{n-1} v_i\cdot w_i$) zurück.
- ```norm()```: Berechnet die Norm ($\sqrt[]{\sum_{i=0}^{n-1} v_i^2}$) des Vektors $v = (v_0,\ldots,v_{n-1})$ und liefert diese zurück.

Für die Methoden ```add```, ```substract``` und ```dot``` können Sie davon ausgehen, dass beide Vektoren immer dieselbe Länge haben.

<!-- grade_id: vector_sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Vector(object):
    """Die Klasse Vektor realisiert einen mathematischen Vektor, auf welchem Vektoroperationen ausgeführt werden können."""
    def __init__(self, vec):
        self.vec = vec

    def add(self, v):
        erg = [0]*self.length()
        for i in range(0, v.length()):
            erg[i] = self.vec[i] + v.index(i)
        return Vector(erg)
    
    def subtract(self, v):
        erg = [0]*self.length()
        for i in range(0, v.length()):
            erg[i] = self.vec[i] - v.index(i)
        return Vector(erg)
    
    def dot(self, v):
        erg = 0
        for i in range(0, v.length()):
            erg = erg + self.vec[i] * v.index(i)
        return erg
    
    def norm(self):
        """norm-Methode der Klasse Vektor"""
        erg = 0
        for i in range(0, len(self.vec)):
            erg = erg + self.vec[i] ** 2
        erg = erg ** 0.5
        return erg

    def length(self):
        return len(self.vec)
    
    def index(self, i):
        return self.vec[i]
### END SOLUTION
```

<!-- grade_id: vector_test_0, points: 1, grade: true, solution: false -->
```python
a = Vector([4, 3, 7])
assert(a.index(0) == 4)
assert(a.index(1) == 3)
assert(a.index(2) == 7)

### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    for i in range(0,n):
        assert(a.index(i) == i+1)
### HIDEEND
```

<!-- grade_id: vector_test_1, points: 1, grade: true, solution: false -->
```python
a = Vector([1, 2, 3])
b = Vector([3, 4, 5])

e = a.add(b)
assert(e.index(0) == 4)
assert(e.index(1) == 6)
assert(e.index(2) == 8)

### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    b = Vector(range(0,n))

    e = a.add(b)

    for i in range(0,n):
        assert(e.index(i) == (a.index(i) + b.index(i)))
### HIDEEND
```

<!-- grade_id: vector_test_2, points: 1, grade: true, solution: false -->
```python
a = Vector([1, 2, 3])
b = Vector([3, 4, 5])
s = a.subtract(b)
assert(s.index(0) == -2)
assert(s.index(1) == -2)
assert(s.index(2) == -2)
### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    b = Vector(range(0,n))

    s = a.subtract(b)

    for i in range(0,n):
        assert(s.index(i) == (a.index(i) - b.index(i)))
### HIDEEND
```

<!-- grade_id: vector_test_3, points: 1, grade: true, solution: false -->
```python
a = Vector([1, 2, 3])
b = Vector([3, 4, 5])
dot = a.dot(b)
assert(dot == 26)
### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    b = Vector(range(0,n))
    dot = 0
    for i in range(0,n):
        dot = dot + a.index(i) * b.index(i)
    assert(a.dot(b) == dot)
### HIDEEND
```

<!-- grade_id: vector_test_4, points: 1, grade: true, solution: false -->
```python
a = Vector([1, 2, 3])
norm = a.norm()
assert(norm == (14**0.5))
### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    norm = 0
    for i in range(0, n):
        norm = norm + a.index(i) ** 2
    norm = norm ** 0.5
    assert(a.norm() == norm)    
### HIDEEND
```

# Aufgabe 3: Objektorientierung

Betrachten Sie für diese Aufgabe die gegebene Klasse `Patient`.

```python
class Patient():
    anz_patienten = 0
    
    def __init__(self, name, alter, privat_patient):
        self.name = name
        self.alter = alter
        self.privat_patient = privat_patient
        
    def get_name(self):
        return self.name
    
    def get_alter(s):
        return s.alter
    
    def is_privat_patient(self):
        return self.privat_patient
    
    def __eq__(self, p):
        return self.name == p.name and self.alter == p.alter
```

### a)
Erläutern Sie zunächst den Unterschied der Operationen '==' und 'is'. 
Weisen Sie dann den beiden gegebenen Variablen `p1` und `p2` Objekte der Klasse `Patient` zu, sodass `p1 == p2` gilt, jedoch nicht `p1 is p2`.

<!-- grade_id: patient_sol_1_1, points: 1, grade: true, solution: true -->
Der Operator `==` führt die Methode `__eq__` des Objektes aus, welches auf der linken Seite des Operators steht.
Der Operator `is` hingegen überprüft, ob die referenzierten Objekte auf beiden Seiten dieselben sind.

<!-- grade_id: patient_sol_1_2, points: 0, grade: false, solution: true -->
```python
p1 = None # None durch entsprechende Belegung ersetzen!
p2 = None # None durch entsprechende Belegung ersetzen!
### BEGIN SOLUTION
p1 = Patient("Max Meier", 26, False)
p2 = Patient ("Max Meier", 26, False)
### END SOLUTION
```

<!-- grade_id: patient_test_1, points: 1, grade: true, solution: false -->
```python
assert(isinstance(p1, Patient))
assert(isinstance(p2, Patient))
assert(p1 == p2)
assert(not p2 is p1)
```

### b)
Wie Sie sehen, verfügt die Klasse `Patient` auch über eine Klassenvariable `anz_patienten`.
Der Wert dieser Variablen soll nun immer um 1 erhöht werden, wenn ein neuer Patient erstellt wurde, d.h. nach Ausführung des folgenden Programms sollte `anz_patienten` den Wert 2 haben:

```Python
p1 = Patient("Max Meier", 26, False)
p2 = Patient ("Max Mustermann", 12, False)
```
An welcher Stelle muss dies realisiert werden?
Geben Sie auch den entsprechenden Befehl an, welcher dies realisiert?

<!-- grade_id: patient_sol_2, points: 1, grade: true, solution: true -->
Es muss im Konstruktor der Befehl 
``` Python
Patient.anz_patienten += 1
```
ausgeführt werden.

### c)
Wie Sie sehen, heißt der erste Parameter der Methode `get_alter()` nicht `self`.
Würde das folgende Programm einen Fehler erzeugen?
Wenn ja, an welcher Stelle und warum, wenn nein, erklären Sie warum nicht?

``` Python
p1 = Patient("Max Meier", 26, False)
print(p1.get_alter())
```

<!-- grade_id: patient_sol_3, points: 1, grade: true, solution: true -->
Das Programm funktioniert, da der Name für den ersten Parameter einer Objektmethode nur per Konvention `self` lautet.

### d)
Angenommen, die obige Klasse `Patient` würde über keinen Konstruktor verfügen.
Würde das folgende Programm einen Fehler erzeugen?
Wenn ja, an welcher Stelle und warum, wenn nein, erklären Sie warum nicht?
``` Python
p1 = Patient()
print(p1.get_name())
```

<!-- grade_id: patient_sol_4, points: 1, grade: true, solution: true -->
Zunächst einmal wird der fehlende Konstruktor von Python durch einen leere Standardkonstruktor ersetzt.
Das Programm wirft einen Fehler wenn versucht wird die Methode p1.get_name() aufzurufen, da der Patient über keine Objektvariable `name` verfügt (diese wurde im Standardkonstruktor nicht definiert).

# Aufgabe 4: PAUL
In dieser Aufgabe sollen Sie eine einfache Version von PAUL implementieren.
Das System soll aus 3 Klassen bestehen: `Student`, `Vorlesung` und `Paul`.
Die Klasse `Student` ist Ihnen bereits gegeben (siehe untere Codezelle).
Ihre Aufgabe ist nun die Klassen `Vorlesung` und `Paul` zu vervollständigen, sodass die gegebenen Tests funktionieren.

## a)
Erstellen Sie eine Klasse `Vorlesung`, welche über die folgenden Methoden verfügt:

- Einen Kontruktor, welcher die ID der Vorlesung (int), sowie den Namen der Vorlesung (str) übergeben bekommt.
- `get_id`: Gibt die ID der Vorlesung zurück
- `get_name`: Gibt den Namen der Vorlesung zurück
- `add_student`: Registriert den Studenten mit der übergebenen Matrikelnummer in der Vorlesung
- `has_student`: Liefert `True`, falls sich ein Student mit der übergebenen Matrikelnummer in der Vorlesung registriert hat, ansonsten `False`.
- `add_grade`: Setzt die Note des Studenten mit der übergebenen Matrikelnummer auf die übergebene Note. Sollte der Student bereits eine Note haben, so wird diese überschrieben. Wenn der Student nicht in der Vorlesung registriert ist, soll keine Aktion ausgeführt werden.
- `get_grade`: Liefert die Note für diese Vorlesung, welche der Student mit der übergebenen Matrikelnummer erhalten hat. Falls die Matrikelnummer nicht in der Vorlesung registriert ist, soll `-1` zurückgegeben werden. Ist der Student zwar registriert, hat aber noch keine Note erhalten, wird `None` zurückgegeben.


<!-- grade_id: paul_sol_1, points: 0, grade: false, solution: true -->
```python
class Student:
    def __init__(self, mat_nr, name):
        self.mat_nr = mat_nr
        self.name = name
    
    def get_mat_nr(self):
        return self.mat_nr
    
    def get_name(self):
        return self.name
    
    def __str__(self):
        return str(self.mat_nr) + " " + self.name
        
class Vorlesung:
    def __init__(self, v_id, name):
        self.v_id = v_id
        self.name = name
        self.studenten = {}
        
    def get_id(self):
        return self.v_id
    
    def get_name(self):
        return self.name
    
    def add_student(self, mat_nr):
        self.studenten[mat_nr] = None
        
    def has_student(self, mat_nr):
        if mat_nr in self.studenten:
            return True
        return False
    
    def add_grade(self, mat_nr, grade):
        if mat_nr in self.studenten:
            self.studenten[mat_nr] = grade
    
    def get_grade(self, mat_nr):
        if mat_nr in self.studenten:
            return self.studenten[mat_nr]
        return -1
```

<!-- grade_id: paul_test_1_1, points: 1, grade: true, solution: false -->
```python
# Erstelle Studenten:
s1 = Student(9863451, "Ralf B")
s2 = Student(9863452, "Jessica W")
s3 = Student(9863453, "Philipp K")
s4 = Student(9863454, "Andrea B")
s5 = Student(9863455, "Sebastian A")
s6 = Student(9863456, "Luca B")
studenten = [s1,s2,s3,s4,s5,s6]

# Kurs-Test
gp1 = Vorlesung(1, "Grundlagen der Programmierung 1")

assert(gp1.get_id() == 1)
assert(gp1.get_name() == "Grundlagen der Programmierung 1")

for s in studenten:
    gp1.add_student(s.get_mat_nr())

for s in studenten:
    assert(gp1.has_student(s.get_mat_nr()))

### HIDESTART
mod = Vorlesung(2, "Modellierung")

assert(mod.get_id() == 2)
assert(mod.get_name() == "Modellierung")

mod.add_student(s1.get_mat_nr())
mod.add_student(s2.get_mat_nr())
mod.add_student(s3.get_mat_nr())

assert(mod.has_student(s1.get_mat_nr()))
assert(mod.has_student(s2.get_mat_nr()))
assert(mod.has_student(s3.get_mat_nr()))
assert(not mod.has_student(s4.get_mat_nr()))
assert(not mod.has_student(s5.get_mat_nr()))
assert(not mod.has_student(s6.get_mat_nr()))
### HIDEEND
```

<!-- grade_id: paul_test_1_2, points: 2, grade: true, solution: false -->
```python
gp1.add_grade(s1.get_mat_nr(), 1.0)
gp1.add_grade(s2.get_mat_nr(), 1.3)
gp1.add_grade(s3.get_mat_nr(), 3.0)
gp1.add_grade(s4.get_mat_nr(), 2.7)
gp1.add_grade(s5.get_mat_nr(), 1.3)
gp1.add_grade(s6.get_mat_nr(), 5.0)

assert(gp1.get_grade(s1.get_mat_nr()) == 1.0)
assert(gp1.get_grade(s2.get_mat_nr()) == 1.3)
assert(gp1.get_grade(s3.get_mat_nr()) == 3.0)
assert(gp1.get_grade(s4.get_mat_nr()) == 2.7)
assert(gp1.get_grade(s5.get_mat_nr()) == 1.3)
assert(gp1.get_grade(s6.get_mat_nr()) == 5.0)

gp1.add_grade(1234567, 1.0) # Nicht in der Vorlesung registrierter Student will eine Note haben
assert(gp1.get_grade(1234567) == -1) # Matrikelnummer nicht im Kurs vorhanden

### HIDESTART
assert(mod.get_grade(s4.get_mat_nr()) == -1)
assert(mod.get_grade(s5.get_mat_nr()) == -1)
assert(mod.get_grade(s6.get_mat_nr()) == -1)
### HIDEEND
```

## b)
Ergänzen Sie nun die Klasse `PAUL`, welche über die folgenden Methoden verfügt:
- `add_vorlesung`: Erstellt eine neue Vorlesung mit übergebener ID und Namen im System. Sie können davon ausgehen, dass `add_vorlesung` niemals mehrfach für dieselbe Vorlesungs-ID aufgerufen wird.
- `add_student`: Erstellt einen neuen Studenten mit übergebener Matrikelnummer und Namen im System. Sie können davon ausgehen, dass `add_student` niemals mehrfach für dieselbe Matrikelnummer aufgerufen wird.
- `registriere_student_in_vorlesung`: Fügt den Studenten mit der übergebenen Matrikelnummer der Vorlesung mit der übergebenen ID hinzu. Sie können davon ausgehen, dass `registriere_student_in_vorlesung` niemals mehrfach für ein Paar (`matr_nr`, `v_id`) aufgerufen wird.
- `get_studenten_in_vorlesung`: Liefert eine Liste, welche alle Matrikelnummern enthält, die in der Vorlesung mit der übergebenen ID registriert sind. Falls das System keine Vorlesung mit der übergebenen ID enthält, soll `None` zurückgegeben werden.
- `add_grade`: Fügt dem Studenten mit der übergebenen Matrikelnummer eine Note für die Vorlesung mit ID `v_id` hinzu.
- `get_grades_for_student`: Liefert ein Dictionary mit allen Vorlesungen, in welcher der übergebene Student registriert ist und die entsprechenden Noten, die der Student für die Vorlesung erhalten hat.

Eine beispielhafte Anwendung können Sie an den gegebenen Tests ablesen.

<!-- grade_id: paul_sol_2, points: 0, grade: false, solution: true -->
```python
class Paul:
    def __init__(self):
        self.studenten = []
        self.vorlesungen = {}
    
    def add_vorlesung(self, v_id, name):
        v = Vorlesung(v_id, name)
        self.vorlesungen[v_id] = v
        
    def add_student(self, mat_nr, name):
        s = Student(mat_nr, name)
        self.studenten.append(s)
        
    def registriere_in_vorlesung(self, mat_nr, v_id):
        self.vorlesungen[v_id].add_student(mat_nr)
        
    def get_studenten_in_vorlesung(self, v_id):
        if v_id in self.vorlesungen:
            L = []
            for i in self.vorlesungen[v_id].studenten.keys():
                L.append(i)
            return L
        return None
    
    def add_grade(self, mat_nr, grade, v_id):
        self.vorlesungen[v_id].add_grade(mat_nr, grade)
        
    def get_grades_for_student(self, mat_nr):
        d = {}
        for v_id in self.vorlesungen:
            v = self.vorlesungen[v_id]
            if v.has_student(mat_nr):
                d[v.get_name()] = v.get_grade(mat_nr)
        return d
```

<!-- grade_id: paul_test_2_1, points: 1, grade: true, solution: false -->
```python
paul = Paul()

paul.add_vorlesung(1, "Grundlagen der Programmierung 1")
paul.add_vorlesung(2, "Modellierung")
paul.add_student(9863451, "Ralf B")
paul.add_student(9863452, "Jessica W")
paul.add_student(9863453, "Philipp K")

paul.registriere_in_vorlesung(9863451, 1)
paul.registriere_in_vorlesung(9863452, 1)
paul.registriere_in_vorlesung(9863453, 1)

paul.registriere_in_vorlesung(9863451, 2)
paul.registriere_in_vorlesung(9863452, 2)
paul.registriere_in_vorlesung(9863453, 2)
```

<!-- grade_id: paul_test_2_2, points: 2, grade: true, solution: false -->
```python
paul.add_grade(9863451, 3.0, 1)
paul.add_grade(9863452, 1.3, 1)
paul.add_grade(9863453, 2.0, 1)

paul.add_grade(9863451, 5.0, 2)
paul.add_grade(9863452, 2.3, 2)
paul.add_grade(9863453, 1.7, 2)

assert(paul.get_studenten_in_vorlesung(1) == [9863451, 9863452, 9863453])
assert(paul.get_studenten_in_vorlesung(2) == [9863451, 9863452, 9863453])
assert(paul.get_grades_for_student(9863451) == {'Grundlagen der Programmierung 1': 3.0, 'Modellierung': 5.0})
assert(paul.get_grades_for_student(9863452) == {'Grundlagen der Programmierung 1': 1.3, 'Modellierung': 2.3})
assert(paul.get_grades_for_student(9863453) == {'Grundlagen der Programmierung 1': 2.0, 'Modellierung': 1.7})
```

<!-- grade_id: doc_test, points: 1, grade: true, solution: false -->
```python
### HIDESTART
assert(s_verbinder.__doc__ != None)
assert(s_suche.__doc__ != None)
assert(s_klein.__doc__ != None)
# TODO: Kann man testen, ob Klassen (z.B. die Klasse Vektor) bzw. Objektmethoden dokumentiert wurden?
### HIDEEND
```