<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

Format:

MatrikelNummer E-Mail-Adresse
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
234567 testStudent2@mail.upb.de
345678 vorname.nachname@mail.upb.de

Hinweis: Das sieht nicht nach separaten Zeilen aus, wenn Sie die Markdownzelle anzeigen lassen. Das ist egal; wichtig ist Formatierung der Eingabe.


## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile), für das entsprechende Übungsblatt! (Wenn Sie Blatt 7 abgeben wolllen hilft es nicht, bei Blatt 1 auf Submit zu drücken!) Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue07.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen, auch nicht Dateiname hue7.ipynb !). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 18.12.2016 um 23:59 Uhr

# Aufgabe 1: String-API
Diese Aufgabe lässt sich eventuell einfacher mithilfe der Python-Dokumentation lösen:
https://docs.python.org/3/library/stdtypes.html.

## a) 
Schreiben Sie eine Funktion `s_verbinder(l,z)`, welche eine Liste mit Strings `l` zu einem einzigen String verbindet und diesen zurückgibt, wobei zwischen die einzelnen Listenelemente jeweils das der Funktion übergebene Zeichen `z` eingefügt werden soll. Behandeln Sie Randfälle (leere Liste, leeres Zeichen) sinnvoll (siehe Tests)! 

<!-- grade_id: string_sol_1, points: 0, grade: false, solution: true -->
```python
def s_verbinder(l,z):
    ### BEGIN SOLUTION
    """Funktion, welche eine Liste mit Strings zu einem einzigen String verbindet."""
    return z.join(l)
    ### END SOLUTION
```

<!-- grade_id: string_test_1, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
assert(s_verbinder(['1','2','3','4'], '-') == '1-2-3-4')
assert(s_verbinder(['a','b','c','d'], '') == 'abcd')
assert(s_verbinder([], '') == '')
assert(s_verbinder([], 'x') == '')
### HIDESTART
assert(s_verbinder(['a','b','c','d'], '-') == 'a-b-c-d')
### HIDEEND
```

## b) 
Schreiben Sie eine Funktion `s_suche(s1,s2)`, welche den niedrigsten Index zurückgibt, bei dem der String `s2` in `s1` auftaucht.
Ist s2 nicht in s1 enthalten, soll die Funktion `-1` zurückgeben.
Behandeln Sie Randfälle (siehe Tests, insbesondere Suche nach leerem String) sinnvoll!

<!-- grade_id: string_sol_2, points: 0, grade: false, solution: true -->
```python
def s_suche(s1,s2):
    ### BEGIN SOLUTION
    """Funktion, welche den niedrigsten Index ausgibt, 
    bei dem der String s2 in s1 auftaucht."""
    return s1.find(s2)
    ### END SOLUTION
```

<!-- grade_id: string_test_2, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
assert(s_suche('Grundlagen der Programmierung', 'der') == 11)
assert(s_suche('Grundlagen der Programmierung', 'r') == 1)
assert(s_suche('', 'P') == -1)
assert(s_suche('Grundlagen der Programmierung', 'X') == -1)
assert(s_suche('Grundlagen der Programmierung', '') == 0)
### HIDESTART
assert(s_suche('Test', 'T') == 0)
assert(s_suche('Test', 'e') == 1)
assert(s_suche('Test', 's') == 2)
assert(s_suche('Test', 't') == 3)
assert(s_suche('Test', 'es') == 1)
assert(s_suche('TestT', 'T') == 0)
### HIDEEND
```

## c) 
Schreiben Sie eine Funktion `split_3(s)`, welche einen string `s` an dem ersten Index `i` aufteilt, an welcher ein Punkt ('.') vorkommt. Es soll ein Triple `(s_links, i, s_rechts)` zurückgegeben werden, wobei `s_links` der Teilstring von `s` ist, welcher alle Zeichen links von `s[i]` enthält und `s_rechts` der Teilstring von `s` ist, welcher alle Zeichen rechts von `s[i]` enthält.
Falls der übergebene String keinen Punkt enthält, soll ein leeres Tupel zurückgegeben werden.

<!-- grade_id: string_sol_3, points: 0, grade: false, solution: true -->
```python
def split_3(s):
    ### BEGIN SOLUTION
    """ Methode welche welche einen string s an dem ersten Index i aufteilt, an welcher ein Punkt ('.') vorkommt."""
    if not '.' in s:
        return ()
    l = s.partition('.')
    sol = (l[0], s.index('.'), l[2])
    return sol
    ### END SOLUTION
```

<!-- grade_id: string_test_3, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
assert(split_3('Dies ist ein Satz. Dies ist noch ein Satz.') == ('Dies ist ein Satz', 17, ' Dies ist noch ein Satz.'))
assert(split_3('Ein Satz ohne Punkt') == ())
### HIDESTART
assert(split_3('Dies ist ein geheimer Satz. Dies ist noch ein geheimer Satz.') == ('Dies ist ein geheimer Satz', 26, ' Dies ist noch ein geheimer Satz.'))
### HIDEEND
```

# Aufgabe 2: Objektorientierung

Betrachten Sie für diese Aufgabe die gegebene Klasse `Patient`.

```python
class Patient:
    """Realisiert einen Patienten"""
    
    """Klassenvariable, die die Anzahl der erstellten Patienten zählen soll"""
    anz_patienten = 0
    
    def __init__(self, name, alter, privat_patient):
        """Konstruktor
        
        Args:
            name: Name des Patienten (str)
            alter: Alter des Patienten (int)
            privat_patient: True, falls der Patient ein Privatpatient ist, ansonsten False (bool)
        """
        self.name = name
        self.alter = alter
        self.privat_patient = privat_patient
        
    def get_name(self):
        """Liefert den Namen des Patienten
        
        Returns:
            str Name des Patienten
        """
        return self.name
    
    def get_alter(s):
        """Liefert das Alter des Patienten
        
        Returns:
            int Alter des Patienten
        """
        return s.alter
    
    def is_privat_patient(self):
        """Liefert True, falls der Patient ein Privatpatient ist, ansonsten False
        
        Returns:
            bool True, falls der Patient ein Privatpatient ist, ansonsten False
        """
        return self.privat_patient
    
    def __eq__(self, p):
        """Vergleicht den Patient self mit dem Patient p auf Basis ihrer Attribute."""
        return self.name == p.name and self.alter == p.alter
```

## a)
Erläutern Sie zunächst den Unterschied der Operationen `==` und `is` anhand der Klasse `Patient`. 
Was passiert, wenn zwei Objekte der Klasse `Patient` miteinander verglichen werden?

<!-- grade_id: patient_sol_1_1, points: 1, grade: true, solution: true -->
Der Operator `==` führt die Methode `__eq__` des Objektes aus, welches auf der linken Seite des Operators steht.
Der Operator `is` hingegen überprüft, ob die referenzierten Objekte auf beiden Seiten dieselben sind.

Weisen Sie dann den beiden gegebenen Variablen `p1` und `p2` Objekte der Klasse `Patient` zu, sodass `p1 == p2` gilt, jedoch nicht `p1 is p2`.

<!-- grade_id: patient_sol_1_2, points: 0, grade: false, solution: true -->
```python
p1 = None # None durch entsprechende Belegung ersetzen!
p2 = None # None durch entsprechende Belegung ersetzen!
### BEGIN SOLUTION
p1 = Patient("Max M", 26, False)
p2 = Patient ("Max M", 26, False)
### END SOLUTION
```

<!-- grade_id: patient_test_1, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
assert(isinstance(p1, Patient))
assert(isinstance(p2, Patient))
assert(p1 == p2)
assert(not p2 is p1)
```

## b)
Wie Sie sehen, heißt der erste Parameter der Methode `get_alter()` nicht `self`.
Würde das folgende Programm einen Fehler erzeugen?
Wenn ja, an welcher Stelle und warum; wenn nein, erklären Sie warum nicht?

```Python
p1 = Patient("Max Meier", 26, False)
print(p1.get_alter())
```

*Hinweis*: Ein simples Copy-Paste einer Fehlermeldung reicht für eine Beschreibung des Fehlers nicht aus.

<!-- grade_id: patient_sol_3, points: 1, grade: true, solution: true -->
Das Programm funktioniert, da der Name für den ersten Parameter einer Objektmethode nur per Konvention `self` lautet. Es geht nur um die Anzahl der Parameter; die Namen sind egal.

## c)
Angenommen, die obige Klasse `Patient` würde nicht die oben definierte Methode `__init__` enthalten.
Würde das folgende Programm einen Fehler erzeugen?
Wenn ja, an welcher Stelle und warum; wenn nein, erklären Sie warum nicht?   
```Python
p1 = Patient()
print(p1.get_name())
```
*Hinweis*: Ein simples Copy-Paste einer Fehlermeldung reicht für eine Beschreibung des Fehlers nicht aus.

<!-- grade_id: patient_sol_4, points: 1, grade: true, solution: true -->
Zunächst einmal wird der fehlende Konstruktor von Python durch einen leeren Standardkonstruktor ersetzt.
Das Programm wirft einen Fehler, wenn versucht wird, die Methode p1.get_name() aufzurufen, da der Patient über keine Objektvariable `name` verfügt (diese wurde im Standardkonstruktor nicht definiert).

# Aufgabe 3: Vektoren
In dieser Aufgabe sollen Sie eine Klasse `Vector` implementieren.
Diese soll Objekte in Anlehnung an mathematische Vektoren repräsentieren.
Ein Vektor $v$ besteht (in dieser Aufgabe) aus $n$ einzelnen Zahlen, also $v = (v_0,\ldots,v_{n-1})$.
Die Dimension $n$ kann hier eine beliebige Zahl aus den natürlichen Zahlen $\mathbb N = \{1,2,3,\ldots\}$ sein.
Die Klasse `Vector` soll für ein Objekt $v = (v_0,\ldots,v_{n-1})$ die folgenden (Objekt-)Methoden bereitstellen:
- Einen Konstruktor, welcher über optionale Parameter verfügt: die Dimension `n` und eine Liste von Integers `vec`. 
Falls beide Parameter angegeben sind, können Sie davon ausgehen, dass `n == len(vec)` gilt.
- `__setitem__(v, i, val)`: Ersetzt das Element $v_i$ des Vektors $v$ mit dem Wert `val`. Sie können davon ausgehen, dass `i` immer kleiner als die Dimension des Vektors ist.
- `__getitem__(v, i)`: Gibt das Element $v_i$ des Vektors $v$ zurück. Falls `i` größer oder gleich der Dimension `n` des Vektors ist, so soll `None` zurückgegeben werden.

Sobald Sie die Methoden `__setitem__(v, i, val)` und `__getitem__(v, i)` implementiert haben, sollten Sie auf die einzelnen Vektorelemente mittels Indizierung zugreifen können (siehe Tests). Sie können dies für die Implementierung der folgenden Methoden benutzen: 

- `add(v, a)`: Verändert den Vektor $v = (v_0,\ldots,v_{n-1})$ indem der übergebene Vektor $a = (a_0,\ldots,a_{n-1})$ zu $v$ addiert wird.
- `difference(v, a)`: Subtrahiert den Vektor $a = (a_0,\ldots,a_{n-1})$ von dem Vektor $v = (v_0,\ldots,v_{n-1})$ und liefert das Ergebnis als neuen Vektor $s = (v_0-a_0,\ldots,v_{n-1}-a_{n-1})$ zurück. Die beiden Vektoren `v` und `a` sollen dabei *nicht* verändert werden.
- `dot(v, w)`: Berechnet das Vektorprodukt von Vektor $v = (v_0,\ldots,v_{n-1})$ und dem übergebenen Vektor $w = (w_0,\ldots,w_{n-1})$ und liefert das Ergebnis ($\sum_{i=0}^{n-1} v_i\cdot w_i$) zurück.
- `norm(v)`: Berechnet die Norm ($\sqrt[]{\sum_{i=0}^{n-1} v_i^2}$) des Vektors $v = (v_0,\ldots,v_{n-1})$ und liefert diese zurück.


Für die Methoden `add`, `difference` und `dot` können Sie davon ausgehen, dass beide Vektoren immer dieselbe Länge haben.

Es ist Ihnen selbstverständlich erlaubt, zusätzliche Hilfsmethoden (wie z.B. `__str__`) zu implementieren, um Ihren Code besser auf Fehler überprüfen zu können.

<!-- grade_id: vector_sol, points: 0, grade: false, solution: true -->
```python
class Vector:
    ### BEGIN SOLUTION
    """Die Klasse Vektor realisiert einen mathematischen Vektor, auf welchem Vektoroperationen ausgeführt werden können."""
    
    def __init__(self, n=None, vec=None):
        """Erstelle leeren Vektor der Länge n. Falls Initiatiliserungsvektor oder -liste
        gegeben, kopiere dessen Inhalt."""

        self.vec = []
        if n:
            self.vec = [0] * n
        if vec:
            self.vec = [0] * len(vec)
            for i in range(len(vec)):
                self.vec[i] = vec[i]

    def add(self, a):
        """Addiere zu self, self verändern"""
        n = len(self)
        for i in range(n):
            self[i] = self[i] + a[i]
    
    def difference(self, a):
        """Subtrahiere von self, aber self nicht verändern"""
        erg = [0]*len(self)
        for i in range(0, len(a)):
            erg[i] = self.vec[i] - a[i]
        return Vector(vec = erg)
    
    def dot(self, w):
        """Berechnet das Vektorprodukt mit dem übergebenen Vektor"""
        return sum([x*o for x, o  in zip(self.vec, w.vec)])
    
    def norm(self):
        """Berechnet die Norm für den Vektor"""
        return sum([x**2 for x in self.vec])**0.5
    
    def __len__(self):
        return len(self.vec)
    
    def __setitem__(self, i, val):
        """Ersetzt das i-te Element des Vektors"""
        self.vec[i] = val
        
    def __getitem__(self, i):
        """Liefert das i-te Element des Vektors"""
        if i >= len(self.vec):
            return None
        else:
            return self.vec[i]
    
    def __str__(self):
        return str(self.vec)
    ### END SOLUTION
```

<!-- grade_id: vector_test_0, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Konstruktor-Test:
v0 = Vector(n=2)
assert(v0[0] == 0) # Zugriff hier mittels Indizierung!
assert(v0[1] == 0)

v1 = Vector(n=2, vec=[1,2])
assert(v1[0] == 1)
assert(v1[1] == 2)
assert(v1[2] == None)

v2 = Vector(vec=[1,2])
assert(v2[0] == 1)
assert(v2[1] == 2)

# Ersetze ein Element von v2:
v2[0] = 2
assert(v2[0] == 2)

### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    for i in range(0,n):
        assert(a[i] == i+1)
    a[0] = 3
    assert(a[0] == 3)
### HIDEEND
```

<!-- grade_id: vector_test_1, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
a = Vector(vec = [1, 2, 3])
b = Vector(vec = [3, 4, 5])

a.add(b) # Addiere b zu a
assert(a[0] == 4)
assert(a[1] == 6)
assert(a[2] == 8)

# b unverändert:
assert(b[0] == 3)
assert(b[1] == 4)
assert(b[2] == 5)

### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    b = Vector(vec = range(0,n))

    a.add(b)
    for i in range(1,n+1):
        assert(a[i-1] == (i + b[i-1]))
    for i in range(0,n):
        assert(b[i] == i)
### HIDEEND
```

<!-- grade_id: vector_test_2, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
a = Vector(vec = [1, 2, 3])
b = Vector(vec = [3, 4, 5])

s = a.difference(b) # Berechne die Differenz a-b und speichere den neuen Vektor in s
assert(s[0] == -2)
assert(s[1] == -2)
assert(s[2] == -2)

# a, b unverändert:
assert(a[0] == 1)
assert(a[1] == 2)
assert(a[2] == 3)
assert(b[0] == 3)
assert(b[1] == 4)
assert(b[2] == 5)

### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    b = Vector(vec = range(0,n))

    s = a.difference(b)

    for i in range(0,n):
        assert(s[i] == (a[i] - b[i]))
    for i in range(1,n):
        assert(a[i] == i+1)
    for i in range(0,n):
        assert(b[i] == i)
### HIDEEND
```

<!-- grade_id: vector_test_3, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
a = Vector(vec = [1, 2, 3])
b = Vector(vec = [3, 4, 5])
dot = a.dot(b) # Berechne das Vektorprodukt und speichere dies in dot
assert(dot == 26)

# a, b unverändert:
assert(a[0] == 1)
assert(a[1] == 2)
assert(a[2] == 3)
assert(b[0] == 3)
assert(b[1] == 4)
assert(b[2] == 5)

### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    b = Vector(vec = range(0,n))
    dot = 0
    for i in range(0,n):
        dot = dot + a[i] * b[i]
    assert(a.dot(b) == dot)
    for i in range(1,n):
        assert(a[i] == i+1)
    for i in range(0,n):
        assert(b[i] == i)
### HIDEEND
```

<!-- grade_id: vector_test_4, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
a = Vector(vec  = [1, 2, 3])
norm = a.norm() # Berechne die Norm und speichere diese in norm
assert(norm == (14**0.5))

# a unverändert:
assert(a[0] == 1)
assert(a[1] == 2)
assert(a[2] == 3)

### HIDESTART
for n in range(2,10):
    a = Vector(vec = range(1,n+1))
    norm = 0
    for i in range(0, n):
        norm = norm + a[i] ** 2
    norm = norm ** 0.5
    assert(a.norm() == norm)    
    for i in range(1,n):
        assert(a[i] == i+1)
### HIDEEND
```

# Aufgabe 4: PAUL
In dieser Aufgabe sollen Sie eine einfache Version von PAUL implementieren.
Das System soll aus 3 Klassen bestehen: `Student`, `Vorlesung` und `Paul`.
Die Klasse `Student` ist Ihnen bereits gegeben (siehe untere Codezelle).
Ihre Aufgabe ist nun, die Klassen `Vorlesung` und `Paul` zu vervollständigen.
Die Anforderungen an die Klassen sowie Tests sind in den jeweiligen Teilaufgaben gegeben.

```python
class Student:
    """Diese Klasse modelliert einen Studenten mit Attributen mat_nr und name."""
    
    def __init__(self, mat_nr, name):
        """Konstruktor

        Args:
            mat_nr: Matrikelnummer des Studenten (int)
            name: Name des Studenten (str)
        """
        self.mat_nr = mat_nr
        self.name = name
    
    def get_mat_nr(self):
        """Get-Methode für das Matrikelnummer-Attribut

        Returns:
            int Matrikelnummer des Studenten
        """
        return self.mat_nr
    
    def get_name(self):
        """Get-Methode für das Name-Attribut

        Returns:
            str Namen des Studenten
        """
        return self.name
    
    def __str__(self):
        """String-Repräsentation eines Studenten

        Returns:
            str Matrikelnummer und name des Studenten
        """
        return str(self.mat_nr) + " " + self.name
```

## a)
Erstellen Sie eine Klasse `Vorlesung`, welche über die folgenden Methoden verfügt:

- Einen Kontruktor, welcher die ID der Vorlesung (int), sowie den Namen der Vorlesung (str) übergeben bekommt.
- `get_id`: Gibt die ID der Vorlesung zurück.
- `get_name`: Gibt den Namen der Vorlesung zurück.
- `add_student`: Registriert den übergebenen Studenten in der Vorlesung.
- `has_student`: Gibt `True` zurück, falls sich ein Student mit der übergebenen Matrikelnummer in der Vorlesung registriert hat, ansonsten `False`. Sie dürfen davon ausgehen, dass sich jeder Student nur einmal für ein und dieselbe Vorlesung registriert.
- `add_grade`: Setzt die Note des Studenten mit der übergebenen Matrikelnummer auf die übergebene Note. Sollte der Student bereits eine Note haben, so wird diese überschrieben. Wenn der Student nicht in der Vorlesung registriert ist, soll keine Aktion ausgeführt werden.
- `get_grade`: Liefert die Note für diese Vorlesung, welche der Student mit der übergebenen Matrikelnummer erhalten hat. Falls die Matrikelnummer nicht in der Vorlesung registriert ist, soll `-1` zurückgegeben werden. Ist der Student zwar registriert, hat aber noch keine Note erhalten, wird `None` zurückgegeben.


<!-- grade_id: paul_sol_1, points: 0, grade: false, solution: true -->
```python
class Vorlesung:
    ### BEGIN SOLUTION
    """Eine Klasse mit Daten für eine Vorlesung"""
    
    def __init__(self, v_id, name):
        """Konstruktor

        Args:
            v_id: ID der Vorlesung (int)
            name: Name der Vorlesung (str)
        """
        self.v_id = v_id
        self.name = name
        self.studenten = {}
        
    def get_id(self):
        """Get-Methode für das v_id-Attribut

        Returns:
            int ID der Vorlesung
        """
        return self.v_id
    
    def get_name(self):
        """Get-Methode für das name-Attribut

        Returns:
            str Name der Vorlesung
        """
        return self.name
    
    def add_student(self, mat_nr):
        """Registriert Studenten mit der übergebenen Matrikelnummer in der Vorlesung
        
        Args:
            mat_nr: Die Matrikelnummer des Studenten (int)
        """
        self.studenten[mat_nr] = None
        
    def has_student(self, mat_nr):
        """Liefert True, falls sich ein Student mit der übergebenen Matrikelnummer in der Vorlesung registriert hat, 
        ansonsten False.
        
        Args:
            mat_nr: Matrikelnummer des zu suchenden Studenten (int)
            
        Returns:
            bool Suchergebnis als Wahrheitswert
        """
        return mat_nr in self.studenten
    
    def add_grade(self, mat_nr, grade):
        """Setzt die Note des Studenten mit der übergebenen Matrikelnummer auf die übergebene Note. 
        Sollte der Student bereits eine Note haben, so wird diese überschrieben. Wenn der Student nicht in der Vorlesung 
        registriert ist, soll keine Aktion ausgeführt werden.
        
        Args:
            mat_nr: Matrikelnummer des zu registrierenden Studenten (int)
            grade: Note, welche für den Studenten eingetragen werden soll (float)
        """
        if mat_nr in self.studenten:
            self.studenten[mat_nr] = grade

    def get_grade(self, mat_nr):
        """Liefert die Note für diese Vorlesung, welche der Student mit der übergebenen Matrikelnummer erhalten hat. 
        Falls die Matrikelnummer nicht in der Vorlesung registriert ist, soll -1 zurückgegeben werden. Ist der Student 
        zwar registriert, hat aber noch keine Note erhalten, wird None zurückgegeben.
        
        Returns:
            float Note des Studenten, None oder -1
        """
        if mat_nr in self.studenten:
            return self.studenten[mat_nr]
        return -1
    ### END SOLUTION
```

<!-- grade_id: paul_test_1_1, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Erstelle Studenten:
studenten = [Student(9863451, "Ralf B"), Student(9863452, "Jessica W"), Student(9863453, "Philipp K"),
             Student(9863454, "Andrea B"), Student(9863455, "Sebastian A"), Student(9863456, "Luca B")]

# Erstelle Vorlesung:
gp1 = Vorlesung(1, "Grundlagen der Programmierung 1")
assert(gp1.get_id() == 1)
assert(gp1.get_name() == "Grundlagen der Programmierung 1")

# Füge Studenten zur Vorlesung hinzu:
for s in studenten:
    gp1.add_student(s.get_mat_nr())
for s in studenten:
    assert(gp1.has_student(s.get_mat_nr()))

### HIDESTART
# Erstelle Vorlesung:
mod = Vorlesung(2, "Modellierung")
assert(mod.get_id() == 2)
assert(mod.get_name() == "Modellierung")

# Füge Studenten zur Vorlesung hinzu:
mod.add_student(studenten[0].get_mat_nr())
mod.add_student(studenten[1].get_mat_nr())
mod.add_student(studenten[2].get_mat_nr())

assert(mod.has_student(studenten[0].get_mat_nr()))
assert(mod.has_student(studenten[1].get_mat_nr()))
assert(mod.has_student(studenten[2].get_mat_nr()))
assert(not mod.has_student(studenten[3].get_mat_nr()))
assert(not mod.has_student(studenten[4].get_mat_nr()))
assert(not mod.has_student(studenten[5].get_mat_nr()))
### HIDEEND
```

<!-- grade_id: paul_test_1_2, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Verteile Noten:
gp1.add_grade(studenten[0].get_mat_nr(), 1.0)
gp1.add_grade(studenten[1].get_mat_nr(), 1.3)
gp1.add_grade(studenten[2].get_mat_nr(), 3.0)
gp1.add_grade(studenten[3].get_mat_nr(), 2.7)
gp1.add_grade(studenten[4].get_mat_nr(), 1.3)

assert(gp1.get_grade(studenten[0].get_mat_nr()) == 1.0)
assert(gp1.get_grade(studenten[1].get_mat_nr()) == 1.3)
assert(gp1.get_grade(studenten[2].get_mat_nr()) == 3.0)
assert(gp1.get_grade(studenten[3].get_mat_nr()) == 2.7)
assert(gp1.get_grade(studenten[4].get_mat_nr()) == 1.3)
assert(gp1.get_grade(studenten[5].get_mat_nr()) == None) # hat noch keine Note erhalten

# Ein nicht in der Vorlesung registrierter Student will eine Note haben:
gp1.add_grade(1234567, 1.0) 
assert(gp1.get_grade(1234567) == -1) # Matrikelnummer nicht im Kurs vorhanden

### HIDESTART
assert(mod.get_grade(studenten[3].get_mat_nr()) == -1)
assert(mod.get_grade(studenten[4].get_mat_nr()) == -1)
assert(mod.get_grade(studenten[5].get_mat_nr()) == -1)
### HIDEEND
```

## b)
Ergänzen Sie nun die Klasse `PAUL`, welche über die folgenden Methoden verfügt:
- `add_vorlesung`: Erstellt eine neue Vorlesung mit dem übergebenen Namen im System und liefert die ID der Vorlesung zurück. Die ID der Vorlesung ist gleich der Anzahl der bisherigen Aufrufe von `add_vorlesung` + 1, d.h. beim ersten Aufruf wird die ID `1` zurückgegeben (siehe Tests).
- `add_student`: Erstellt einen neuen Studenten mit dem übergebenen Namen im System und liefert dessen Matrikelnummer zurück. Die Matrikelnummer des Studenten ist gleich der Anzahl der bisherigen Aufrufe von `add_student` + 1, d.h. beim ersten Aufruf wird die Matrikelnummer `1` zurückgegeben (siehe Tests).
- `registriere_in_vorlesung`: Fügt den Studenten mit der übergebenen Matrikelnummer der Vorlesung mit der übergebenen ID hinzu. Sie können davon ausgehen, dass `registriere_student_in_vorlesung` niemals mehrfach für ein Paar (`matr_nr`, `v_id`) aufgerufen wird.
- `get_studenten_in_vorlesung`: Liefert eine Liste, welche alle Matrikelnummern enthält, die in der Vorlesung mit der übergebenen ID registriert sind. Falls das System keine Vorlesung mit der übergebenen ID enthält, soll `None` zurückgegeben werden.
- `get_vorlesung`: Liefert die Vorlesung mit der übergebenen ID als Objekt der Klasse `Vorlesung`. Ist keine solche Vorlesung im System gespeichert, soll `None` zurückgegeben werden.
- `get_student`: Liefert den Studenten mit der übergebenen Matrikelnummer als Objekt der Klasse `Student`. Ist kein solcher Student im System gespeichert, soll `None` zurückgegeben werden.
- `get_grades_for_student`: Liefert ein Dictionary mit allen Vorlesungen, in welcher der übergebene Student registriert ist und die entsprechenden Noten, die der Student für die Vorlesung erhalten hat.

Eine beispielhafte Anwendung können Sie an den gegebenen Tests ablesen.

<!-- grade_id: paul_sol_2, points: 0, grade: false, solution: true -->
```python
class Paul:
    ### BEGIN SOLUTION
    """ Modelliert eine Art Benutzerschnittstelle zur Erstellung bzw. Verwaltung von Studenten und Vorlesungen """
            
    def __init__(self):
        """ Konstruktor """
        self.studenten = []
        self.vorlesungen = []
        
    def add_vorlesung(self, name):
        """Erstellt eine neue Vorlesung mit dem übergebenen Namen im System und liefert die ID der Vorlesung zurück.

        Args:
            name: Name der Vorlesung (str)
            
        Returns:
            (int) ID der Vorlesung
        """
        v = Vorlesung(len(self.vorlesungen) + 1, name)
        self.vorlesungen.append(v)
        return len(self.vorlesungen)
        
    def add_student(self, name):
        """Erstellt einen neuen Studenten mit dem übergebenen Namen im System und liefert dessen Matrikelnummer zurück.

        Args:
            name: Name des Studenten (str)
            
        Returns:
            (int) Matrikelnummer des Studenten
        """
        s = Student(len(self.studenten) + 1, name)
        self.studenten.append(s)
        return len(self.studenten)
    
    def registriere_in_vorlesung(self, mat_nr, v_id):
        """Fügt den Studenten mit der übergebenen Matrikelnummer der Vorlesung mit der übergebenen ID hinzu. 
        Sie können davon ausgehen, dass registriere_student_in_vorlesung niemals mehrfach für ein Paar (matr_nr, v_id) 
        aufgerufen wird.

        Args:
            mat_nr: Matrikelnummer des zu registrierenden Studenten (int)
            name: ID der Vorlesung, in welcher der Student registriert werden soll (str)
        """
        if mat_nr <= len(self.studenten):
            self.vorlesungen[v_id - 1].add_student(mat_nr)
        
    def get_studenten_in_vorlesung(self, v_id):
        """Liefert eine Liste, welche alle Matrikelnummern enthält, die in der Vorlesung mit der übergebenen ID 
        registriert sind. Falls das System keine Vorlesung mit der übergebenen ID enthält, soll None zurückgegeben werden.

        Args:
            v_id: ID der Vorlesung (int)
        
        Returns:
            list Matrikelnummern aller Studenten, die in der Vorlesung registriert sind
        """
        if v_id <= len(self.vorlesungen):
            return list(self.vorlesungen[v_id - 1].studenten.keys())
        return None
    
    def get_vorlesung(self, v_id):
        """ Liefert die Vorlesung mit der übergebenen ID als Objekt der Klasse Vorlesung.
        
        Args:
            v_id: ID der Vorlesung, welche gesucht ist
        """
        if v_id <= len(self.vorlesungen):
            return self.vorlesungen[v_id - 1]
        return None
    
    def get_student(self, mat_nr):
        """ Liefert den Studenten mit der übergebenen Matrikelnummer als Objekt der Klasse Student.
        
        Args:
            mat_nr: Matrikelnummer des gesuchten Studenten
        """
        if mat_nr <= len(self.studenten):
            return self.studenten[mat_nr - 1]
        return None

    def get_grades_for_student(self, mat_nr):
        """Liefert ein Dictionary mit allen Vorlesungen, in welcher der übergebene Student registriert ist und 
        die entsprechenden Noten, die der Student für diese Vorlesungen erhalten hat.
        
        Args:
            mat_nr: Matrikelnummer des Studenten (int)
        
        Returns:
            dict Dictionary mit allen Noten für den gegebenen Studenten
        """
        d = {}
        for v in self.vorlesungen:
            if v.has_student(mat_nr):
                d[v.get_name()] = v.get_grade(mat_nr)
        return d
    ### END SOLUTION
```

<!-- grade_id: paul_test_2_1, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Erstelle System:
paul = Paul()

# Füge Vorlesungen und Studenten hinzu:
gp1_id = paul.add_vorlesung("Grundlagen der Programmierung 1")
mod_id = paul.add_vorlesung("Modellierung")
assert(gp1_id == 1)
assert(mod_id == 2)
matrikelnummern = [paul.add_student("Ralf B") ,paul.add_student("Jessica W") ,paul.add_student("Philipp K")]
assert(matrikelnummern[0] == 1 and matrikelnummern[1] == 2 and matrikelnummern[2] == 3)

assert(paul.get_student(matrikelnummern[0]).get_name() == 'Ralf B')
assert(paul.get_student(matrikelnummern[1]).get_name() == 'Jessica W')
assert(paul.get_student(matrikelnummern[2]).get_name() == 'Philipp K')

# Registriere Studenten in Vorlesung:
paul.registriere_in_vorlesung(matrikelnummern[0], 1)
paul.registriere_in_vorlesung(matrikelnummern[1], 1)
paul.registriere_in_vorlesung(matrikelnummern[2], 1)

paul.registriere_in_vorlesung(matrikelnummern[0], 2)
paul.registriere_in_vorlesung(matrikelnummern[1], 2)
paul.registriere_in_vorlesung(matrikelnummern[2], 2)

assert(1 in paul.get_studenten_in_vorlesung(1))
assert(2 in paul.get_studenten_in_vorlesung(1))
assert(3 in paul.get_studenten_in_vorlesung(1))

assert(1 in paul.get_studenten_in_vorlesung(2))
assert(2 in paul.get_studenten_in_vorlesung(2))
assert(3 in paul.get_studenten_in_vorlesung(2))

### HIDESTART
dua_id = paul.add_vorlesung("Datenstrukturen und Algorithmen")
assert(dua_id == 3)
andrea_id = paul.add_student("Andrea B")
assert(paul.get_student(andrea_id).get_name() == 'Andrea B')
assert(andrea_id == 4)
paul.registriere_in_vorlesung(andrea_id, 3)
assert(paul.get_studenten_in_vorlesung(3) == [andrea_id])
### HIDEEND
```

<!-- grade_id: paul_test_2_2, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Füge Noten hinzu:
paul.get_vorlesung(1).add_grade(matrikelnummern[0], 3.0)
paul.get_vorlesung(1).add_grade(matrikelnummern[1], 1.3)
paul.get_vorlesung(1).add_grade(matrikelnummern[2], 2.0)

paul.get_vorlesung(2).add_grade(matrikelnummern[0], 5.0)
paul.get_vorlesung(2).add_grade(matrikelnummern[1], 2.3)
assert(paul.get_grades_for_student(matrikelnummern[0]) == {'Grundlagen der Programmierung 1': 3.0, 'Modellierung': 5.0})
assert(paul.get_grades_for_student(matrikelnummern[1]) == {'Grundlagen der Programmierung 1': 1.3, 'Modellierung': 2.3})
assert(paul.get_grades_for_student(matrikelnummern[2]) == {'Grundlagen der Programmierung 1': 2.0, 'Modellierung': None})

### HIDESTART
paul.get_vorlesung(3).add_grade(andrea_id, 2.0)
assert(paul.get_grades_for_student(andrea_id) == {'Datenstrukturen und Algorithmen': 2.0})
### HIDEEND
```

<!-- grade_id: doc_test, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Alle Dokumentationen vorhanden? 

# Aufgabe 1
assert(s_verbinder.__doc__ != None)
assert(s_suche.__doc__ != None)
assert(split_3.__doc__ != None)

# Aufgabe 3
assert(Vector.__doc__ != None)
assert(Vector.__init__.__doc__ != None)
assert(Vector.__setitem__.__doc__ != None)
assert(Vector.__getitem__.__doc__ != None)
assert(Vector.add.__doc__ != None)
assert(Vector.difference.__doc__ != None)
assert(Vector.dot.__doc__ != None)
assert(Vector.norm.__doc__ != None)

# Aufgabe 4
assert(Vorlesung.__doc__ != None)
assert(Vorlesung.__init__.__doc__ != None)
assert(Vorlesung.get_id.__doc__ != None)
assert(Vorlesung.get_name.__doc__ != None)
assert(Vorlesung.add_student.__doc__ != None)
assert(Vorlesung.has_student.__doc__ != None)
assert(Vorlesung.add_grade.__doc__ != None)
assert(Vorlesung.get_grade.__doc__ != None)
assert(Paul.__doc__ != None)
assert(Paul.add_vorlesung.__doc__ != None)
assert(Paul.add_student.__doc__ != None)
assert(Paul.registriere_in_vorlesung.__doc__ != None)
assert(Paul.get_studenten_in_vorlesung.__doc__ != None)
assert(Paul.get_vorlesung.__doc__ != None)
assert(Paul.get_grades_for_student.__doc__ != None)
```