<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile). Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue07.ipynb ab (nicht umbenennen, keine Kopie erstellen!). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie  die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigenen Lösungen. 

## Kommentare (1 Punkt!)

Kommentieren Sie Code-Abgaben! Insbesondere muss jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Das Vorhandensein von ```docstrings``` auf allen Code-Abgaben wird mit einem Punkt gewertet!


# Abgabetermin: 18.12.2016 um 23:59 Uhr



# Aufgabe 1: String-API (3 Punkte)
Diese Aufgabe lässt sich eventuell einfacher mithilfe der Python-Dokumentation lösen:
https://docs.python.org/3/library/stdtypes.html.

### a) 
Schreiben Sie eine Funktion `s_verbinder(l,z)`, welches eine Liste mit Strings `l` zu einem einzigen String verbindet, wobei zwischen die einzelnen Listenelemente jeweils das Zeichen `z` eingefügt werden soll.

<!-- grade_id: string_sol_1, points: 0, grade: false, solution: true -->
```python
def s_verbinder(l,z):
    # Your Code here
    ### BEGIN SOLUTION
    """Methode, welche eine Liste mit Strings zu einem einzigen STring verbindet."""
    return z.join(l)
    ### END SOLUTION
```

<!-- grade_id: string_test_1, points: 1, grade: true, solution: false -->
```python
assert(s_verbinder(['1','2','3','4'], '-') == '1-2-3-4')
### HIDESTART
assert(s_verbinder(['a','b','c','d'], '-') == 'a-b-c-d')
### HIDEEND
```

### b) 
Schreiben Sie eine Funktion `s_suche(s1,s2)`, welche den niedrigsten Index ausgibt, an welcher der String `s2` in `s1` auftaucht.

<!-- grade_id: string_sol_2, points: 0, grade: false, solution: true -->
```python
def s_suche(s1,s2):
    # Your Code here
    ### BEGIN SOLUTION
    """Methode, welche den Niedrigsten Index ausgibt, an welchem der String s2 in s1 auftaucht."""
    return s1.find(s2)
    ### END SOLUTION
```

<!-- grade_id: string_test_2, points: 1, grade: true, solution: false -->
```python
assert(s_suche('Grundlagen der Programmierung', 'der') == 11)
assert(s_suche('Grundlagen der Programmierung', 'P') == 15)
### HIDESTART
assert(s_suche('Test', 'T') == 0)
assert(s_suche('Test', 'e') == 1)
assert(s_suche('Test', 's') == 2)
assert(s_suche('Test', 't') == 3)
assert(s_suche('Test', 'es') == 1)
assert(s_suche('TestT', 'T') == 0)
### HIDEEND
```

### c) 
Schreiben Sie eine Funktion `s_klein(s)`, welche alle Großbuchstaben im string `s` durch entsprechende Kleinbuchstaben ersetzt.

<!-- grade_id: string_sol_3, points: 0, grade: false, solution: true -->
```python
def s_klein(s):
    # Your Code here
    ### BEGIN SOLUTION
    """Methode, welche alle Großbuchstaben des übergebenen Strings durch entsprechende Kleinbuchstaben ersetzt."""
    return s.lower()
    ### END SOLUTION
```

<!-- grade_id: string_test_3, points: 1, grade: true, solution: false -->
```python
assert(s_klein('Hallo Welt!') == 'hallo welt!')
### HIDESTART
assert(s_klein('Dies ist ein geHEimer Satz.') == 'dies ist ein geheimer satz.')
### HIDEEND
```

# Aufgabe 2: Vektorklasse (5 Punkte)
In dieser Aufgabe sollen Sie eine Klasse ```Vektor``` implementieren.
Ein Vektor $v$ besteht aus $n$ einzelnen Zahlen, also $v = (v_0,\ldots,v_{n-1})$.
Das $n$ kann hier eine beliebige Zahl aus den natürlichen Zahlen $\mathbb N = \{1,2,3,\ldots\}$ sein.
Die Klasse ```Vektor``` soll für ein Objekt $v = (v_0,\ldots,v_{n-1})$ die folgenden Methoden bereitstellen:
- ```index(i)```: Liefert die Zahl an der $i$-ten Stelle des Vektors, d.h. für $v = (v_0,\ldots,v_{n-1})$ wird $v_i$ zurückgegeben.
- ```add(a)```: Addiert den Vektor $v = (v_0,\ldots,v_{n-1})$ mit dem übergebenen Vektor $a = (a_0,\ldots,a_{n-1})$ und liefert das Ergebnis als neuen Vektor $e = (v_0+a_0,\ldots,v_{n-1}+a_{n-1})$ zurück.
- ```substract(a)```: Subtrahiert den Vektor $a = (a_0,\ldots,a_{n-1})$ von dem Vektor $v = (v_0,\ldots,v_{n-1})$ und liefert das Ergebnis als neuen Vektor $s = (v_0-a_0,\ldots,v_{n-1}-a_{n-1})$ zurück.
- ```dot(w)```: Berechnet das Vektorprodukt für den Vektor $v = (v_0,\ldots,v_{n-1})$ und dem übergebenen Vektor $w = (w_0,\ldots,w_{n-1})$ und liefert das Ergebnis ($\sum_{i=0}^{n-1} v_i\cdot w_i$) zurück.
- ```norm()```: Berechnet die Norm ($\sqrt[]{\sum_{i=0}^{n-1} v_i^2}$) des Vektors $v = (v_0,\ldots,v_{n-1})$ und liefert diese zurück.

Für die Methoden ```add```, ```substract``` und ```dot``` können Sie davon ausgehen, dass beide Vektoren immer dieselbe Länge haben.

<!-- grade_id: vector_sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Vector(object):
    """Die Klasse Vektor realisiert einen mathematischen Vektor, auf welchem Vektoroperationen ausgeführt werden können."""
    def __init__(self, vec):
        self.vec = vec

    def add(self, v):
        erg = [0]*self.length()
        for i in range(0, v.length()):
            erg[i] = self.vec[i] + v.index(i)
        return Vector(erg)
    
    def subtract(self, v):
        erg = [0]*self.length()
        for i in range(0, v.length()):
            erg[i] = self.vec[i] - v.index(i)
        return Vector(erg)
    
    def dot(self, v):
        erg = 0
        for i in range(0, v.length()):
            erg = erg + self.vec[i] * v.index(i)
        return erg
    
    def norm(self):
        """norm-Methode der Klasse Vektor"""
        erg = 0
        for i in range(0, len(self.vec)):
            erg = erg + self.vec[i] ** 2
        erg = erg ** 0.5
        return erg

    def length(self):
        return len(self.vec)
    
    def index(self, i):
        return self.vec[i]
### END SOLUTION
```

<!-- grade_id: vector_test_0, points: 1, grade: true, solution: false -->
```python
a = Vector([4, 3, 7])
assert(a.index(0) == 4)
assert(a.index(1) == 3)
assert(a.index(2) == 7)

### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    for i in range(0,n):
        assert(a.index(i) == i+1)
### HIDEEND
```

<!-- grade_id: vector_test_1, points: 1, grade: true, solution: false -->
```python
a = Vector([1, 2, 3])
b = Vector([3, 4, 5])

e = a.add(b)
assert(e.index(0) == 4)
assert(e.index(1) == 6)
assert(e.index(2) == 8)

### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    b = Vector(range(0,n))

    e = a.add(b)

    for i in range(0,n):
        assert(e.index(i) == (a.index(i) + b.index(i)))
### HIDEEND
```

<!-- grade_id: vector_test_2, points: 1, grade: true, solution: false -->
```python
a = Vector([1, 2, 3])
b = Vector([3, 4, 5])
s = a.subtract(b)
assert(s.index(0) == -2)
assert(s.index(1) == -2)
assert(s.index(2) == -2)
### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    b = Vector(range(0,n))

    s = a.subtract(b)

    for i in range(0,n):
        assert(s.index(i) == (a.index(i) - b.index(i)))
### HIDEEND
```

<!-- grade_id: vector_test_3, points: 1, grade: true, solution: false -->
```python
a = Vector([1, 2, 3])
b = Vector([3, 4, 5])
dot = a.dot(b)
assert(dot == 26)
### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    b = Vector(range(0,n))
    dot = 0
    for i in range(0,n):
        dot = dot + a.index(i) * b.index(i)
    assert(a.dot(b) == dot)
### HIDEEND
```

<!-- grade_id: vector_test_4, points: 1, grade: true, solution: false -->
```python
a = Vector([1, 2, 3])
norm = a.norm()
assert(norm == (14**0.5))
### HIDESTART
for n in range(2,10):
    a = Vector(range(1,n+1))
    norm = 0
    for i in range(0, n):
        norm = norm + a.index(i) ** 2
    norm = norm ** 0.5
    assert(a.norm() == norm)    
### HIDEEND
```

# Aufgabe 3: Objektorientierung (5 Punkte)

Betrachten Sie für diese Aufgabe die gegebene Klasse `Patient`.

```python
class Patient():
    anz_patienten = 0
    
    def __init__(self, name, alter, privat_patient):
        self.name = name
        self.alter = alter
        self.privat_patient = privat_patient
        
    def get_name(self):
        return self.name
    
    def get_alter(s):
        return s.alter
    
    def is_privat_patient(self):
        return self.privat_patient
    
    def __eq__(self, p):
        return self.name == p.name and self.alter == p.alter
```

### a)
Erläutern Sie zunächst den Unterschied der Operationen '==' und 'is'. 
Weisen Sie dann den beiden gegebenen Variablen `p1` und `p2` Objekte der Klasse `Patient` zu, sodass `p1 == p2` gilt, jedoch nicht `p1 is p2`.

<!-- grade_id: patient_sol_1_1, points: 1, grade: true, solution: true -->
Der Operator `==` führt die Methode `__eq__` des Objektes aus, welches auf der linken Seite des Operators steht.
Der Operator `is` hingegen überprüft, ob die referenzierten Objekte auf beiden Seiten dieselben sind.

<!-- grade_id: patient_sol_1_2, points: 0, grade: false, solution: true -->
```python
p1 = None # None durch entsprechende Belegung ersetzen!
p2 = None # None durch entsprechende Belegung ersetzen!
### BEGIN SOLUTION
p1 = Patient("Max Meier", 26, False)
p2 = Patient ("Max Meier", 26, False)
### END SOLUTION
```

<!-- grade_id: patient_test_1, points: 1, grade: true, solution: false -->
```python
assert(isinstance(p1, Patient))
assert(isinstance(p2, Patient))
assert(p1 == p2)
assert(not p2 is p1)
```

### b)
Wie Sie sehen, verfügt die Klasse `Patient` auch über eine Klassenvariable `anz_patienten`.
Der Wert dieser Variablen soll nun immer um 1 erhöht werden, wenn ein neuer Patient erstellt wurde, d.h. nach Ausführung des folgenden Programms sollte `anz_patienten` den Wert 2 haben:

```Python
p1 = Patient("Max Meier", 26, False)
p2 = Patient ("Max Mustermann", 12, False)
```
An welcher Stelle muss dies realisiert werden?
Geben Sie auch den entsprechenden Befehl an, welcher dies realisiert?

<!-- grade_id: patient_sol_2, points: 1, grade: true, solution: true -->
Es muss im Konstruktor der Befehl 
``` Python
Patient.anz_patienten += 1
```
ausgeführt werden.

### c)
Wie Sie sehen, heißt der erste Parameter der Methode `get_alter()` nicht `self`.
Würde das folgende Programm einen Fehler erzeugen?
Wenn ja, an welcher Stelle und warum, wenn nein, erklären Sie warum nicht?

``` Python
p1 = Patient("Max Meier", 26, False)
print(p1.get_alter())
```

<!-- grade_id: patient_sol_3, points: 1, grade: true, solution: true -->
Das Programm funktioniert, da der Name für den ersten Parameter einer Objektmethode nur per Konvention `self` lautet.

### d)
Angenommen, die obige Klasse `Patient` würde über keinen Konstruktor verfügen.
Würde das folgende Programm einen Fehler erzeugen?
Wenn ja, an welcher Stelle und warum, wenn nein, erklären Sie warum nicht?
``` Python
p1 = Patient()
print(p1.get_name())
```

<!-- grade_id: patient_sol_4, points: 1, grade: true, solution: true -->
Zunächst einmal wird der fehlende Konstruktor von Python durch einen leere Standardkonstruktor ersetzt.
Das Programm wirft einen Fehler wenn versucht wird die Methode p1.get_name() aufzurufen, da der Patient über keine Objektvariable `name` verfügt (diese wurde im Standardkonstruktor nicht definiert).

# Aufgabe 3: Arztpraxis (6 Punkte)
Sie können in dieser Aufgabe die Klasse `Patient` (siehe Beginn Aufgabe 3) verwenden!   
Stellen Sie sich vor, Sie arbeiten in einer Arztpraxis und sollen ein System zur Verwaltung der Patienten im Wartezimmer implementieren.
Das System soll Patienten, die in die Praxis kommen, entweder in das Büro des Arztes schicken (falls das Büro gerade frei ist) oder ins Wartezimmer bitten.
Im Büro des Arztes darf sich maximal ein Patient gleichzeitig befinden (der Patient, der gerade vom Arzt behandelt wird).
Sobald das Büro des Arztes wieder frei wird (d.h. die Behandlung abgeschlossen ist), kann nun ein Patient im Wartezimmer aufgerufen und in das Büro des Arztes gebeten werden.
Für jeden Aufruf wird immer der Patient gewählt, der die längste Zeit im Wartezimmer verbracht hat.

### a)
Implementieren Sie eine Klasse ```Arztpraxis```, welche die folgenden Methoden bereitstellt:
- ```betretePraxis(p)```: Der gegebene Patient ```p``` betritt die Praxis: Falls sich momentan kein anderer Patient in Behandlung befindet, wird ```p``` direkt zum Arzt geschickt, ansonsten wird ```p``` in das Wartezimmer gebeten.
- ```naechsterPatientZumArzt()```: Entfernt den Patienten aus dem Wartezimmer, welcher die längste Zeit dort verbracht hat, schickt diesen zum Arzt und gibt diesen zurück. Falls das Wartezimmer leer ist oder sich noch ein Patient in Behandlung befindet, werden keine Aktionen ausgeführt und lediglich ```None``` zurückgegeben.
- ```beimArzt()```: Gibt den Patienten zurück, welcher sich gerade beim Arzt befindet. Falls sich kein Patient gerade beim Arzt befindet, wird ```None``` zurückgegeben.
- ```behandlungAbgeschlossen()```: Der Patient, der sich aktuell in Behandlung befindet verlässt die Praxis, sodass das Zimmer des Arztes nun wieder leer ist.

<!-- grade_id: praxis_2_sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Arztpraxis:
    
    def __init__(self):
        self.warteschlange = []
        self.inBehandlung = None
        
    def betretePraxis(self, p):
        # print(p.get_name(), 'hat die Praxis betreten.')
        if self.inBehandlung is None and len(self.warteschlange) == 0: # Achtung: self.inBehandlung == None wäre hier falsch!
            self.inBehandlung = p
            # print(p.get_name(), 'ist nun in Behandlung.')
        else:
            self.warteschlange.append(p)
    
    def naechsterPatientZumArzt(self):
        if not self.beimArzt() is None:
            return None
        if len(self.warteschlange) == 0:
            return None
        p = self.warteschlange.pop(0)
        self.inBehandlung = p
        # print(p.get_name(), 'ist nun in Behandlung.')
        return p
    
    def beimArzt(self):
        return self.inBehandlung
    
    def behandlungAbgeschlossen(self):
        # print(self.inBehandlung.get_name(), 'hat die Praxis verlassen.')
        self.inBehandlung = None
    
### END SOLUTION    
```

<!-- grade_id: praxis_2_test, points: 1, grade: true, solution: false -->
```python
# Disclaimer: Die folgende Anwendung spiegelt nicht die gesamte Funktionalität der Arztpraxis wieder. 
# Beachten Sie daher genau die in der Aufgabenstellung beschriebenen Funktionsweisen der Methoden.

praxis = Arztpraxis()
p1 = Patient("Max Meier", 26, False)
p2 = Patient("Max Mustermann", 13, False)
p3 = Patient("Michaela Musterfrau", 20, False)
praxis.betretePraxis(p2)
praxis.betretePraxis(p1)
praxis.betretePraxis(p3)

assert(praxis.beimArzt() == p2)

praxis.behandlungAbgeschlossen()

assert(praxis.beimArzt() == None)

praxis.naechsterPatientZumArzt()

assert(praxis.beimArzt() == p1)
```

<!-- grade_id: praxis_2_test_hidden, points: 2, grade: true, solution: false -->
```python
### HIDESTART
praxis = Arztpraxis()
p1 = Patient("Max Meier", 26, False)
p2 = Patient("Max Mustermann", 13, False)
p3 = Patient("Michaela Musterfrau", 20, False)
p4 = Patient("Michael F", 87, False)
praxis.betretePraxis(p2)
praxis.betretePraxis(p1)
praxis.betretePraxis(p3)

assert(praxis.beimArzt() == p2)

praxis.behandlungAbgeschlossen()
praxis.betretePraxis(p4)

assert(praxis.beimArzt() == None)

praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == p1)

praxis.behandlungAbgeschlossen()
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == p3)

praxis.behandlungAbgeschlossen()
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == p4)

praxis.behandlungAbgeschlossen()
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == None)
### HIDEEND
```

Unsere Arztpraxis soll nun zwischen Kassen- und Privatpatienten unterscheiden. Privatpatienten werden von der Praxis bevorzugt behandelt, d.h. für die Auswahl des als nächstes zu behandelden Patienten wird immer erst geprüft, ob sich ein Privatpatient im Wartezimmer befindet.
Ist dies der Fall so wird dieser den Kassenpatienten vorgezogen, unabhängig von der Zeit, die er im Wartezimmer verbracht hat.
Falls sich mehrere Privatpatienten im Wartezimmer befinden, so wird der sich am längsten im Wartezimmer befindliche Privatpatient als nächstes aufgerufen.

d) Erweitern Sie nun Ihre Klasse ```Arztpraxis``` aus Aufgabenteil b) zu einer Klasse ```ArztpraxisUnfair```, sodass diese nun Privatpatienten wie oben beschrieben bevorzugt behandelt.

<!-- grade_id: praxis_4_sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
class Arztpraxis:

    def __init__(self):
        self.warteschlange = []
        self.inBehandlung = None
        self.privatSchlange = []
        
    def betretePraxis(self, p):
        # print(p.getName(), 'hat die Praxis betreten.')
        if self.inBehandlung is None and len(self.warteschlange) == 0 and len(self.privatSchlange) == 0:
            # print(p.getName(), 'ist nun in Behandlung.')
            self.inBehandlung = p
        elif p.is_privat_patient():
            self.privatSchlange.append(p)
        else:
            self.warteschlange.append(p)
    
    def naechsterPatientZumArzt(self):
        if self.beimArzt() != None:
            return None
        if len(self.privatSchlange) == 0 and len(self.warteschlange) == 0:
            return None
        
        if len(self.privatSchlange) > 0:
            p = self.privatSchlange.pop(0)
        else:
            p = self.warteschlange.pop(0)
        self.inBehandlung = p
        # print(p.getName(), 'ist nun in Behandlung.')
        return p
    
    def beimArzt(self):
        return self.inBehandlung
    
    def behandlungAbgeschlossen(self):
        # print(self.inBehandlung.get_name(), 'hat die Praxis verlassen.')
        self.inBehandlung = None    
### END SOLUTION
```

<!-- grade_id: praxis_4_test, points: 2, grade: true, solution: false -->
```python
praxis = Arztpraxis()
p1 = Patient("Max Meier", 26, False)
p2 = Patient("Max Mustermann", 13, False)
p3 = Patient("Michaela Musterfrau", 20, True)

praxis.betretePraxis(p2)
praxis.betretePraxis(p1)
praxis.betretePraxis(p3)

assert(praxis.beimArzt() == p2)
praxis.behandlungAbgeschlossen()

praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == p3)
```

<!-- grade_id: praxis_3_test_hidden, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Derselbe Test wie in Teil a, um zu sehen, ob die Arztpraxis noch ihre normalen Funktionen beibehalten hat.
praxis = Arztpraxis()
p1 = Patient("Max Meier", 26, False)
p2 = Patient("Max Mustermann", 13, False)
p3 = Patient("Michaela Musterfrau", 20, False)
p4 = Patient("Michael F", 87, False)
praxis.betretePraxis(p2)
praxis.betretePraxis(p1)
praxis.betretePraxis(p3)

assert(praxis.beimArzt() == p2)

praxis.behandlungAbgeschlossen()
praxis.betretePraxis(p4)

assert(praxis.beimArzt() == None)

praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == p1)

praxis.behandlungAbgeschlossen()
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == p3)

praxis.behandlungAbgeschlossen()
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == p4)

praxis.behandlungAbgeschlossen()
praxis.naechsterPatientZumArzt()
assert(praxis.beimArzt() == None)
### HIDEEND
```

<!-- grade_id: docu_test, points: 1, grade: true, solution: false -->
```python
### HIDESTART
assert(s_verbinder.__doc__ != None)
assert(s_suche.__doc__ != None)
assert(s_klein.__doc__ != None)
# TODO: Kann man testen, ob Klassen (z.B. die Klasse Vektor) bzw. Objektmethoden dokumentiert wurden?
### HIDEEND
```