<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

<!--- Wenn Sie diese Ansicht verwenden (was Sie nicht tun sollten; klicken Sie am besten einmal auf "Run cell"), dann steht die Formatierung hier richtig:

Format:

MatrikelNummer E-Mail-Adresse
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
234567 testStudent2@mail.upb.de
345678 vorname.nachname@mail.upb.de

... und nicht hier: -->

Format:

MatrikelNummer E-Mail-Adresse
<br \>
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
<br \>
234567 testStudent2@mail.upb.de
<br \>
345678 vorname.nachname@mail.upb.de

Hinweis: Wenn Sie das so in die obere Zelle eingeben, wird in der ausgeführten Markdownzelle (Strg+Enter) alles in einer Zeile angezeigt. Das ist egal; wichtig ist Formatierung der Eingabe.


## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile), für das entsprechende Übungsblatt! (Wenn Sie Blatt 9 abgeben wolllen hilft es nicht, bei Blatt 1 auf Submit zu drücken!) Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue09.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen, auch nicht Dateiname hue9.ipynb !). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen sinnvollen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 22.01.2017 um 23:59 Uhr

# Aufgabe 1:

## a)
Betrachten Sie die folgenden beiden Codevarianten in Python:
```Python 
#Variante A:
for i in [x for x in range(100)]:
    print(i, end=', ')
    
#Variante B:
for i in (x for x in range(100)):
    print(i, end=', ')

```
Erläutern Sie den Unterschied zwischen den beiden Varianten unter Verwendung der richtigen Fachbegriffe.

<!-- grade_id: a1-sol, points: 1, grade: true, solution: true -->
Bei Variante A handelt es sich um eine List Comprehension, bei welcher zunächst die gesamte Liste (mit Elementen von 0 bis 99) erzeugt und anschließend darüber iteriert wird.

Bei Variante B handelt es sich um eine Generator Expression, bei welcher nicht im Vorhinein eine ganze Liste erzeugt wird, sondern es wird bei jeder Iteration das aktuelle Element erzeugt.

## b)

Im obigen Beispiel führen beide Varianten zu demselben Ergebnis.
Nennen Sie Fälle, in denen Sie auf jeden Fall ```[x for ...]``` verwenden müssen und dieses nicht einfach durch ```(x for ...)``` ersetzen können.

<!-- grade_id: a1b-sol, points: 1, grade: true, solution: true -->
Wenn ich eine Liste erzeugen möchte, die ich anschließend noch bearbeiten, weitergeben kann, etc., muss man die List Comprehension-Syntax verwenden.

## c)
In welchen Fällen würden Sie auf jeden Fall ```(x for ...)``` anstelle von ```[x for ...]``` verwenden müssen?

<!-- grade_id: a1c-sol, points: 1, grade: true, solution: true -->
Wenn die Anzahl der Elemente, die mit diesen Ausdrücken erzeugt wird, zu groß für den Speicher ist (z. B. unendlich), müsste man die Generator Expression verwenden.

# Aufgabe 2:

## a)
Schreiben Sie eine Funktion namens ```count_occurences(s1, s2)``` die für zwei übergebene String-Parameter ```s1``` und ```s2``` zurückgibt, wie oft ```s2``` in ```s1``` vollständig enthalten ist.

Hinweis: Sie brauchen nicht zu überprüfen, ob es sich bei ```s1``` und ```s2``` um Strings handelt.

TODO: Diese Aufgabe wird vermutlich gestrichen (warte noch auf Rückmeldung). War ursprünglich als map-reduce-Beispiel gedacht.

<!-- grade_id: a2b-sol, points: 0, grade: false, solution: true -->
```python
def count_occurences(s1, s2):
    pass
```

<!-- grade_id: a2b-test, points: 1, grade: true, solution: false -->
```python

```

## b)

Schreiben Sie eine Funktion namens ```transform_dict(d1, d2)```, welches bei Eingabe zweier `dict`s `d1` und `d2` ein neues `dict` zurückgibt, dessen Schlüssel dieselben sind wie die Schlüssel von `d1`, dessen Werte jedoch ersetzt sind durch die Werte, welche in d2 unter dem Schlüssel des bisherigen Wertes gespeichert sind.

Seien z. B. die folgenden beiden `dict`s gegeben:
```Python
matrikelnummer_zu_note = { 35135135: 5, 60103513: 3, 10981179: 2, 18637724 : 4 }
note_zahl_zu_string = { 1: "summa cum laude", 2: "magna cum laude", 3: "cum laude", 4: "rite", 5: "non probatum" }
```
Dann soll ein Aufruf von ```transform_dict(matrikelnummer_zu_note, note_zahl_zu_string)``` folgendes `dict` zurückgeben:
```Python
{60103513: 'cum laude', 10981179: 'magna cum laude', 18637724: 'rite', 35135135: 'non probatum'}
```

**Hinweis:** Verwenden Sie hierzu die Funktion ```map()```.

<!-- grade_id: a2c-sol, points: 0, grade: false, solution: true -->
```python
def transform_dict(d1, d2):
    ### BEGIN SOLUTION
    """Gibt bei Eingabe zweier dicts d1 und d2 ein neues dict zurück,
    dessen Schlüssel dieselben sind wie die Schlüssel von d1, 
    dessen Werte jedoch ersetzt sind durch die Werte, 
    welche in d2 unter dem Schlüssel des bisherigen Wertes gespeichert sind.
    
    Args:
        d1: das erste dict, dessen Schlüssel den Schlüsseln des zurückgegebenen dicts entsprechen (dict)
        d2: das zweite dict, dessen Werte in das zurückgegebene dict eingefügt werden (dict)
        
    Returns:
        dict: das Ziel-dict, wie oben definiert
    """
    return dict(map(lambda k: (k, d2[d1[k]]), d1))
    ### END SOLUTION
```

<!-- grade_id: a2c-test, points: 1, grade: true, solution: false -->
```python
matrikelnummer_zu_note = {35135135: 5, 60103513: 3, 10981179: 2, 18637724 : 4}
note_zahl_zu_string = {1: "summa cum laude", 2: "magna cum laude", 3: "cum laude", 4: "rite", 5: "non probatum"}

assert transform_dict(matrikelnummer_zu_note, note_zahl_zu_string) == {60103513: 'cum laude', 10981179: 'magna cum laude', 18637724: 'rite', 35135135: 'non probatum'}


d1 = {"a": "b", 5: 2, True: False}
d2 = {"o": "a", 10: 5, "b" : True}

assert transform_dict(d2, d1) == {'o': 'b', 10: 2, 'b': False}

### HIDESTART
import builtins

class MyException(Exception):
    pass

def map(*args, **kwargs):
    raise MyException()
        
try:
    transform_dict(matrikelnummer_zu_note, note_zahl_zu_string)
except MyException:
    used_map = True
else:
    used_map = False
finally:
    map = builtins.map
    assert used_map
### HIDEEND
```

## c)

Schreiben Sie eine Funktion ```kleiner_gleich(d, wert)```, die bei Übergabe eines dicts `d` und einer Zahl `wert` ein ```filter```-Objekt zurückgibt, welches ein Iterator auf die Schlüssel aus `d` ist, deren Werte kleiner oder gleich ```wert``` sind.

Beispielsweise soll ```kleiner_gleich(matrikelnummer_zu_note, 3)``` (mit `dict` ```matrikelnummer_zu_note``` aus der letzten Aufgabe) ein ```filter```-Objekt zurückgeben, welches die Werte `60103513` und `10981179` liefert.

Hinweis: Verwenden Sie hierzu die Funktion ```filter()```. Diese gibt ein Objekt vom Typ ```filter``` zurück.

<!-- grade_id: a2d-sol, points: 1, grade: false, solution: true -->
```python
def kleiner_gleich(d, wert):
    """Gibt für ein übergebenes dict d und einer übergebenen Zahl wert alle Schlüssel von d zurück,
    deren Wert kleiner oder gleich wert sind.
    
    Args:
        dict: Das dictionary (dict)
        wert: der Wert mit dem verglichen wird (zahl)
    
    Returns:
        filter-Objekt mit den gewünschten Schlüsseln
    """
    return filter(lambda k: d[k] <= wert, d)
```

<!-- grade_id: a2d-test, points: 1, grade: true, solution: false -->
```python
matrikelnummer_zu_note = {35135135: 5, 60103513: 3, 10981179: 2, 18637724 : 4}
mindestens_befriedigend = kleiner_gleich(matrikelnummer_zu_note, 3)
assert type(mindestens_befriedigend) == filter
s = {10981179, 60103513}
for i in mindestens_befriedigend:
    assert i in s
    s.remove(i)
    
assert len(s) == 0


d1 = {"a": 15, "c": 29, True: 13, 15 : 12}
res = set(kleiner_gleich(d1, 25))
assert set(res) == {True, 'a', 15} and len(list(res)) == 3
#TODO: Kann mir irgendwer sagen, warum das nicht funktioniert?


### HIDESTART
def filter(*args, **kwargs):
    raise MyException()
       
try:
    kleiner_gleich(matrikelnummer_zu_note, 3)
except MyException:
    used_filter = True
else:
    used_filter = False
finally:
    filter = builtins.filter
    assert used_filter
### HIDEEND
```

## d)
Schreiben Sie eine Funktion ```kleiner_gleich_gen```, welche dasselbe leistet wie die Funktion aus Teilaufgabe c), jedoch nicht die Funktion ```filter()```, sondern die ```(x for ...)```-Syntax (s. Aufgabe 1) verwendet und anstelle eines ```filter```-Objekts einen ```generator``` zurückgibt.

<!-- grade_id: a2e-sol, points: 1, grade: false, solution: true -->
```python
def kleiner_gleich_gen(d, wert):
    """Gibt für ein übergebenes dict d und einer übergebenen Zahl wert alle Schlüssel von d zurück,
    deren Wert kleiner oder gleich wert sind.
    
    Args:
        dict: Das dictionary (dict)
        wert: der Wert mit dem verglichen wird (zahl)
    
    Returns:
        generator-Objekt mit den gewünschten Schlüsseln
    """    
    return (k for (k,v) in d.items() if v <= wert)
```

<!-- grade_id: a2e-test, points: 1, grade: true, solution: false -->
```python
import types

matrikelnummer_zu_note = { 35135135: 5, 60103513: 3, 10981179: 2, 18637724 : 4 }

mindestens_befriedigend_gen = kleiner_gleich_gen(matrikelnummer_zu_note, 3)

assert isinstance(mindestens_befriedigend_gen, types.GeneratorType)

s = {10981179, 60103513}

for i in mindestens_befriedigend_gen:
    assert i in s
    s.remove(i)
    
assert len(s) == 0


d1 = {"a": 15, "c": 29, True: 13, 15 : 12}
res = set(kleiner_gleich_gen(d1, 25))
print(res)
print({True, 'a', 15})
assert set(res) == {True, 'a', 15} #and len(list(res)) == 3
#TODO: Siehe oben
```

Output:
```
{True, 'a', 15}
{True, 'a', 15}
```

## e)
Gegeben sei eine Liste von Listen von Zahlen, z. B. so:
```Python
L = [[1, 2, 3], [11, 13], [9, 9, 0], [13, 0, 15, 16]]
```

Schreiben Sie eine Funktion namens ```no_zeros(lst)```, welche für eine übergebene solche Liste ```lst``` eine neue Liste von Listen von Zahlen zurückgibt, welche der übergebenen Liste gleicht mit dem Unterschied, dass alle Nullen herausgefiltert (also entfernt) wurden.
Verwenden Sie hierzu `map()` und `filter()`.

Im Beispiel hieße das, dass der Aufruf ```no_zeros(L)``` die Liste ```[[1, 2, 3], [11, 13], [9, 9], [13, 15, 16]]``` zurückgeben würde.

Hinweis: Sie brauchen nicht zu überprüfen, ob eine Liste im richtigen Format übergeben wurde. Beachten Sie aber, dass sowohl die äußere als auch jede der innerem Listen leer sein können.

Hinweis: Die Aufgabe lässt sich mit einem Ausruck erledigen, der ```map``` und ```filter``` mit geeigneten ```lambda```-Ausdrücken nutzt. 

<!-- grade_id: a2x-sol, points: 0, grade: false, solution: true -->
```python
def no_zeros(lst):
    """ Erzeugt für eine übergebene Liste von Listen eine neue Liste, die der übergebenen gleicht, 
    mit dem Unterschied, dass alle Nullen daraus entfernt wurden.
    
    Args:
        lst: Eine Liste von Listen (list)
        
    Returns:
        eine neue Liste (innere Listen von lst um Nullen bereinigt)
    
    """
    return list(map (lambda d: list(filter(lambda x: x> 0, d)), lst))
```

<!-- grade_id: a2x-test, points: 1, grade: true, solution: false -->
```python
L1 = [[1, 2, 3], [11, 13], [9, 9, 0], [13, 0, 15, 16]]
assert no_zeros(L1) == [[1, 2, 3], [11, 13], [9, 9], [13, 15, 16]] 

L2 = [[], [0, 0], [9, 9, 0], [13, 0, 15, 16]]
assert no_zeros(L2) == [[], [], [9, 9], [13, 15, 16]] 

### HIDESTART
import builtins

class MyException(Exception):
    pass

def map(*args, **kwargs):
    raise MyException()
      
try:
    no_zeros(L1)
except MyException:
    used_map = True
else:
    used_map = False
finally:
    map = builtins.map
    assert used_map

    
def filter(*args, **kwargs):
    raise MyException()
       
try:
    no_zeros(L1)
except MyException:
    used_filter = True
else:
    used_filter = False
finally:
    filter = builtins.filter
    assert used_filter
### HIDEEND
```

# Aufgabe 3:

## a)
Gegeben ist der folgende Python-Code:
```python
u = 1

def method1(x):
    def method2(n):
        def method3(m):
            return x + n + m
        def method4():
            global u
            u += 1
        method4()
        return method3
    def method5(y):
        return x + y
    return method2, method5            
```
Ersetzen Sie so viele ```def```-Anweisungen wie möglich durch entsprechende ```lambda```-Ausdrücke, ohne die Semantik zu verändern.

<!-- grade_id: a3a-ans, points: 1, grade: true, solution: true -->
```python
u = 1

def method1(x):
    def method2(n):
        def method4():
            global u
            u += 1
        method4()
        return lambda m : x + n + m

    return method2, lambda y : x + y
```

## b)

Schreiben Sie eine Funktion ```brief_generator(...)```, so dass folgender Code:
```python
brief_max = brief_generator("Dies ist der Kopfteil", "Dies ist der Fußteil")("Max Mustermann")
brief1 = brief_max("Dies ist ein Betreff", "Dies ist ein Brieftext")
print(brief1)
```
folgendes ausgibt:

```
Dies ist der Kopfteil

Max Mustermann

Dies ist der Betreff

Dies ist ein Brieftext

Dies ist der Fußteil
```

Verwenden Sie hierzu Closures, aber noch keine Decorators.

Sie müssen hier (ausnahmsweise) keine ```docstring```s ergänzen!

<!-- grade_id: a3b-sol, points: 0, grade: false, solution: true -->
```python
def brief_generator(kopfteil, fussteil):
    """Funktion, der man zwei Strings kopfteil und fussteil übergeben kann,
    die an eine Funktion gebunden werden, welche zurückgegeben wird. 
    Diese Funktion erhält wiederum einen Parameter absender, welcher an einer Funktion
    gebunden wird, welche diese Funktion zurückgibt.
    Diese Funktion wiederum gibt kopfteil, absender, betreff, text und fussteil, jeweils
    getrennt durch Leerzeilen, aus.
    
    Args:
        kopfteil: eine Zeichenkette (str)
        fussteil: eine Zeichenkette (str)
        
    Returns:
        Funktion wie oben beschrieben (function)
    """
    def _mit_name(absender):
        def _mit_inhalt(betreff, text):
            return kopfteil + "\n\n" + absender + "\n\n" + betreff + "\n\n" + text + "\n\n" + fussteil
        return _mit_inhalt
    return _mit_name
```

<!-- grade_id: a3b-test, points: 1, grade: true, solution: false -->
```python
brief_max = brief_generator("Dies ist der Kopfteil", "Dies ist der Fußteil")("Max Mustermann")
brief1 = brief_max("Dies ist ein Betreff", "Dies ist ein Brieftext")
assert brief1 == "Dies ist der Kopfteil\n\nMax Mustermann\n\nDies ist ein Betreff\n\nDies ist ein Brieftext\n\nDies ist der Fußteil"

### HIDESTART
kopf_fuss = brief_generator("Kopf", "Fuss")
mit_name = kopf_fuss("Name")
brief2 = mit_name("Betreff", "Text")
assert brief2 == "Kopf\n\nName\n\nBetreff\n\nText\n\nFuss"
### HIDEEND
```

## c)
Betrachten Sie das Programm in der nachfolgenden Zelle und ergänzen Sie vor der Definition der Methode ```brief(...)``` Decorator-Annotationen, so dass die Variable ```brief``` am Ende den gleichen Inhalt hat wie in der vorhergehenden Aufgabe (bzw. die Tests funktionieren).

Hinweis: Lesen Sie nach, wie man in Python Parameter an Dekoratoren übergeben kann und vollziehen Sie anhand dessen den bereits gegebenen Code-Teil nach. 
Das Verständnis dieses Teils wird Ihnen bei einer späteren Aufgabe von Nutzen sein.

Auch bei dieser Aufgabe dürfen Sie (ausnahmsweise) auf das Angeben von ```docstring```s verzichten.

<!-- grade_id: a3c-sol, points: 0, grade: false, solution: true -->
```python
def kopf_und_fuss(kopfteil, fussteil):
    def kopf_und_fuss_decorator(func):
        def inner(*args, **kwargs):
            return kopfteil + "\n\n" + func(*args, **kwargs) + "\n\n" + fussteil
        return inner
    return kopf_und_fuss_decorator
        
def name_block(absender):
    def name_block_decorator(func):
        def inner(*args, **kwargs):
            return absender + "\n\n" + func(*args, **kwargs)
        return inner
    return name_block_decorator
        
### BEGIN SOLUTION
@kopf_und_fuss("Dies ist der Kopfteil", "Dies ist der Fußteil")
@name_block("Max Mustermann")
### END SOLUTION
def brief(betreff, text):
    return betreff + "\n\n" + text
```

<!-- grade_id: a3c-test, points: 1, grade: true, solution: false -->
```python
brief_string = brief("Dies ist ein Betreff", "Dies ist ein Brieftext")
assert brief_string == "Dies ist der Kopfteil\n\nMax Mustermann\n\nDies ist ein Betreff\n\nDies ist ein Brieftext\n\nDies ist der Fußteil"

### HIDESTART
brief_string2 = brief("Betreff", "Text")
assert brief_string2 == "Dies ist der Kopfteil\n\nMax Mustermann\n\nBetreff\n\nText\n\nDies ist der Fußteil"
### HIDEEND
```

## d)
Vergleichen Sie Ihre Lösung von Aufgabe c) mit der von Aufgabe d): Welche Variante halten Sie in passenden Anwendungsszenarien für sinnvoller? Begründen Sie kurz.

<!-- grade_id: a3d-sol, points: 1, grade: true, solution: true -->
Bei der Lösung von Aufgabe d) ist der String für den Kopf- und Fußteil sowie für den Namen fest an die Funktion brief gebunden - sie lässt sich nicht mehr Nutzen, um z. B. Briefe mit anderen Absendern zu schreiben.
Bei Aufgabe c) ist dies noch flexibel anpassbar. Man kann auch zunächst nur Kopf- und Fußteil festlegen und die dadurch zurückgegebene Funktion dann mit verschiedenen Absendern aufrufen.

## e)
Gegeben sei die Funktion 
```Python
def gruesse(name):
    return "Ich grüße " + name + "!"
```

Diese soll nun mit Hilfe von Dekoratoren so erweitert werden, dass der Rückgabewert noch in `div`- und `p`-HTML-Tags gepackt wird.
Das heißt,
```Python
gruesse("dich")
```
soll
```Python
<div><p>Ich grüße dich!</p></div>
```
zurückgeben.

Schreiben Sie dazu einen Dekorator namens `html_tag`, dem Sie einen Parameter übergeben können und wenden Sie diesen auf die Funktion ```gruesse``` (mehrfach, mit verschiedenen aktuellen Parametern) an.
Der zu übergebende Parameter soll dabei nur der Name des HTML-Tags (z. B. `"div"` oder `"p"`) sein, nicht z. B. `"<div>"` oder `"<p>"`.

Hinweis: Schreiben Sie den Dekorator so allgemein, dass er auch andere Funktionen (mit möglicherweise anderen Parametern) dekorieren kann.

<!-- grade_id: a3e-sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def html_tag(tag_name):
    """Dekorator, welcher die von der hiermit dekorierterten Funktion zurückgegebenen Strings in
    öffnende und schließende HTML-Tags verpackt. 
    Die HTML-Tags, mit denen dekoriert werden soll, müssen als Parameter an den Dekorator übergeben werden.
    """
    def html_tag_decorator(func):
        def inner(*args, **kwargs):
            return "<{0}>{1}</{0}>".format(tag_name, func(*args, **kwargs))
        return inner
    return html_tag_decorator


@html_tag("div")
@html_tag("p")
### END SOLUTION
def gruesse(name):
    return "Ich grüße " + name + "!"
```

<!-- grade_id: a3e-test, points: 1, grade: true, solution: false -->
```python
assert gruesse("dich") == "<div><p>Ich grüße dich!</p></div>"

### HIDESTART
@html_tag("html")
def test_func():
    return "TEST"

assert test_func() == "<html>TEST</html>"

@html_tag("html")
def test_func2(a, b):
    return a + " TEST " + b

assert test_func2("1", "2") == "<html>1 TEST 2</html>"
### HIDEEND
```

## f)
Schreiben Sie einen Decorator namens `call_count`, welcher in einer Objektvariable namens `calls` zählt, wie oft die Funktion, die er dekoriert, aufgerufen wurde.

Wenden Sie diesen anschließend auf die unten vorgegebene Funktion `f()` an.

Hinweis: Der Decorator soll auf Funktionen mit beliebiger Art und Anzahl von Parametern anwendbar sein.

Tipp: Die Variable `calls` sollte eine Objektvariable der inneren Funktion der Dekoratorfunktion sein.

<!-- grade_id: a3f-sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def call_count(func):
    """Dekorator zum Zaehlen der Anzahl der Funktionsaufrufe (wird in <Funktionsname>.calls gespeichert)."""
    
    def inner(*args, **kwargs):
        inner.calls += 1        
        return func(*args, **kwargs)
   
    inner.calls = 0

    return inner        

@call_count
### END SOLUTION
def f():
    pass
```

<!-- grade_id: a3f-test, points: 1, grade: true, solution: false -->
```python
for i in range(15):
    f()
    
assert f.calls == 15

### HIDESTART
@call_count
def g():
    pass

for i in range(10):
    g()
    
assert g.calls == 10
### HIDEEND
```

# Aufgabe 4

## a)
Betrachten Sie den Code in der folgenden Python-Zelle:

```python
import time 
import functools

def benchmark(fct):
  def _benchmark(*args, **kwargs):
      tstart = time.time()
      r = fct(*args, **kwargs)
      tend = time.time()

      print("{} ({}, {}) verbrauchte {} Millisekunden".format(
          fct.__name__, args, kwargs, (tend-tstart)*1000))

      return r
      
  return _benchmark

def _fib_old(num):
    if num < 2:
        return num
    else:
        return _fib_old(num-1) + _fib_old(num-2) 

@benchmark
def fib_old(num):
    return _fib_old(num)


@functools.lru_cache(maxsize=100)
def _fib_new(num):
    if num < 2:
        return num
    else:
        return _fib_new(num-1) + _fib_new(num-2) 
    
@benchmark    
def fib_new(num):
    return _fib_new(num)

fib1 = fib_old(32)
fib2 = fib_new(32)
```

Output:
```
fib_old ((32,), {}) verbrauchte 1123.6979961395264 Millisekunden
fib_new ((32,), {}) verbrauchte 0.03409385681152344 Millisekunden
```

Erklären Sie, wieso die Ausführung der letzten beiden Code-Zeilen zu unterschiedlichen Ausgaben führt.

Schlagen Sie dazu in der Python-Referenz die Dokumentation von `functools.lru_cache` nach.

<!-- grade_id: a4a2-sol, points: 1, grade: true, solution: true -->
Die Berechnung der Fibonacci-Zahlen findet hier rekursiv in den Funktionen `_fib_old` bzw. `_fib_new` statt. 
Bei der hier gewählten Implementierung werden üblicherweise Ergebnisse doppelt berechnet (so ruft `fib(10)` `fib(8)` und `fib(7)` auf und der Aufruf `fib(8)` ruft wiederum `fib(7)` auf.
Die Funktion `fib_new` hat eine kürzere Laufzeit, da bei den rekursiven Aufrufen von `_fib_new` Ergebnisse, die bereits bei anderen Funktionsaufrufen berechnet wurden, aus einem Cache geholt werden, welchen der Dekorator `functools.lru_cache` verwaltet.
Dieser Cache speichert für die Funktion für jeden Parameter, mit dem sie aufgerufen wurde, den Rückgabewert des Funktionsaufrufs.
Die Größe dieses Caches wird durch den Dekorator-Parameter `maxsize` bestimmt (welcher hier mit 100 die Anzahl der verschiedenen Funktionsaufrufe deutlich übersteigt).

# b)

In welchen Fällen ist eine Verwendung von functools.lru_cache (als Decorator) ratsam? Wann ist davon abzuraten?

<!-- grade_id: a4b-sol, points: 1, grade: true, solution: true -->
Ratsam ist die Verwendung des Dekorators, wenn die Funktion, die er dekoriert, mehrfach mit demselben Parameter aufgerufen wird. 
Allerdings ist von seiner Verwendung abzuraten, wenn die Funktion, die mit ihm dekoriert wird, Seiteneffekte hat (also z. B. eine `print`-Anweisung enthält oder globale Variablen verändert) und es somit nicht nur auf den Rückgabewert der Funktion ankommt.

# c)
Betrachten Sie den folgenden Python-Code und die entstehende Ausgabe, wenn Sie den Code ausführen:

```python
import time 
import functools

def benchmark(fct):
  def _benchmark(*args, **kwargs):
      tstart = time.time()
      r = fct(*args, **kwargs)
      tend = time.time()

      print("{} ({}, {}) verbrauchte {} Millisekunden".format(
          fct.__name__, args, kwargs, (tend-tstart)*1000))

      return r
      
  return _benchmark

@benchmark
def f1(num):
    if num < 2:
        return num
    else:
        return f1(num-1) + f1(num-2)



@benchmark
@functools.lru_cache(maxsize=100)
def f2(num):
    if num < 2:
        return num
    else:
        return f1(num-1) + f1(num-2)
    



@functools.lru_cache(maxsize=100)
@benchmark
def f3(num):
    if num < 2:
        return num
    else:
        return f3(num-1) + f3(num-2)

num = 13

print(f1(num))
print ("------------")
print(f2(num))
print ("------------")
print(f3(num))
```

Output:
```
f1 ((1,), {}) verbrauchte 0.00095367431640625 Millisekunden
f1 ((0,), {}) verbrauchte 0.00095367431640625 Millisekunden
f1 ((2,), {}) verbrauchte 0.125885009765625 Millisekunden
f1 ((1,), {}) verbrauchte 0.0 Millisekunden
f1 ((3,), {}) verbrauchte 0.14495849609375 Millisekunden
f1 ((1,), {}) verbrauchte 0.0 Millisekunden
f1 ((0,), {}) verbrauchte 0.0 Millisekunden
f1 ((2,), {}) verbrauchte 0.015020370483398438 Millisekunden
f1 ((4,), {}) verbrauchte 0.17595291137695312 Millisekunden
f1 ((1,), {}) verbrauchte 0.0 Millisekunden
f1 ((0,), {}) verbrauchte 0.0 Millisekunden
f1 ((2,), {}) verbrauchte 0.015020370483398438 Millisekunden
f1 ((1,), {}) verbrauchte 0.0 Millisekunden
f1 ((3,), {}) verbrauchte 0.028848648071289062 Millisekunden
f1 ((5,), {}) verbrauchte 0.2319812774658203 Millisekunden
f1 ((1,), {}) verbrauchte 0.00095367431640625 Millisekunden
f1 ((0,), {}) verbrauchte 0.00095367431640625 Millisekunden
f1 ((2,), {}) verbrauchte 0.015974044799804688 Millisekunden
f1 ((1,), {}) verb
```

Die Funktionen f1, f2 und f3 sind (bis auf ihre Dekoratoren) gleich.
Erklären Sie, wie die unterschiedlichen Ausgaben bei den drei Funktionsaufrufen zustande kommen.

Hinweis: Beziehen Sie sich hierbei bitte nicht auf die konkreten Zahlen, sondern auf die unterschiedliche Anzahl der Ausgaben.

<!-- grade_id: a4c-sol, points: 1, grade: true, solution: true -->
Die rekursive Funktion `f1` wird lediglich mit dem Dekorator `benchmark` ausgestattet.
Dadurch wird für jeden Aufruf dieser Funktion die Laufzeit dieses Aufrufs gemessen.
Die Anzahl der Aufrufe ist hier am größten, da jeder der rekursiven Aufrufe wirklich ausgeführt wird (und somit Funktionswerte der Fibonacci-Funktion "doppelt berechnet" werden).

Die rekursive Funktion `f2` wird mit den Dekoratoren `benchmark` und `functools.lru_cache` ausgestattet, wobei in jedem (rekursiven) Funktionsaufruf zuerst die dekorierende Funktion von `benchmark` zuerst ausgeführt wird und danach erst die von `functools.lru_cache`.
Dies hat zur Folge, dass wegen der Aufrufe in der Zeile `return f1(num-1) + f1(num-2)` zwar die Benchmark-Ausgaben von beiden Funktionsaufrufen auf die Konsole geschrieben werden.
Die aufgerufenen Funktionen werden aber nicht wirklich ausgeführt, sondern deren Ergebnis aus dem Cache geholt (sofern der Wert schon einmal berechnet wurde).
Damit finden dann auch keine weiteren rekursiven Aufrufe aufgrund solcher bereits zuvor berechneten Ergebnisse statt.

Bei `f3` schließlich ist das Verhalten zwar im Grunde so wie bei `f2`, aufgrund der geänderten Reihenfolge der Dekoratoren finden aber Ausgaben für Funktionsaufrufe, deren Ergebnisse ohnehin aus dem Cache geholt werden, nicht mehr statt.
Daher gibt es hier nur eine Ausgabe pro zu berechnendem Funktionswert (welcher jeder der Funktionswerte von 0 bis 13 ist).

<!-- grade_id: docstring-test, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Alle Dokumentationen vorhanden? 

l = []

# Aufgabe 2:
# TODO: evtl. a)
l.append("transform_dict")
l.append("kleiner_gleich")
l.append("kleiner_gleich_gen")
l.append("no_zeros")



# Aufgabe 3
l.append("html_tag")
l.append("call_count")

for f in l:
    try:
        func = eval(f)
        assert(func.__doc__ != None)
        print(func.__doc__ != None, "-->", f)    
    except AttributeError:
        print("Not Implemented", "-->", f)
    except AssertionError:
        print(func.__doc__ != None, "-->", f)
```

Output:
```
True --> transform_dict
True --> kleiner_gleich
True --> kleiner_gleich_gen
True --> no_zeros
True --> html_tag
True --> call_count
```