<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

<!--- Wenn Sie diese Ansicht verwenden (was Sie nicht tun sollten; klicken Sie am besten einmal auf "Run cell"), dann steht die Formatierung hier richtig:

Format:

MatrikelNummer E-Mail-Adresse
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
234567 testStudent2@mail.upb.de
345678 vorname.nachname@mail.upb.de

... und nicht hier: -->

Format:

MatrikelNummer E-Mail-Adresse
<br \>
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
<br \>
234567 testStudent2@mail.upb.de
<br \>
345678 vorname.nachname@mail.upb.de

Hinweis: Wenn Sie das so in die obere Zelle eingeben, wird in der ausgeführten Markdownzelle (Strg+Enter) alles in einer Zeile angezeigt. Das ist egal; wichtig ist Formatierung der Eingabe.


## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile), für das entsprechende Übungsblatt! (Wenn Sie Blatt 9 abgeben wolllen hilft es nicht, bei Blatt 1 auf Submit zu drücken!) Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue09.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen, auch nicht Dateiname hue9.ipynb !). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 22.01.2017 um 23:59 Uhr

# Aufgabe 1:

## a)
Betrachten Sie die folgenden beiden Codevarianten in Python:
```Python 
#Variante A:
for i in [x for x in range(100)]:
    print(i, end=', ')
    
#Variante B:
for i in (x for x in range(100)):
    print(i, end=', ')

```
Erläutern Sie den Unterschied zwischen den beiden Varianten unter Verwendung der richtigen Fachbegriffe.

## b)
In welchen Fällen würden Sie auf jeden Fall ```[x for ...]``` anstelle von ```(x for ...)``` verwenden?

## c)
In welchen Fällen würden Sie auf jeden Fall ```(x for ...)``` anstelle von ```[x for ...]``` verwenden?

# Aufgabe 2:

## a)
Gegeben sei eine Liste von Listen von Zahlen, z. B. so:
```Python
L = [[1, 2, 3], [11, 13], [9, 9, 0], [13, 0, 15, 16]]
```

Schreiben Sie eine Funktion namens ```no_zeros(lst)```, welche für eine übergebene solche Liste ```lst``` eine neue Liste von Listen von Zahlen zurückgibt, welche der übergebenen Liste gleicht mit dem Unterschied, dass alle Nullen herausgefiltert (also entfernt) wurden.

Im Beispiel hieße das, dass der Aufruf ```no_zeros(L)``` die Liste ```[[1, 2, 3], [11, 13], [9, 9], [13, 15, 16]]``` zurückgeben würde.

Hinweis: Sie brauchen nicht zu überprüfen, ob eine Liste im richtigen Format übergeben wurde. Beachten Sie aber, dass sowohl die äußere, als auch innere Listen leer sein können.

<!-- grade_id: a2a-sol, points: 0, grade: false, solution: true -->
```python
def no_zeros(lst):
    return list(map (lambda d: list(filter(lambda x: x> 0, d)), lst))
```

## b)
Schreiben Sie eine Funktion namens ```count_occurences(s1, s2)``` die für zwei übergebene String-Parameter ```s1``` und ```s2``` zurückgibt, wie oft ```s2``` in ```s1``` vollständig enthalten ist.

Hinweis: Sie brauchen nicht zu überprüfen, ob es sich bei ```s1``` und ```s2``` um Strings handelt.

## c)

Schreiben Sie eine Funktion namens ```transform_dict(d1, d2)```, welches bei Eingabe zweier `dict`s `d1` und `d2` ein neues `dict` zurückgibt, dessen Schlüssel dieselben sind wie die Schlüssel von `d1`, dessen Werte jedoch ersetzt sind durch die Werte, welche in d2 unter dem Schlüssel des bisherigen Wertes gespeichert sind.

Seien z. B. die folgenden beiden `dict`s gegeben:
```Python
matrikelnummer_zu_note = { 35135135: 5, 60103513: 3, 10981179: 2, 18637724 : 4 }
note_zahl_zu_string = { 1: "summa cum laude", 2: "magna cum laude", 3: "cum laude", 4: "rite", 5: "non probatum" }
```
Dann soll ein Aufruf von ```transform_dict(matrikelnummer_zu_note, note_zahl_zu_string)``` folgendes `dict` zurückgeben:
```Python
{60103513: 'cum laude', 10981179: 'magna cum laude', 18637724: 'rite', 35135135: 'non probatum'}
```

**Hinweis:** Verwenden Sie hierzu die Funktion ```map()```.

<!-- grade_id: a2c-sol, points: 0, grade: false, solution: true -->
```python
def transform_dict(d1, d2):
    ### BEGIN SOLUTION
    """Gibt bei Eingabe zweier dicts d1 und d2 ein neues dict zurück,
    dessen Schlüssel dieselben sind wie die Schlüssel von d1, 
    dessen Werte jedoch ersetzt sind durch die Werte, 
    welche in d2 unter dem Schlüssel des bisherigen Wertes gespeichert sind.
    Args:
        d1: das erste dict, dessen Schlüssel den Schlüsseln des zurückgegebenen dicts entsprechen (dict)
        d2: das zweite dict, dessen Werte in das zurückgegebene dict eingefügt werden (dict)
        
    Returns:
        dict: das Ziel-dict, wie oben definiert
    """
    return dict(map(lambda k: (k, d2[d1[k]]), d1))
    ### END SOLUTION
```

<!-- grade_id: a2c-test, points: 1, grade: true, solution: false -->
```python
matrikelnummer_zu_note = { 35135135: 5, 60103513: 3, 10981179: 2, 18637724 : 4 }
note_zahl_zu_string = { 1: "summa cum laude", 2: "magna cum laude", 3: "cum laude", 4: "rite", 5: "non probatum" }

assert transform_dict(matrikelnummer_zu_note, note_zahl_zu_string) == {60103513: 'cum laude', 10981179: 'magna cum laude', 18637724: 'rite', 35135135: 'non probatum'}
```

## d)

Schreiben Sie eine Funktion ```kleiner_gleich(d, wert)```, die bei Übergabe eines dicts `d` und einer Zahl `wert` einen ```iterator``` zurückgibt, welche die Schlüssel aus `d` zurückgibt, deren Wert kleiner oder gleich ```wert``` ist.

Beispielsweise soll ```kleiner_gleich(matrikelnummer_zu_note, 3)``` (mit `dict` ```matrikelnummer_zu_note``` aus der letzten Aufgabe) einen ```iterator``` zurückgeben, welcher die Werte `60103513` und `10981179` liefert.

Hinweis: Verwenden Sie hierzu die Funktion ```filter()```.

<!-- grade_id: a2d-sol, points: 1, grade: false, solution: true -->
```python
def kleiner_gleich(d, wert):
    ### BEGIN SOLUTION
    return filter(lambda k: d[k] <= wert, d)
    ### END SOLUTION
```

<!-- grade_id: a2d-test, points: 1, grade: true, solution: false -->
```python
matrikelnummer_zu_note = { 35135135: 5, 60103513: 3, 10981179: 2, 18637724 : 4 }
mindestens_befriedigend = kleiner_gleich(matrikelnummer_zu_note, 3)
assert type(mindestens_befriedigend) == filter
s = {10981179, 60103513}
for i in mindestens_befriedigend:
    assert i in s
    s.remove(i)
    
assert len(s) == 0
```

## e)
Schreiben Sie eine Funktion ```kleiner_gleich_gen```, welche dasselbe leistet wie die Funktion aus Teilaufgabe d), jedoch nicht die Funktion ```filter()```, sondern die ```(x for ...)```-Syntax (s. Aufgabe 1) verwendet und anstelle eines ```iterators``` einen ```generator``` zurückgibt.

<!-- grade_id: a2e-sol, points: 1, grade: false, solution: true -->
```python
def kleiner_gleich_gen(d, wert):
    ### BEGIN SOLUTION
    return (k for (k,v) in d.items() if v <= wert)
    ### END SOLUTION
```

<!-- grade_id: a2e-test, points: 1, grade: true, solution: false -->
```python
import types

matrikelnummer_zu_note = { 35135135: 5, 60103513: 3, 10981179: 2, 18637724 : 4 }

mindestens_befriedigend_gen = kleiner_gleich_gen(matrikelnummer_zu_note, 3)

assert isinstance(mindestens_befriedigend_gen, types.GeneratorType)

s = {10981179, 60103513}

for i in mindestens_befriedigend_gen:
    assert i in s
    s.remove(i)
    
assert len(s) == 0
```

# Aufgabe 3:

## a)
Gegeben ist der folgende Python-Code:
```python
u = 1

def method1(x):
    def method2(n):
        def method3(m):
            return x + n + m
        def method4():
            global u
            u += 1
        method4()
        return method3
    def method5(y):
        return x + y
    return method2, method5            
```
Ersetzen Sie so viele ```def```-Anweisungen wie möglich durch entsprechende ```lambda```-Ausdrücke, ohne die Semantik zu verändern.

<!-- grade_id: a3a-ans, points: 1, grade: true, solution: true -->
```python
u = 1

def method1(x):
    def method2(n):
        def method4():
            global u
            u += 1
        method4()
        return lambda m : x + n + m

    return method2, lambda y : x + y
```

## b)

Schreiben Sie eine Funktion ```brief_generator(...)```, so dass folgender Code:
```python
brief_max = brief_generator("Dies ist der Kopfteil", "Dies ist der Fußteil")("Max Mustermann")
brief1 = brief_max("Dies ist ein Betreff", "Dies ist ein Brieftext")
print(brief1)
```
folgendes ausgibt:

```
Dies ist der Kopfteil

Max Mustermann

Dies ist der Betreff

Dies ist ein Brieftext

Dies ist der Fußteil
```

Verwenden Sie hierzu Closures, aber noch keine Decorators.

<!-- grade_id: a3b-sol, points: 0, grade: false, solution: true -->
```python
def brief_generator(kopfteil, fussteil):
    #TODO docstring
    def _mit_name(name):
        def _mit_inhalt(betreff, text):
            return kopfteil + "\n\n" + name + "\n\n" + betreff + "\n\n" + text + "\n\n" + fussteil
        return _mit_inhalt
    return _mit_name
```

<!-- grade_id: a3b-test, points: 0, grade: true, solution: false -->
```python
brief_max = brief_generator("Dies ist der Kopfteil", "Dies ist der Fußteil")("Max Mustermann")
brief1 = brief_max("Dies ist ein Betreff", "Dies ist ein Brieftext")
assert brief1 == "Dies ist der Kopfteil\n\nMax Mustermann\n\nDies ist ein Betreff\n\nDies ist ein Brieftext\n\nDies ist der Fußteil"
```

## c)
Betrachten Sie das Programm in der nachfolgenden Zelle und ergänzen Sie vor der Definition der Methode ```brief(...)``` Decorator-Annotationen, so dass die Variable ```brief``` am Ende den gleichen Inhalt hat wie in der vorhergehenden Aufgabe (bzw. die Tests funktionieren).

Hinweis: Lesen Sie nach, wie man in Python Parameter an Dekoratoren übergeben kann und vollziehen Sie anhand dessen den bereits gegebenen Code-Teil nach. 
Das Verständnis dieses Teils wird Ihnen bei einer späteren Aufgabe von Nutzen sein.

<!-- grade_id: a3c-sol, points: 0, grade: false, solution: true -->
```python
def kopf_und_fuss(kopfteil, fussteil):
    def kopf_und_fuss_decorator(func):
        def inner(*args, **kwargs):
            return kopfteil + "\n\n" + func(*args, **kwargs) + "\n\n" + fussteil
        return inner
    return kopf_und_fuss_decorator
        
def name_block(name):
    def name_block_decorator(func):
        def inner(*args, **kwargs):
            return name + "\n\n" + func(*args, **kwargs)
        return inner
    return name_block_decorator
        
### BEGIN SOLUTION
@kopf_und_fuss("Dies ist der Kopfteil", "Dies ist der Fußteil")
@name_block("Max Mustermann")
### END SOLUTION
def brief(betreff, text):
    return betreff + "\n\n" + text
```

<!-- grade_id: a3c-test, points: 1, grade: true, solution: false -->
```python
brief_string = brief("Dies ist ein Betreff", "Dies ist ein Brieftext")
assert brief_string == "Dies ist der Kopfteil\n\nMax Mustermann\n\nDies ist ein Betreff\n\nDies ist ein Brieftext\n\nDies ist der Fußteil"
```

## d)
Vergleichen Sie Ihre Lösung von Aufgabe c) mit der von Aufgabe d): Welche Variante halten Sie in möglichen Anwendungsszenarien für sinnvoller? Begründen Sie kurz.

<!-- grade_id: a3d-sol, points: 1, grade: true, solution: true -->
```python
#TODO
```

## e)
Gegeben sei die Funktion 
```Python
def gruesse(name):
    return "Ich grüße " + name + "!"
```

Diese soll nun mit Hilfe von Dekoratoren so erweitert werden, dass der Rückgabewert noch in `div`- und `p`-HTML-Tags gepackt wird.
Das heißt,
```Python
gruesse("dich")
```
soll
```Python
<div><p>Ich grüße dich!</p></div>
```
zurückgeben.

Schreiben Sie dazu einen Dekorator namens `html_tag`, dem Sie einen Parameter übergeben können und wenden Sie diesen auf die Funktion (mehrfach, mit verschiedenen aktuellen Parametern) an.
Der zu übergebende Parameter soll dabei nur der Name des HTML-Tags (z. B. `"div"` oder `"p"`) sein, nicht z. B. `"<div>"` oder `"<p>"`

<!-- grade_id: a3e-sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def html_tag(tag_name):
    def html_tag_decorator(func):
        def inner(name):
            return "<{0}>{1}</{0}>".format(tag_name, func(name))
        return inner
    return html_tag_decorator


@html_tag("div")
@html_tag("p")
### END SOLUTION
def gruesse(name):
    return "Ich grüße " + name + "!"
```

<!-- grade_id: a3e-test, points: 0, grade: true, solution: false -->
```python
assert gruesse("dich") == "<div><p>Ich grüße dich!</p></div>"
```

Schreiben Sie einen Decorator namens `call_count`, welcher zählt, wie oft die Funktion, die er dekoriert, aufgerufen wurde.

Wenden Sie diesen anschließend auf die unten vorgegebene Funktion `f()` an.

Hinweis: Der Decorator soll auf Funktionen mit beliebiger Art und Anzahl von Parametern anwendbar sein.

<!-- grade_id: a3f-sol, points: 0, grade: false, solution: true -->
```python
#TODO: hier muss noch einiges anders...

### BEGIN SOLUTION
def call_count(func):
    a = 0
    
    def inner(*args, **kwargs):
        nonlocal a #das muss noch anders
        a += 1        
        func(*args, **kwargs)
    
    return inner
        
@call_count
### END SOLUTION
def f(param):
    print("f wurde aufgerufen mit Parameter ", param)
    
for i in range(15):
    f(i)
```

Output:
```
f wurde aufgerufen mit Parameter  0
f wurde aufgerufen mit Parameter  1
f wurde aufgerufen mit Parameter  2
f wurde aufgerufen mit Parameter  3
f wurde aufgerufen mit Parameter  4
f wurde aufgerufen mit Parameter  5
f wurde aufgerufen mit Parameter  6
f wurde aufgerufen mit Parameter  7
f wurde aufgerufen mit Parameter  8
f wurde aufgerufen mit Parameter  9
f wurde aufgerufen mit Parameter  10
f wurde aufgerufen mit Parameter  11
f wurde aufgerufen mit Parameter  12
f wurde aufgerufen mit Parameter  13
f wurde aufgerufen mit Parameter  14
```