<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

<!--- Wenn Sie diese Ansicht verwenden (was Sie nicht tun sollten; klicken Sie am besten einmal auf "Run cell"), dann steht die Formatierung hier richtig:

Format:

MatrikelNummer E-Mail-Adresse
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
234567 testStudent2@mail.upb.de
345678 vorname.nachname@mail.upb.de

... und nicht hier: -->

Format:

MatrikelNummer E-Mail-Adresse
<br \>
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
<br \>
234567 testStudent2@mail.upb.de
<br \>
345678 vorname.nachname@mail.upb.de

Hinweis: Wenn Sie das so in die obere Zelle eingeben, wird in der ausgeführten Markdownzelle (Strg+Enter) alles in einer Zeile angezeigt. Das ist egal; wichtig ist Formatierung der Eingabe.


## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile), für das entsprechende Übungsblatt! (Wenn Sie Blatt 11 abgeben wolllen hilft es nicht, bei Blatt 1 auf Submit zu drücken!) Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue11.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen sinnvollen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 29.01.2017 um 23:59 Uhr

# Aufgabe 1:

## a)
Betrachten Sie den folgenden Python-Code:

```Python 
class Tier:
    def friss(self):
        print("Das Tier frisst.")
    
class Pflanzenfresser(Tier):
    def friss(self):
        print("Das Tier frisst vegetarisch.")

class Fleischfresser(Tier):
    def friss(self):
        print("Das Tier frisst Fleisch.")
       
class Allesfresser(Pflanzenfresser, Fleischfresser):
    pass

baer = Allesfresser()
baer.friss()
```

Welche Ausgabe wird erzeugt und warum?

<!-- grade_id: a1-sol, points: 2, grade: true, solution: true -->
Es wird "Das Tier isst vegetarisch" ausgegeben.
Das liegt daran, dass die Methode `iss()` von `Pflanzenfresser` aufgerufen wird, da diese Oberklasse als erstes in der Liste der Oberklassen aufgeführt ist.

## b)
Betrachten Sie das folgende, komplexere, Beispiel:

```Python
class Tier:
    def friss(self):
        print("Das Tier frisst.")

class Pflanzenfresser(Tier):
    def friss(self):
        print("Das Tier frisst vegetarisch.")
        
class Fleischfresser(Tier):
    def extra_portion(self):
        print("Der Fleischfresser frisst noch eine Portion.")
        
class Allesfresser(Pflanzenfresser, Fleischfresser):
    def friss(self):
        print("Der Allesfresser frisst.")
        self.extra_portion()

class Raubtier(Fleischfresser):
    def extra_portion(self):
        print("Das Raubtier frisst noch eine Portion.")        

class Baer(Allesfresser, Raubtier): pass

baer = Baer()
baer.friss()
```

Welche Ausgabe erzeugt dieser Code? Erklären Sie!

<!-- grade_id: a1b-sol, points: 1, grade: true, solution: true -->
Es wird die Ausgabe 

```Der Allesfresser frisst.
Das Raubtier frisst noch eine Portion.
```

erzeugt.

Dies liegt an der Method Resolution Order, welche die Klasse Raubtier vor Fleischfresser enthält (aufgrund des C3-Algorithmus).
Würde die Klasse Fleischfresser vor Raubtier in der MRO vorkommen, wäre die Monotonie verletzt.

## c)
Betrachten Sie nun die folgende Abwandlung des Beispiels aus Aufgabe a):

```Python 
class Tier:
    def friss(self):
        print("Das Tier frisst.")
    
class Pflanzenfresser(Tier):
    def __init__(self):
        print("Pflanzenfresser erstellt.")

    def friss(self):
        print("Das Tier frisst vegetarisch.")

class Fleischfresser(Tier):
    def __init__(self):
        print("Fleischfresser erstellt.")
        
    def friss(self):
        print("Das Tier frisst Fleisch.")
       
class Allesfresser(Pflanzenfresser, Fleischfresser):
    def __init__(self):
        super().__init__()
        print("Allesfresser erstellt.")

baer = Allesfresser()
```

Kopieren Sie den Code in die folgende Zelle und führen Sie ihn aus.

Wie Sie sehen, werden lediglich die Ausgaben der Konstruktoren der Klassen `Allesfresser` und `Pflanzenfresser` aufgerufen.

Wandeln Sie den Code (in der folgenden Zelle) so ab, dass auch der Konstruktor der Klasse `Fleischfresser` aufgerufen wird.
Vermeiden Sie dabei unnötige Wiederholungen.

<!-- grade_id: a1c-sol, points: 2, grade: true, solution: true -->
```python
### BEGIN SOLUTION

class Tier:
    def friss(self):
        print("Das Tier frisst.")
    
class Pflanzenfresser(Tier):
    def __init__(self):
        super().__init__()
        print("Pflanzenfresser erstellt.")

    def friss(self):
        print("Das Tier frisst vegetarisch.")

class Fleischfresser(Tier):
    def __init__(self):
        super().__init__()    # Zeile für Aufgabenerfüllung nicht erforderlich, aber guter Stil
        print("Fleischfresser erstellt.")
        
    def friss(self):
        print("Das Tier frisst Fleisch.")
       
class Allesfresser(Pflanzenfresser, Fleischfresser):
    def __init__(self):
        super().__init__()
        print("Allesfresser erstellt.")

baer = Allesfresser()

### END SOLUTION
```

## d)

Betrachten Sie den folgenden Python-Code:

```Python
class F: pass 
class G: pass 
class H: pass
class E(G,H): pass
class D(E,F): pass 
class C(E,G): pass
class B(C,H): pass
class A(D,B,E): pass
```

Ist es möglich, dass in der Linearisierung von A die Klasse E vor C vorkommt? Begründung!

<!-- grade_id: a1d-sol, points: 2, grade: true, solution: true -->
Nein, dies würde die Monotonie verletzen, da C eine Unterklasse von E ist.

## e)

Betrachten Sie noch einmal das Beispiel aus Teilaufgabe d):

Wäre es möglich, dass in der Linearisierung von A die Klasse F vor E vorkommt? Begründung!

<!-- grade_id: a1e-sol, points: 2, grade: true, solution: true -->
Dies wäre nicht möglich.
Schließlich ist D eine Oberklasse von A und die Reihenfolge der direkten Oberklassen von D lautet: zuerst E, dann F.
Sonst wäre die "Local precedence ordering"-Eigenschaft verletzt.

Betrachten Sie die folgende Klassenhierarchie:
    
```Python
class G: pass
class F: pass
class E: pass
class D(F): pass
class C(G,D,E): pass
class B(E,F): pass
class A(B,C): pass
```

Ist diese Klassenhierarchie linearisierbar? 
Wenn ja, geben Sie eine Linearisierung an.
Wenn nein, warum nicht?

<!-- grade_id: a1f-sol, points: 2, grade: true, solution: true -->
Die Klassenhierarchie ist nicht linearisierbar. Die Oberklasse B von A verlangt aufgrund der Reihenfolge der Oberklassen, dass E vor F kommt. Allerdings verlangt die Oberklasse C von A, dass D und damit dessen Oberklasse, F, vor E kommt.

Genauer:
* Lin(B) = EF
* Lin(D) = DF
* Lin(C) = C + merge (G, DF, E) = C + G + merge(DF, E) = C + G + D + merge(F, E) = C + G + D + F + E
* Lin(A) = A + merge(EF, CGDFE) = ... = A + C + G + D + merge(EF, FE) --> kein good head mehr vorhanden

# Aufgabe 2:

## a)

Im Ordner, in dem auch dieses Übungsblatt liegt, finden Sie eine Datei namens `simple_example.csv`.
Es handelt sich dabei um eine sogenannte Comma-separated values (CSV) - Datei.
CSV-Dateien lassen sich beispielsweise nutzen, um einfache Tabellen in einer menschenlesbaren Form abzuspeichern.
Mehr Informationen hierzu können Sie beispielsweise [hier auf Wikipedia](https://en.wikipedia.org/wiki/Comma-separated_values) nachlesen.
Die Datei `simple_example.csv` stellt eine Tabelle mit zwei Spalten dar: Zeit (in min seit Start eines Experiments) und Temperatur.
Die erste Zeile enthält jeweils den Titel jeder Spalte.
Das verwendete Trennzeichen ist das Komma.

Dateien dieser Art sollen nun in Python eingelesen werden.
Definieren Sie dazu eine Klasse namens `Table` mit folgenden Attributen und Methoden:
* Attribut titles: eine Liste, welche alle Spaltenüberschriften enthält
* Attribut cols: eine Liste von Listen, welche die (Fließkomma-)Einträge der Spalten enthält (col[0] steht für die erste Spalte, usw.)
* Methode `__str__(self)`, um die Attribute sinnvoll auf der Konsole ausgeben zu können
* Methode `__eq__(self, other)`, die nur dann `True` zurückgibt, wenn die Attributwerte beider als Parameter übergebener Objekte gleich sind, sonst `False`
* Methode `parse_csv(self, filename)`, welche den Dateinamen einer CSV-Datei (wie z. B. `simple_example.csv`) als Parameter übergeben bekommt, diese einliest und die eingelesenen Daten (Spaltenüberschriften und Zelleneinträge) im Objekt selbst speichert
* einen Konstruktor, dem man entweder die beiden Attribute direkt übergeben kann oder den Dateinamen einer CSV-Datei (über den Parameter `filename`), welche eingelesen werden soll.

Wird beispielsweise der Konstruktor durch `Table(filename = "simple_example.csv")` aufgerufen, wobei `simple_example.csv` so aussieht:
```
Zeit,Temperatur
0.0,15.2
3.0,16.1
6.0,14.8
9.0,20.2
```

dann soll ein Objekt vom Typ `Table` mit folgenden Attributwerten erstellt werden:
* `titles: ['Zeit', 'Temperatur']`
* `cols: [[0.0, 3.0, 6.0, 9.0], [15.2, 16.1, 14.8, 20.2]]`

Verwenden Sie für die Implementierung der Funktion `parse_csv` *unbedingt* [das CSV-Modul von Python](https://docs.python.org/3/library/csv.html)!

Sie dürfen davon ausgehen, dass die Datei mit dem als Parameter übergebenen Dateinamen tatsächlich existiert.

<!-- grade_id: a2a-sol, points: 0, grade: false, solution: true -->
```python
import csv

### BEGIN SOLUTION
class Table:
    """Klasse zur Repräsentation von Tabellen beliebigen Spalten und zugehörigen Spaltentiteln.
    
    Attributes:
        titles (list): Liste der Spaltentitel.
        cols (list): Liste von Listen mit den Einträgen der Spalten.
    """
    def __init__(self, titles = [], cols = [], filename = ''):
        """Konstruktor, der alle Attribute initialisiert (mit den als Parameter übergebenen Werten)."""
        self.titles = titles
        self.cols = cols
        if filename != '':
            self.parse_csv(filename)
        
    def __str__(self):
        return "titles: {}\ncols: {}".format( 
                self.titles, self.cols)

    def __eq__(self, other):
        return self.titles == other.titles and self.cols == other.cols

    def parse_csv(self, filename):
        """Liest eine CSV-Datei mit Spaltentiteln ein und füllt die Werte des aktuellen Objekts 
        mit den eingelesenen Werten.

        Args:
            filename: Dateiname der einzulesenden CSV-Datei (str)
        """
        with open(filename, newline='') as csvfile:
            reader = csv.reader(csvfile, delimiter=',', quotechar='|')
            self.titles = next(reader)    
            self.cols = [ [] for i in self.titles ]
            for row in reader:                
                for i, value in enumerate(row):
                    self.cols[i].append(float(value))
    ### END SOLUTION
```

<!-- grade_id: a2a0-test, points: 3, grade: true, solution: false -->
```python
### HIDESTART
def table_to_float(table):
    return Table(table.titles, [ [float(row) for row in col] for col in table.cols ] )
    
def write_table(filename, table):
    with open(filename, 'w', newline='') as csvfile:
        writer = csv.writer(csvfile, delimiter=',')
        writer.writerow(table.titles)
        for x in zip(*table.cols):
            writer.writerow(x)  
### HIDEEND            

# erster, einfacher Test:
my_table1 = Table(["title1", "title2"], [["l1"], ["l2"]])

print(my_table1)

assert my_table1.titles[0] == "title1"
assert my_table1.titles[1] == "title2"
assert my_table1.cols[0] == ["l1"]
assert my_table1.cols[1] == ["l2"]

my_table2 = Table(["title1", "title2"], [["l1"], ["l2"]])
assert my_table2 == my_table1

# zweiter Test, mit lesen:
titles = ["Zeit", "Temperatur"]
filename1 = "simple_example.csv"
my_table3 = Table(titles, [[0.0, 3.0, 6.0, 9.0], [15.2, 16.1, 14.8, 20.2]])
my_table3_parsed = Table(filename = "simple_example.csv")
### HIDESTART
print(my_table3_parsed)
assert table_to_float(my_table3_parsed) == my_table3 # Teilpunkte, falls immerhin dies funktioniert
### HIDEEND
assert my_table3_parsed == my_table3 # Teilpunkte, falls immerhin dies funktioniert

# dritter Test       
my_table4 = Table(["Zeit", "Temperatur"], [[0.0, 3.0, 6.0, 9.0, 12.0, 15.0, 18.0, 21.0, 24.0, 27.0, 30.0, 33.0, 36.0, 39.0, 42.0, 45.0, 48.0, 51.0, 54.0, 57.0, 60.0, 63.0, 66.0, 69.0, 72.0, 75.0, 78.0, 81.0, 84.0, 87.0, 90.0, 93.0, 96.0, 99.0, 102.0, 105.0, 108.0, 111.0], [35.90574523051787, 27.732801078808876, 42.035892511616005, 58.000772870249925, 96.94160853867335, 68.18506402498244, 14.671104120321022, 89.46105104216167, 47.27588706818212, 39.20070242229503, 26.714222293869348, 94.23349196918205, 51.77845568366778, 51.20161232801617, 93.30105788349742, 1.716984636989427, 8.755447352930679, 35.609379762311974, 24.876780742139005, 83.22149466739975, 65.20304271184895, 12.436205877934714, 90.37672735944263, 42.01655399602063, 83.90332779517878, 22.413857456299844, 51.83371482019683, 23.375603100801754, 52.15712567238341, 0.23171943708498732, 35.620773790224824, 1.8610917723884013, 50.37944193914525, 14.31962104350778, 60.72421338243065, 15.53821017629663, 91.90231143517047, 16.374637453126528]])
### HIDESTART
write_table("results1.csv", my_table4)
### HIDEEND
my_table4_parsed = Table(filename = "results1.csv")
### HIDESTART
print(my_table4_parsed)
assert table_to_float(my_table4_parsed) == my_table4 # Teilpunkte, falls immerhin dies funktioniert
### HIDEEND
assert my_table4_parsed == my_table4 # Teilpunkte, falls immerhin dies funktioniert
```

## b)

Schreiben Sie eine Funktion namens ```simple_plot```, die zwei gleich lange Listen als Parameter erhält und, unter Verwendung der Bibliothek `matplotlib`, ein einfaches Diagramm (engl. plot) erzeugt, welches die Elemente der zweiten Liste auf der y-Achse über den Elementen der ersten Liste (x-Achse) aufträgt.
Beispielsweise könnte ein Plot der in Teilaufgabe a) von der Funktion `parse_csv` zurückgegebenen Liste folgendermaßen aussehen:
![Beispiel-Plot](./plot_example.png)

Die weiteren Parameter der Funktion sind dabei durch den bereits gegebenen docstring vorgegeben und sollen berücksichtigt werden.

Verwenden Sie zur Implementierung die Methode [pyplot.plot aus der Bibliothek matplotlib](http://matplotlib.org/api/pyplot_api.html#matplotlib.pyplot.plot).

<!-- grade_id: a2b-sol, points: 1, grade: false, solution: true -->
```python
import matplotlib.pyplot

def simple_plot(xvalues, yvalues, title, xlabel, ylabel):
    """Erstellt ein Diagramm auf Basis von Werten für die x- und die y-Achse und gibt dieses aus. 
       Dabei wird der erste Eintrag der zweiten Liste dem ersten Eintrag der ersten Liste zugeordnet, usw.
  
    Args:
        xvalues: Liste mit Werten, die auf der x-Achse auftragen werden. (list)
        yvalues: Liste mit Werten, die auf der y-Achse aufgetragen werden. (list)
        title:  Der Titel, der über dem Plot stehen soll (str)
        xlabel: Die Achsenbeschriftung für die x-Achse (str)
        ylabel: Die Achsenbeschriftung für die y-Achse (str)
    """
    ### BEGIN SOLUTION    
    matplotlib.pyplot.plot(xvalues, yvalues)

    matplotlib.pyplot.title(title)
    matplotlib.pyplot.xlabel(xlabel)
    matplotlib.pyplot.ylabel(ylabel)
    matplotlib.pyplot.grid(True)
    matplotlib.pyplot.show()
    ### END SOLUTION
```

<!-- grade_id: a2b-test, points: 2, grade: true, solution: false -->
```python
my_values = ([0.0, 3.0, 6.0, 9.0], [15.2, 16.1, 14.8, 20.2])

simple_plot(my_values[0], my_values[1], "Simulationsergebnisse", "Zeit (s)", "Temperatur (°C)")

### HIDESTART
#assert false # hier muss manuell geprueft werden
### HIDEEND
```

## c)
Die in der übernächsten Codezelle vorgegebene Funktion ```parse_and_plot()``` dient dazu, den Inhalt einer zweispaltigen CSV-Datei einzulesen und als Diagramm auszugeben.
Die einzulesende Datei wird beim Aufruf des Skriptes, in der diese Funktion steht (z. B: `script.py`, als Argument übergeben (z. B. `python3 script.py simple_example.csv`).
Anschließend wird diese CSV-Datei in ein `Table`-Objekt (Teilaufgabe a)) eingelesen und aus dieser anschließend mittels der Funktion `simple_plot` (Teilaufgabe b)) ein Diagramm erzeugt.

Für das Verarbeiten der Kommandozeilenargumente soll die Funktion ```parse()``` dienen, welche von Ihnen zu schreiben ist. Die Kommandozeilenparameter sind, wie üblich, in `sys.argv` gespeichert.

Die Kommandozeilenparameter, die übergeben werden können, sind folgende:
* Der Dateiname der einzulesenden CSV-Datei als (nicht optionales) Positionsargument (an erster Stelle) (string)
* Der Titel des Diagramms als optionales Argument mit dem Namen "title" (string)
* Die Achsenbeschriftung auf der x-Achse als optionales Argument mit dem Namen "xlabel" (string)
* Die Achsenbeschriftung auf der y-Achse als optionales Argument mit dem Namen "ylabel" (string)

Beispielsweise soll bei Aufruf des Skripts mit den Argumenten "simple_example.csv --title=Simulationsergebnisse --xlabel=Zeit (s) --ylabel=Temperatur (°C) das Diagramm aus dem Beispiel von Aufgabe b) erzeugt werden.

Die Funktion `parse()` soll ein Objekt mit den Attributen `filename`, `title`, `xlabel` und `ylabel` zurückgeben, deren Werte den obigen Kommandozeilenparametern entsprechen. Ist einer der optionalen Parameter nicht angegeben, so soll jeweils der leere String in dem jeweiligen Attributwert gespeichert werden.

Verarbeiten Sie *unbedingt* die Liste der Kommandozeilenparameter mit Hilfe des Moduls [argparse](https://docs.python.org/3/library/argparse.html) und verwenden Sie die Funktionen `parse_csv` und `simple_plot` aus den Teilaufgaben a) und b).

<!-- grade_id: a2c-sol, points: 1, grade: false, solution: true -->
```python
import sys
import argparse

def parse():
    ### BEGIN SOLUTION
    """Verarbeitet die Kommandozeilenparameter filename, title (optional), xlabel (optional) und ylabel (optional)
    
    Returns:
        Ein Objekt mit den Attributen filename, title, xlabel und ylabel, welche die beim Skriptaufruf
        übergebenen Argumente speichern. 
        title, xlabel und ylabel sind der leere String, falls diese optionalen Parameter nicht übergeben wurden.
    """
    parser = argparse.ArgumentParser(description='Process some integers.')
    parser.add_argument('filename', metavar='filename',
                    help='Der Dateiname, der anzugeben ist.')
    parser.add_argument('--title', dest='title', default="", type=str, 
                    help='Der Titel, der über dem Diagramm erscheinen soll.')
    parser.add_argument('--xlabel', dest='xlabel', default="",
                    help='x-Achsenbeschriftung (überschreibt Überschrift der ersten Spalte)')
    parser.add_argument('--ylabel', dest='ylabel', default="",
                    help='y-Achsenbeschriftung (überschreibt Überschrift der zweiten Spalte)')

    args = parser.parse_args()

    return args
    ### END SOLUTION
```

```python
# in dieser Codezelle sollen Sie nichts ändern

def parse_and_plot():
    """Liest eine CSV-Datei gemäß Kommandozeilenparameter ein und erstellt aus den Werten der 
    ersten beiden Spalten ein Diagramm.
    Wird kein Titel angegeben, wird standardmäßig der leere String als Titel verwendet werden.
    Wird für die x- und/oder die y-Achse keine Achsenbeschriftung angegeben, 
    so wird der Titel der ersten (bzw. zweiten) Spalte für die Beschriftung der x- (bzw. y-) Achse verwendet.    
    """
    args = parse()
    table = Table(filename = args.filename)

    # kleiner Trick: Der leere String evaluiert zu False ...  
    simple_plot(table.cols[0], table.cols[1],
                    args.title,
                    args.xlabel or table.titles[0],
                    args.ylabel or table.titles[1],                    
                    )
```

```python
# Diese Zelle ist nur für die Kontrolle Ihrerseits gedacht.

from unittest.mock import patch

# die folgenden Zeilen simulieren, dass aller Python-Code in diesem Notebook in einer Datei
# namens script.py liegt und dieses Skript mit dem Aufruf 
# > python3 script.py simpleexamle.csv --title=Simulationsergebnisse, --xlabel=Zeit (s) --ylabel=Temperatur (°C) 
# ausgeführt wird.
vargs = ["script.py", "simple_example.csv", '--title=Simulationsergebnisse', '--xlabel=Zeit (s)', '--ylabel=Temperatur (°C)']
with patch.object(sys, 'argv', vargs):
    parse_and_plot()
    
# die folgenden Zeilen simulieren, dass aller Python-Code in diesem Notebook in einer Datei
# namens script.py liegt und dieses Skript mit dem Aufruf 
# > python3 script.py results1.csv
# ausgeführt wird.    
vargs = ["script.py", "results1.csv"]
with patch.object(sys, 'argv', vargs):
    parse_and_plot()
```

<!-- grade_id: a2c-test, points: 2, grade: true, solution: false -->
```python
vargs = ["script.py", "simple_example.csv", '--title=Simulationsergebnisse', '--xlabel=Zeit (s)', '--ylabel=Temperatur (°C)']
with patch.object(sys, 'argv', vargs):
    args = parse()
    assert args.filename == "simple_example.csv"
    assert args.title == "Simulationsergebnisse"
    assert args.xlabel == "Zeit (s)"
    assert args.ylabel == "Temperatur (°C)"
    
vargs = ["script.py", "results1.csv"]
with patch.object(sys, 'argv', vargs):
    args = parse()
    assert args.filename == "results1.csv"
    assert args.title == ""
    assert args.xlabel == ""
    assert args.ylabel == ""
```

## d)
Angenommen, der gesamte obige Code wäre nicht in einem Jupyter-Notebook, sondern in einem Python-Skript geschrieben.
Was müssten Sie schreiben, damit bei Ausführung des Skriptes (allerdings nicht bei der Einbindung als Modul!) automatisch die Methode `parse_and_plot()` aufgerufen wird?

<!-- grade_id: a2d-sol, points: 1, grade: true, solution: true -->
```python
#if __name__ == "__main__":
#    parse_and_plot()
# TODO: auskommentieren, wenn das Review des Blattes fertig ist
```

<!-- grade_id: docstring_test, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Prüft, ob alle Dokumentationen vorhanden sind.

vargs = ["script.py", "results1.csv"]
with patch.object(sys, 'argv', vargs):
    l = []

    # Aufgabe 2a)
    l.append(("Table", False))
    l.append(("Table", True, "[], []"))
    l.append(("Table([], []).parse_csv", True, '"simple_example.csv"'))
    # Aufgabe 2b)
    l.append(("simple_plot", True, '[], [], "", "", ""'))
    # Aufgabe 2c)
    l.append(("parse", True, ""))

    for f in l:
        try:
            func = eval(f[0])
            if f[1]:
                eval(f[0] + "(" + f[2] + ")")
            # Schlägt der folgende Test fehl, wurde die Funktion zwar implementiert, hat aber keinen docstring.
            assert(func.__doc__ != None)
        except NotImplementedError:
            print("Folgende Funktion wurde nicht implementiert (daher wird nicht geprüft, ob sie einen docstring hat):", "-->", f[0])
        except AssertionError:
            print("Bei folgender Funktion fehlt der docstring: ", func.__doc__ != None, "-->", f[0])    
            assert(False)
```