<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

<!--- Wenn Sie diese Ansicht verwenden (was Sie nicht tun sollten; klicken Sie am besten einmal auf "Run cell"), dann steht die Formatierung hier richtig:

Format:

MatrikelNummer E-Mail-Adresse
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
234567 testStudent2@mail.upb.de
345678 vorname.nachname@mail.upb.de

... und nicht hier: -->

Format:

MatrikelNummer E-Mail-Adresse
<br \>
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
<br \>
234567 testStudent2@mail.upb.de
<br \>
345678 vorname.nachname@mail.upb.de

Hinweis: Wenn Sie das so in die obere Zelle eingeben, wird in der ausgeführten Markdownzelle (Strg+Enter) alles in einer Zeile angezeigt. Das ist egal; wichtig ist Formatierung der Eingabe.


## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile), für das entsprechende Übungsblatt! (Wenn Sie Blatt 10 abgeben wolllen hilft es nicht, bei Blatt 1 auf Submit zu drücken!) Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue11.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen sinnvollen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 29.01.2017 um 23:59 Uhr

# Aufgabe 1:

## a)
Betrachten Sie den folgenden Python-Code:

```Python 
class Tier:
    def friss(self):
        print("Das Tier frisst.")
    
class Pflanzenfresser(Tier):
    def friss(self):
        print("Das Tier frisst vegetarisch.")

class Fleischfresser(Tier):
    def friss(self):
        print("Das Tier frisst Fleisch.")
       
class Allesfresser(Pflanzenfresser, Fleischfresser):
    pass

baer = Allesfresser()
baer.friss()
```

Welche Ausgabe wird erzeugt und warum?

<!-- grade_id: a1-sol, points: 2, grade: true, solution: true -->
Es wird "Das Tier isst vegetarisch" ausgegeben, da die Methode `iss()` von `Pflanzenfresser` aufgerufen wird, da diese Oberklasse als erstes in der Liste der Oberklassen aufgeführt ist.

## b)
Betrachten Sie das folgende, komplexere, Beispiel:

```Python
class Tier:
    def friss(self):
        print("Das Tier frisst.")

class Pflanzenfresser(Tier):
    def friss(self):
        print("Das Tier frisst vegetarisch.")
        
class Fleischfresser(Tier):
    def extra_portion(self):
        print("Der Fleischfresser frisst noch eine Portion.")
        
class Allesfresser(Pflanzenfresser, Fleischfresser):
    def friss(self):
        print("Der Allesfresser frisst.")
        self.extra_portion()

class Raubtier(Fleischfresser):
    def extra_portion(self):
        print("Das Raubtier frisst noch eine Portion.")        

class Baer(Allesfresser, Raubtier): pass

baer = Baer()
baer.friss()
```

Welche Ausgabe erzeugt dieser Code? Erklären Sie!

<!-- grade_id: a1b-sol, points: 1, grade: true, solution: true -->
Es wird die Ausgabe 

```Der Allesfresser frisst.
Das Raubtier frisst noch eine Portion.
```

erzeugt.

Dies liegt an der Method Resolution Order, welche die Klasse Raubtier vor Fleischfresser enthält (aufgrund des C3-Algorithmus).
Würde die Klasse Fleischfresser vor Raubtier in der MRO vorkommen, wäre die Monotonie verletzt.

## c)
Betrachten Sie nun die folgende Abwandlung des Beispiels aus Aufgabe a):

```Python 
class Tier:
    def friss(self):
        print("Das Tier frisst.")
    
class Pflanzenfresser(Tier):
    def __init__(self):
        print("Pflanzenfresser erstellt.")

    def friss(self):
        print("Das Tier frisst vegetarisch.")

class Fleischfresser(Tier):
    def __init__(self):
        print("Fleischfresser erstellt.")
        
    def friss(self):
        print("Das Tier frisst Fleisch.")
       
class Allesfresser(Pflanzenfresser, Fleischfresser):
    def __init__(self):
        super().__init__()
        print("Allesfresser erstellt.")

baer = Allesfresser()
```

Kopieren Sie den Code in die folgende Zelle und führen Sie ihn aus.

Wie Sie sehen, werden lediglich die Ausgaben der Konstruktoren der Klassen `Allesfresser` und `Pflanzenfresser` aufgerufen.

Wandeln Sie den Code (in der folgenden Zelle) so ab, dass auch der Konstruktor der Klasse `Fleischfresser` aufgerufen wird.
Vermeiden Sie dabei unnötige Wiederholungen.

<!-- grade_id: a1c-sol, points: 2, grade: true, solution: true -->
```python
### BEGIN SOLUTION

class Tier:
    def friss(self):
        print("Das Tier frisst.")
    
class Pflanzenfresser(Tier):
    def __init__(self):
        super().__init__()
        print("Pflanzenfresser erstellt.")

    def friss(self):
        print("Das Tier frisst vegetarisch.")

class Fleischfresser(Tier):
    def __init__(self):
        super().__init__()    # Zeile für Aufgabenerfüllung nicht erforderlich, aber guter Stil
        print("Fleischfresser erstellt.")
        
    def friss(self):
        print("Das Tier frisst Fleisch.")
       
class Allesfresser(Pflanzenfresser, Fleischfresser):
    def __init__(self):
        super().__init__()
        print("Allesfresser erstellt.")

baer = Allesfresser()

### END SOLUTION
```

## d)

Betrachten Sie den folgenden Python-Code:

```Python
class F: pass 
class G: pass 
class H: pass
class E(G,H): pass
class D(E,F): pass 
class C(E,G): pass
class B(C,H): pass
class A(D,B,E): pass
```

Ist es möglich, dass in einer gültigen Linearisierung von A die Klasse E vor C vorkommt? Begründung!

<!-- grade_id: a1d-sol, points: 2, grade: true, solution: true -->
Nein, dies würde die Monotonie verletzen, da C eine Unterklasse von D ist.

## e)

Betrachten Sie noch einmal das Beispiel aus Teilaufgabe d):

Wäre es möglich, dass in einer gültigen Linearisierung von A die Klasse F vor E vorkommt? Begründung!

<!-- grade_id: a1e-sol, points: 2, grade: true, solution: true -->
Dies wäre nicht möglich.
Schließlich ist D eine Oberklasse von A und die Reihenfolge der direkten Oberklassen von D lautet: zuerst E, dann F.
Sonst wäre die "Local precedence ordering"-Eigenschaft verletzt.

Betrachten Sie die folgende Klassenhierarchie:
    
```Python
class G: pass
class F: pass
class E: pass
class D(F): pass
class C(G,D,E): pass
class B(E,F): pass
class A(B,C): pass
```

Ist diese Klassenhierarchie linearisierbar? 
Wenn ja, geben Sie eine Linearisierung an.
Wenn nein, warum nicht?

<!-- grade_id: a1f-sol, points: 2, grade: true, solution: true -->
Die Klassenhierarchie ist nicht linearisierbar.

(mir fehlt aber gerade noch das entscheidende Argument, warum A, B, C, G, D, E, F keine gültige Linearisierung ist laut VL)

HK: Weil bei Lin(C) FE enthalten ist, bei Lin(B) aber EF enthalten ist. Damit zirkulare Abhängigkeit, wenn Lin(A) berechnet werden soll. Liegt letztlich daran, dass B(E,F) E vor F hat, während bei C F vor E steht, über den Umweg von D. D tut ja eigentlichs nichts zur Sache hier. 

# Aufgabe 2:

## a)

Im Ordner, in dem auch dieses Übungsblatt liegt, finden Sie eine Datei namens `results.csv`.
Es handelt sich dabei um eine sogenannte Comma-separated values (CSV) - Datei.
CSV-Dateien lassen sich beispielsweise nutzen, um einfache Tabellen in einer menschenlesbaren Form abzuspeichern.
Mehr Informationen hierzu können Sie beispielsweise [hier auf Wikipedia](https://en.wikipedia.org/wiki/Comma-separated_values) nachlesen.
Die hier vorliegende Datei stellt eine Tabelle mit zwei Spalten dar: Zeit (in min seit Start eines Experiments) und Temperatur.
Die erste Zeile enthält jeweils den Titel jeder Spalte.
Das verwendete Trennzeichen ist das Komma.

Diese Datei soll nun in Python eingelesen werden.
Definieren Sie zunächst eine Klasse namens `Table`, welche die vier Attribute `title1`, `title2`, `values1` und `values2` enthält, welche alle durch Parameterübergabe an den Konstruktor gesetzt werden können.
Außerdem soll die Klasse eine Methode `__str__(self)` besitzen, um die Attribute sinnvoll auf der Konsole ausgeben zu können und eine Methode `__eq__(self, other)`, die nur dann `True` zurückgibt, wenn die Attributwerte beider als Parameter übergebener Objekte gleich sind, sonst `False`.

Schreiben Sie anschließend eine Funktion namens `parse_csv`, welche den Dateinamen einer CSV-Datei (also z.B. `"results.csv"`) als Parameter übergeben bekommt, welche:
* zwei Spalten enthält,
* in der ersten Zeile die Titel der Spalten enthält (als Zeichenkette),
* in den weiteren Zeilen nur Zahlen (mit Dezimalstellen, mit Punkt . als Trennzeich!) speichert

und die ein Objekt der Klasse `Table` zurückgibt, dessen Attribute folgendermaßen gesetzt sind:
* `title1`: der Titel der ersten Spalte,
* `title2`: der Titel der zweiten Spalte
* `values1`: eine Liste mit allen Werten der ersten Spalte (ohne den Titel), in der Reihenfolge, wie sie in der Datei auftauchen
* `values2`: eine Liste mit allen Werten der zweiten Spalte (ohne den Titel), in der Reihenfolge, wie sie in der Datei auftauchen.

Wird beispielsweise der Dateiname der CSV-Datei `simple_example.csv` übergeben, welche so aussieht: 
```
Zeit,Temperatur
0.0,15.2
3.0,16.1
6.0,14.8
9.0,20.2
```

Dann soll die Rückgabe der Funktion ein Objekt vom Typ `Table` mit folgenden Attributwerten sein:
* `title1 = "Zeit"`
* `title2 = "Temperatur"`
* `values1 = [0.0, 3.0, 6.0, 9.0]`
* `values2 = [15.2, 16.1, 14.8, 20.2]`

Verwenden Sie für die Implementierung der Funktion `parse_csv` *unbedingt* [das CSV-Modul von Python](https://docs.python.org/3/library/csv.html)!

Sie dürfen davon ausgehen, dass die Datei mit dem als Parameter übergebenen Dateinamen tatsächlich existiert.

<!-- grade_id: a2a-sol, points: 0, grade: false, solution: true -->
```python
import csv

### BEGIN SOLUTION
class Table:
    """Klasse zur Repräsentation von Tabellen mit zwei Spalten und zugehörigen Spaltentiteln.
    
    Attributes:
        title1 (str): Titel der ersten Spalte.
        title2 (str): Titel der zweiten Spalte.
        values1 (list): Einträge der ersten Spalte.
        values2 (list): Einträge der zweiten Spalte.
    """
    def __init__(self, title1, title2, values1, values2):
        """Konstruktor, der alle Attribute initialisiert (mit den als Parameter übergebenen Werten)."""
        self.title1 = title1
        self.title2 = title2
        self.values1 = values1
        self.values2 = values2
        
    def __str__(self):
        return "title1: {}\ntitle2: {}\nvalues1: {}\nvalues2: {}".format( 
                self.title1, self.title2, self.values1, self.values2)

    def __eq__(self, other):
        return self.title1 == other.title1 and self.title2 == other.title2 and \
                self.values1 == other.values1 and self.values2 == other.values2
### END SOLUTION

def parse_csv(filename):
    ### BEGIN SOLUTION
    #TODO: docstring
    with open(filename, newline='') as csvfile:
    #    csvfile = open(filename, newline='')
        reader = csv.reader(csvfile, delimiter=',', quotechar='|')
        first_row = next(reader)    
        values1, values2 = [], []
        for row in reader:
            values1.append(float(row[0]))
            values2.append(float(row[1]))
        #csvfile.close()
    
    return Table(first_row[0], first_row[1], values1, values2)
    ### END SOLUTION
```

<!-- grade_id: a2a0-test, points: 1, grade: true, solution: false -->
```python
#TODO: hier die Klasse Table testen
```

<!-- grade_id: a2a-test, points: 2, grade: true, solution: false -->
```python
### HIDESTART
#import random
#TODO: hier die Klasse Table aus ML nochmal definieren...
### HIDEEND
title_column1 = "Zeit"
title_column2 = "Temperatur"

filename1 = "simple_example.csv"
atable = Table(title_column1, title_column2, [0.0, 3.0, 6.0, 9.0], [15.2, 16.1, 14.8, 20.2])
parsed = parse_csv(filename1)
### HIDESTART
print(parsed)

assert parsed.title1 == title_column1
assert parsed.title2 == title_column2
parsed_semi = Table(parsed.title1, parsed.title2, [float(i) for i in parsed.values1], [float(i) for i in parsed.values2] ) 
assert parsed_semi == atable #Teilpunkte, falls immerhin dies funktioniert
### HIDEEND
assert parsed == atable

filename2 = "results1.csv"
atable = Table(title_column1, title_column2, [0.0, 3.0, 6.0, 9.0, 12.0, 15.0, 18.0, 21.0, 24.0, 27.0, 30.0, 33.0, 36.0, 39.0, 42.0, 45.0, 48.0, 51.0, 54.0, 57.0, 60.0, 63.0, 66.0, 69.0, 72.0, 75.0, 78.0, 81.0, 84.0, 87.0, 90.0, 93.0, 96.0, 99.0, 102.0, 105.0, 108.0, 111.0], [35.90574523051787, 27.732801078808876, 42.035892511616005, 58.000772870249925, 96.94160853867335, 68.18506402498244, 14.671104120321022, 89.46105104216167, 47.27588706818212, 39.20070242229503, 26.714222293869348, 94.23349196918205, 51.77845568366778, 51.20161232801617, 93.30105788349742, 1.716984636989427, 8.755447352930679, 35.609379762311974, 24.876780742139005, 83.22149466739975, 65.20304271184895, 12.436205877934714, 90.37672735944263, 42.01655399602063, 83.90332779517878, 22.413857456299844, 51.83371482019683, 23.375603100801754, 52.15712567238341, 0.23171943708498732, 35.620773790224824, 1.8610917723884013, 50.37944193914525, 14.31962104350778, 60.72421338243065, 15.53821017629663, 91.90231143517047, 16.374637453126528])

### HIDESTART
title_column1 = "Zeit"
title_column2 = "Temperatur"

#filename2 = "results_generated.csv"

# Neu schreiben, um Faelschungen zu verhindern
with open(filename2, 'w', newline='') as csvfile:
    writer = csv.writer(csvfile, delimiter=',')
    writer.writerow([title_column1, title_column2])
    #values = (title_column1, title_column2, [], [])
    for i in range(len(atable.values1)): # 
    #for i in range(round(random.random() * 100)):
    #    values[2].append(i * 3)
    #    values[3].append(random.random()*100)
        writer.writerow([atable.values1[i], atable.values2[i]])    
### HIDEEND

parsed = parse_csv(filename2)
### HIDESTART
print(parsed)

assert parsed.title1 == title_column1
assert parsed.title2 == title_column2
parsed_semi = Table(parsed.title1, parsed.title2, [float(i) for i in parsed.values1], [float(i) for i in parsed.values2] ) 
assert parsed_semi == atable #Teilpunkte, falls immerhin dies funktioniert
### HIDEEND

assert parsed == atable
```

## b)

Schreiben Sie eine Funktion namens ```simple_plot```, die zwei gleich lange Listen als Parameter erhält und, unter Verwendung der Bibliothek `matplotlib`, ein einfaches Diagramm (engl. plot) erzeugt, welches die Elemente der zweiten Liste auf der y-Achse über den Elementen der ersten Liste (x-Achse) aufträgt.
Beispielsweise könnte ein Plot der in Teilaufgabe a) von der Funktion `parse_csv` zurückgegebenen Liste folgendermaßen aussehen:
![Beispiel-Plot](./plot_example.png)

Die weiteren Parameter der Funktion sind dabei durch den bereits gegebenen docstring vorgegeben und sollen berücksichtigt werden.

Verwenden Sie zur Implementierung die Methode [pyplot.plot aus der Bibliothek matplotlib](http://matplotlib.org/api/pyplot_api.html#matplotlib.pyplot.plot).

<!-- grade_id: a2b-sol, points: 1, grade: false, solution: true -->
```python
import matplotlib.pyplot

def simple_plot(xvalues, yvalues, title, xlabel, ylabel):
    '''Erstellt ein Diagramm auf Basis von Werten für die x- und die y-Achse und gibt dieses aus. 
       Dabei wird der erste Eintrag der zweiten Liste dem ersten Eintrag der ersten Liste zugeordnet, usw.
  
    Args:
        xvalues: Liste mit Werten, die auf der x-Achse auftragen werden.
        yvalues: Liste mit Werten, die auf der y-Achse aufgetragen werden.
        title:  Der Titel, der über dem Plot stehen soll (str)
        xlabel: Die Achsenbeschriftung für die x-Achse (str)
        ylabel: Die Achsenbeschriftung für die y-Achse (str)
    '''
    ### BEGIN SOLUTION    
    matplotlib.pyplot.plot(xvalues, yvalues)

    matplotlib.pyplot.title(title)
    matplotlib.pyplot.xlabel(xlabel)
    matplotlib.pyplot.ylabel(ylabel)
    matplotlib.pyplot.grid(True)
    #matplotlib.pyplot.savefig("plot_example.png")
    matplotlib.pyplot.show()
    ### END SOLUTION
```

<!-- grade_id: a2b-test, points: 2, grade: true, solution: false -->
```python
my_values = ([0.0, 3.0, 6.0, 9.0], [15.2, 16.1, 14.8, 20.2])

simple_plot(my_values[0], my_values[1], "Simulationsergebnisse", "Zeit (s)", "Temperatur (°C)")

### HIDESTART
#assert false # hier muss manuell geprueft werden
### HIDEEND
```

## c)
Schreiben Sie nun eine Funktion ```parse_and_plot()```, welche die Funktionen `parse_csv` und `simple_plot` aus den Teilaufgaben a) und b) aufruft, um Werte aus einer CSV-Datei einzulesen, welche anschließend in einem Diagramm dargestellt werden. 
Die Funktion nimmt zwar keine Parameter entgegen, verwendet aber die beim Aufruf des Skriptes auf der Konsole übergebenen Kommandozeilenparameter (`sys.argv`).
Die Kommandozeilenparameter, die übergeben werden können, sind folgende:
* Der Dateiname der einzulesenden CSV-Datei als (nicht optionales) Positionsargument (an erster Stelle) (string)
* Der Titel des Diagramms als optionales Argument mit dem Namen "title" (string)
* Die Achsenbeschriftung auf der x-Achse als optionales Argument mit dem Namen "xlabel" (string)
* Die Achsenbeschriftung auf der y-Achse als optionales Argument mit dem Namen "ylabel" (string)

Wird kein Titel angegeben, soll standardmäßig der leere String als Titel verwendet werden.
Wird für die x- und/oder die y-Achse keine Achsenbeschriftung angegeben, so soll der Titel der ersten (bzw. zweiten) Spalte für die Beschriftung der x- (bzw. y-) Achse verwendet werden.

Beispielsweise soll bei Aufruf des Skripts mit den Argumenten "simple_example.csv --title=Simulationsergebnisse --xlabel=Zeit (s) --ylabel=Temperatur (°C) das Diagramm aus dem Beispiel von Aufgabe b) erzeugt werden.

Verarbeiten Sie *unbedingt* die Liste der Kommandozeilenparameter mit Hilfe des Moduls [argparse](https://docs.python.org/3/library/argparse.html) und verwenden Sie die Funktionen `parse_csv` und `simple_plot` aus den Teilaufgaben a) und b).

<!-- grade_id: a2c-sol, points: 1, grade: false, solution: true -->
```python
import sys
import argparse


def parse_and_plot():
    ### BEGIN SOLUTION
    parser = argparse.ArgumentParser(description='Process some integers.')
    parser.add_argument('filename', metavar='filename',
                    help='bla')
    parser.add_argument('--title', dest='title', default="", type=str, 
                    help='blubb')
    parser.add_argument('--xlabel', dest='xlabel', default="",
                    help='blubb')
    parser.add_argument('--ylabel', dest='ylabel', default="",
                    help='blubb')

    args = parser.parse_args()
    table = parse_csv(args.filename)
    
    if args.xlabel == "":
        args.xlabel = table.title1
    if args.ylabel == "":
        args.ylabel = table.title2

    simple_plot(table.values1, table.values2, args.title, args.xlabel, args.ylabel)
    ### END SOLUTION


```

<!-- grade_id: a2c-test, points: 2, grade: true, solution: false -->
```python
# hier im Test noch die ML der vorherigen Aufgaben einfügen, damit die verwendet werden...

from unittest.mock import patch

vargs = ["script.py", "simple_example.csv", '--title=Simulationsergebnisse', '--xlabel=Zeit (s)', '--ylabel=Temperatur (°C)']
with patch.object(sys, 'argv', vargs):
    parse_and_plot()
    
vargs = ["script.py", "results1.csv"]
with patch.object(sys, 'argv', vargs):
    parse_and_plot()
```

## d)
Angenommen, der gesamte obige Code wäre nicht in einem Jupyter-Notebook, sondern in einem Python-Skript geschrieben.
Was müssten Sie schreiben, damit bei Ausführung des Skriptes (allerdings nicht bei der Einbindung als Modul!) automatisch die Methode `parse_and_plot()` aufgerufen wird?

<!-- grade_id: a2d-sol, points: 1, grade: true, solution: true -->
```python
#if __name__ == "__main__":
#    parse_and_plot()
```

<!-- grade_id: docstring_test, points: 1, grade: true, solution: false -->
```python
#TODO: die docstring-Zelle
```