<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

## Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 11.12.2016 um 23:59 Uhr
### Wichtig: Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue06.ipynb ab (nicht umbenennen, keine Kopie erstellen!). Sonst kann eine Bewertung nicht stattfinden.

# Aufgabe 1: String-API (6 Punkte)
Diese Aufgabe lässt sich eventuell einfacher mithilfe der Python-Dokumentation lösen:
https://docs.python.org/3/library/stdtypes.html.

a) Schreiben Sie eine Funktion `s_verbinder(l,z)`, welches eine Liste mit Strings `l` zu einem einzigen String verbindet, wobei zwischen die einzelnen Listenelemente jeweils das Zeichen `z` eingefügt werden soll.

<!-- grade_id: string_sol_1, points: 0, grade: false, solution: true -->
```python
def s_verbinder(l,z):
    # Your Code here
    ### BEGIN SOLUTION
    return z.join(l)
    ### END SOLUTION
```

<!-- grade_id: string_test_1, points: 1, grade: true, solution: false -->
```python
assert(s_verbinder(['1','2','3','4'], '-') == '1-2-3-4')
### HIDESTART
assert(s_verbinder(['a','b','c','d'], '-') == 'a-b-c-d')
### HIDEEND
```

b) Schreiben Sie eine Funktion `s_suche(s1,s2)`, welche den niedrigsten Index ausgibt, an welcher der String `s2` in `s1` auftaucht.

<!-- grade_id: string_sol_2, points: 0, grade: false, solution: true -->
```python
def s_suche(s1,s2):
    # Your Code here
    ### BEGIN SOLUTION
    return s1.find(s2)
    ### END SOLUTION
```

<!-- grade_id: string_test_2, points: 1, grade: true, solution: false -->
```python
assert(s_suche('Test', 'T') == 0)
assert(s_suche('Test', 'e') == 1)
assert(s_suche('Test', 's') == 2)
assert(s_suche('Test', 't') == 3)
assert(s_suche('Test', 'es') == 1)
### HIDESTART
assert(s_suche('Grundlagen der Programmierung', 'der') == 11)
assert(s_suche('Grundlagen der Programmierung', 'P') == 15)
### HIDEEND
```

c) Schreiben Sie eine Funktion `s_klein(s)`, welche alle Großbuchstaben im string `s` durch entsprechende Kleinbuchstaben ersetzt.

<!-- grade_id: string_sol_3, points: 0, grade: false, solution: true -->
```python
def s_klein(s):
    # Your Code here
    ### BEGIN SOLUTION
    return s.lower()
    ### END SOLUTION
```

<!-- grade_id: string_test_3, points: 1, grade: true, solution: false -->
```python
assert(s_klein('Hallo Welt!') == 'hallo welt!')
### HIDESTART
assert(s_klein('Dies ist ein geHEimer Satz.') == 'dies ist ein geheimer satz.')
### HIDEEND
```

d) Schreiben Sie eine Funktion `s_special_substitution(s)`, welche für einen string `s` sowohl den ersten Buchstaben, als auch alle Buchstaben, welche auf ein Leerzeichen folgen, groß schreibt.

<!-- grade_id: string_sol_4, points: 0, grade: false, solution: true -->
```python
def s_special_substitution(s):
    # Your Code here
    ### BEGIN SOLUTION
    return s.title()
    ### END SOLUTION
```

<!-- grade_id: string_test_4, points: 0, grade: true, solution: false -->
```python
assert(s_special_substitution('Der Typ string für Zeichenketten ist als eine Klasse definiert.') == 
       'Der Typ String Für Zeichenketten Ist Als Eine Klasse Definiert.')
### HIDESTART
assert(s_special_substitution('Dies ist ein geheimer Satz.') == 
       'Dies Ist Ein Geheimer Satz.')
### HIDEEND
```

e) Schreiben Sie eine Funktion `get_words(s)`, welche für einen string `s` eine Liste mit den einzelnen Wörtern von `s` zurückgibt.

<!-- grade_id: string_sol_5, points: 0, grade: false, solution: true -->
```python
def get_words(s):
    # Your Code here
    ### BEGIN SOLUTION
    return s.split(" ")
    ### END SOLUTION
```

<!-- grade_id: string_test_5, points: 0, grade: true, solution: false -->
```python
assert(get_words('Grundlagen der Programmierung') == ['Grundlagen', 'der', 'Programmierung'])
assert(get_words('Der Typ string für Zeichenketten ist als eine Klasse definiert.') ==
       ['Der', 'Typ', 'string', 'für', 'Zeichenketten', 'ist', 'als', 'eine', 'Klasse', 'definiert.'])
### HIDESTART
assert(get_words('Dies ist ein geheimer Satz.') == ['Dies', 'ist', 'ein', 'geheimer', 'Satz.'])
### HIDEEND
```

f) Schreiben Sie eine Funktion `split_3(s)`, welche einen string `s` an dem ersten Index `i` aufteilt, an welcher ein Punkt ('.') vorkommt. Es soll ein Triple `(s_links, i, s_rechts)` zurückgegeben werden, wobei `s_links` der Teilstring von `s` ist, welcher alle Zeichen links von `s[i]` enthält und `s_rechts`der Teilstring von `s` ist, welcher alle Zeichen rechts von `s[i]` enthält.

<!-- grade_id: string_sol_6, points: 0, grade: false, solution: true -->
```python
def split_3(s):
    # Your Code here
    ### BEGIN SOLUTION
    l = s.partition('.')
    sol = (l[0], s.index('.'), l[2])
    return sol
    ### END SOLUTION
```

<!-- grade_id: string_test_6, points: 0, grade: true, solution: false -->
```python
assert(split_3('Dies ist ein Satz. Dies ist noch ein Satz.') == ('Dies ist ein Satz', 17, ' Dies ist noch ein Satz.'))
### HIDESTART
assert(split_3('Dies ist ein geheimer Satz. Dies ist noch ein geheimer Satz.') == ('Dies ist ein geheimer Satz', 26, ' Dies ist noch ein geheimer Satz.'))
### HIDEEND
```

# Aufgabe 2: Rekursion (9 Punkte)
a) Gegeben sind die folgenden Funktionen `f` und `g` in Python:
```Python
def f(a):
    if a == 0:
        print('f')
    else:
        g(a)
        
def g(a):
    if a < 0:
        f(a+1)
    elif a > 0:
        print('g')
        f(a-1)
```

Geben Sie für den Aufruf `f(3)` die Folge der Funktionsaufrufe von `f` und `g` an sowie die Ausgabe des Programms.   
(1 Punkt)

<!-- grade_id: rekursion_sol_1, points: 1, grade: true, solution: true -->
Funktionsaufrufe: $f(3) \rightarrow g(3) \rightarrow f(2) \rightarrow g(2) \rightarrow f(1) \rightarrow g(1) \rightarrow f(0)$   
Ausgabe:   
`g`   
`g`   
`g`   
`f`

b) Schreiben Sie eine *rekursive* Funktion `anz_ziffern(x)`, welche für eine übergebene natürliche Zahl `x` die Anzahl der Ziffern von `x` zurückgibt. Sehen Sie sich die gegebenen Tests an für Beispiele.   
(1 Punkt)

<!-- grade_id: rekursion_sol_2, points: 0, grade: false, solution: true -->
```python
def anz_ziffern(x):
    # Your Code here
    ### BEGIN SOLUTION
    if x // 10 == 0:
        return 1
    return anz_ziffern(x//10) + 1
    ### END SOLUTION
```

<!-- grade_id: rekursion_test_2, points: 1, grade: true, solution: false -->
```python
assert(anz_ziffern(2) == 1)
assert(anz_ziffern(42) == 2)
assert(anz_ziffern(873246187456216) == 15)
### HIDESTART
l = [10**x for x in range(10)]
for i in l:
    assert(anz_ziffern(i) == l.index(i)+1)
### HIDEEND
```

c) Gegeben sei die folgende Funktion `f`:
\begin{equation}
   f(a,b) =
   \begin{cases}
     a & \text{falls } b = 1 \\
     a \cdot f(a, (b-1)/2) \cdot f(a, (b-1)/2)& \text{falls } b > 1 \text{ und } b \text{ ungerade} \\
     f(a, b/2) \cdot f(a, b/2)& \text{sonst} \\
   \end{cases}
\end{equation}

Implementieren Sie diese Funktion in Python, indem Sie die vorgegebene Funktion `f(a,b)` vervollständigen. Sie können davon ausgehen, dass `b` immer eine natürliche Zahl ist.
(2 Punkte)

<!-- grade_id: rekursion_sol_3, points: 0, grade: false, solution: true -->
```python
def f(a,b):
    # Your Code here
    ### BEGIN SOLUTION
    if b == 1:
        return a
    if b % 2 == 1:
        return a*f(a,(b-1)//2)*f(a,(b-1)//2) 
    return f(a, b//2)*f(a, b//2)
    ### END SOLUTION
```

<!-- grade_id: rekursion_test_3_1, points: 1, grade: true, solution: false -->
```python
assert(f(13,12) == 23298085122481)
assert(f(17,7) == 410338673)
assert(f(9,21) == 109418989131512359209)
### HIDESTART
# Teste den ersten Fall:
for a in range(1,10):
    assert(f(a,1) == a)
    
# Teste den zweiten Fall:
for a in range(1,10):
    for b in [3,5,7,9,11,13,15,17,19]:
        assert(f(a,b) == a**b)
### HIDEEND
```

<!-- grade_id: rekursion_test_3_2, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste den dritten Fall:
for a in range(1,10):
    for b in [2,4,6,8,10,12,14,16,18,20]:
        assert(f(a,b) == a**b)
### HIDEEND
```

d) Wie kann die durch `f` berechnete Funktion einfach ausgedrückt werden?
Welche Vorteile hat dies im Vergleich zu einer (naiven) iterativen Berechnung?
(1 Punkt)

<!-- grade_id: rekursion_sol_4, points: 2, grade: true, solution: true -->
### BEGIN SOLUTION
Die Funktion berechnet $a^b$ für ein $a \geq 0$ und eine natürliche Zahl $b$.
Die Berechnung wird schneller, da man nun anstelle von $b$ Multiplikationen lediglich *ungefähr* $\log b$ Multiplikationen (mathematisch präzise: $\mathcal O(\log b)$ Multiplikationen) benötigt. Für hohe Werte von $b$ kann dies durchaus einen Unterschied in der Praxis machen.
### END SOLUTION

e) Modifizieren Sie die aus der Vorlesung bekannte Methode `hanoi` (siehe Vorlesung Sektion 7.5.17) so, dass diese zurückgibt, wie oft eine Scheibe bewegt wurde.   
**Wichtig**: Nicht-rekursive Lösungen werden mit 0 Punkten bewertet, d.h. die folgende Lösung ist nicht legal: 

```Python
def hanoi(hoehe, von, nach, ablage):
    return 2**hoehe-1
```
Beachten Sie weiterhin, dass Sie den Methodenkopf der Funktion `hanoi` nicht verändern dürfen (siehe Tests).   
(3 Punkte)

<!-- grade_id: hanoi_sol_5, points: 0, grade: false, solution: true -->
```python
def hanoi(hoehe, von, nach, ablage):
    # Your Code here
    ### BEGIN SOLUTION
    if hoehe==1:
        #print("Lege Scheibe von", von, "nach", nach)
        return 1
    else:
        m = hanoi(hoehe-1, von, ablage, nach)
        #print("Lege Scheibe von", von, "nach", nach)
        n = hanoi(hoehe-1, ablage, nach, von)
        return m + n + 1
    ### END SOLUTION
```

<!-- grade_id: hanoi_test_5, points: 3, grade: true, solution: false -->
```python
assert(hanoi(5, 'Start', 'Ziel', 'Ablage') == 31)
assert(hanoi(8, 'Start', 'Ziel', 'Ablage') == 255)
### HIDESTART
for i in range(1,10):
    assert(hanoi(i, 'Start', 'Ziel', 'Ablage') == 2**i - 1)
### HIDEEND
```

# Aufgabe 3: Sichtbarkeit von Variablen (5 Punkte)
a) Gegeben ist das folgende Programm:
```Python
a_var = 'global value a'
b_var = 'global value b'

def func():
    a_var = 'local value a'
    print(a_var, '[ a_var inside func() ]')
    print(b_var, '[ b_var inside func() ]')

print(a_var, '[ a_var outside func() ]')
a_func()
print(a_var, '[ a_var outside func() ]')
```

Erläutern Sie anhand dieses Programms und dessen Ausgabe die Begriffe *global scope* und *local scope*.

<!-- grade_id: scope_sol_1, points: 2, grade: true, solution: true -->
### BEGIN SOLUTION
- Global Scope: Variablen, an die außerhalb aller Funktionen zugewiesen wurde, sind im global scope. Dies betrifft hier die Variablen `a_var` und `b_var` aus Zeile 1 und 2.
- Local Scope: Variablen, an die innerhalb eines def zugewiesen wurde, sind im local scope dieser Funktion. Dies betrifft hier die Variable a_var an welche direkt nach Zeile `def func():` der Wert 'local value a' zugewiesen wird. Befindet sich eine zu lesende Variabe nicht im local scope einer Funktion, so wird im global scope nachgeschaut, ob sich dort solch eine Variable befindet. Daher kann die Funktion `func` im obigen Beispiel auch auf die Variable `b_var` lesend zugreifen.

### END SOLUTION

b) Gegeben sei das folgende Programm:
```Python
lst = [1,2,3,4]

def f1():
    lst.append(5)
    
f1()
print(lst)

def f2():
    lst += [6]

f2()
print(lst)
```
Erläutern Sie, warum der Aufruf der Funktion `f1` fehlerfrei gelingt, jedoch nicht der Aufruf der Funktion `f2`.

<!-- grade_id: scope_sol_2, points: 2, grade: true, solution: true -->
### BEGIN SOLUTION
Da sich die Variable `lst` im global Scope befindet, ist diese von Funktionen aus lediglich lesbar, d.h. Zuweisungen an `lst` innerhalb von Funktionen sind verboten.
Der Befehl `lst.append(5)` ist keine Zuweisung an die Variable `lst`, der Befehl `lst += [6]` jedoch schon, weswegen letzterer einen Fehler wirft.
### END SOLUTION

c) Gegeben sei das folgende Programm:

```Python
a_var = 'global value'

def outer():
    a_var = 'enclosed value'

    def inner():
        print(a_var)
    
    inner()

outer()
```

Geben Sie die Ausgabe des Programms an und begründen Sie, wie diese zustande gekommen ist.

<!-- grade_id: scope_sol_3, points: 1, grade: true, solution: true -->
### BEGIN SOLUTION
Der `print`-Befehl befindet sich in der Funktion `inner` und greift auf die Variable `a_var` zu, welche sich im local scope der Funktion `outer` befindet. Dies ist möglich, da `inner` nicht neben, sondern in `outer` definiert wurde.
### END SOLUTION

# Aufgabe 4: Flexible Funktionsparameter (3 Punkte)
Gegeben sei das folgende Programm:

```Python
def f(a, b, c):
    print("a:", a, end=", ")
    print("b:", b, end=", ")
    print("c:", c)

f(**{'a': 1, 'b': 2, 'c': 3})
f(*[1,2,3])
```

Offensichtlich bewirken beide Aufrufe `f(**{'a': 1, 'b': 2, 'c': 3})` und `f(*[1,2,3])` für dieses Programm dasselbe.
Dies heißt jedoch nicht, dass beide Aufrufe in ihrer Semantik äquivalent sind.
Erläutern Sie die semantischen Gemeinsamkeiten und Unterschiede der Aufrufe im Allgemeinen.

<!-- grade_id: funktionsparameter_sol, points: 3, grade: true, solution: true -->
### BEGIN SOLUTION
Der Befehl `f(**{'a': 1, 'b': 2, 'c': 3})` ordnet jedem Parameter der Funktion `f` einen Schlüssel aus dem Dictionary zu. Dabei müssen die Variablennamen exakt den Schlüsseln im Dictionary entsprechen, ansonsten wird ein Fehler geworfen.
Bei dem Befehl `f(*[1,2,3])` werden die Funktionsparameter ihrer Reihenfolge nach auf die Listenelemente zugewiesen.
Beide Befehle haben gemeinsam, dass die Anzahl der Elemente im Dictionary bzw. in der Liste exakt der Anzahl an Funktionsparametern entsprechen muss (sofern keine optionalen Parameter in der Funktion definiert sind). Ist dies nicht der Fall wird ein Fehler geworfen.
### END SOLUTION