<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile). Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue06.ipynb ab (nicht umbenennen, keine Kopie erstellen!). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie  die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Kommentare (1 Punkt!)

Kommentieren Sie Code-Abgaben! Insbesondere muss jede Funktion (Klasse, Methode, ...) einen ```docstring``` haben! Das Vorhandensein von ```docstrings``` auf allen Code-Abgaben wird mit einem Punkt gewertet!


# Abgabetermin: 11.12.2016 um 23:59 Uhr



# Aufgabe 1: Rekursive Funktion umwandeln (3 Punkte)
Gegeben ist die folgende (rekursive) Funktion `berechne_rekursiv`, welche eine natürliche Zahl `zahl` als Funktionsparameter übergeben bekommt.

```python
def berechne_rekursiv(zahl):
    """Eine rekursive Funktion, welche entweder True oder False ausgibt.
   
    Args: 
        zahl (int): eine natürliche Zahl
    Returns: 
        bool: siehe Teilaufgabe b)
    """

    if zahl == 3 or zahl == 6 or zahl == 9:
        return True
    if zahl < 10:
        return False
    sZahl = str(zahl)
    qSum = 0
    for s in sZahl:
        qSum = qSum + int(s)
    return berechne_rekursiv(qSum)
```

### a) (2 Punkte)
Wandeln Sie diese rekursive Funktion in eine iterative um, indem Sie eine Funktion `berechne_iterativ` implementieren. Beide Funktionen sollen äquivalent in ihrer Ausgabe sein, d.h. für eine natürliche Zahl `z` soll `berechne_rekursiv(z) == berechne_iterativ(z)` gelten (siehe Tests). 

<!-- grade_id: transformation_sol_1, points: 0, grade: false, solution: true -->
```python
def berechne_iterativ(zahl):
    # YOUR CODE HERE
    ### BEGIN SOLUTION
    """Eine iterative Funktion, welche entweder True oder False ausgibt.
   
    Args: 
        zahl (int): eine natürliche Zahl
    Returns: 
        bool: siehe Teilaufgabe b)
    """
    while zahl >= 10:
        qSum = 0
        sZahl = str(zahl)
        for s in sZahl:
            qSum = qSum + int(s)
        zahl = qSum
        
    if zahl == 3 or zahl == 6 or zahl == 9:
        return True
    return False
    ### END SOLUTION
```

<!-- grade_id: transformation_test_1, points: 2, grade: true, solution: false -->
```python
z = 102064321020643210206432
assert(berechne_rekursiv(z) == berechne_iterativ(z))
z = z - 1
assert(berechne_rekursiv(z) == berechne_iterativ(z))
z = z - 1
assert(berechne_rekursiv(z) == berechne_iterativ(z))
z = z - 1
assert(berechne_rekursiv(z) == berechne_iterativ(z))
```

### b) (1 Punkt)
Welche Eigenschaft müssen die an die Funktionen `berechne_rekursiv` bzw. `berechne_iterativ` übergebenen Zahlen haben, damit `True` zurückgegeben wird?

<!-- grade_id: transformation_sol_2, points: 1, grade: true, solution: true -->
### BEGIN SOLUTION
Die Zahl muss durch 3 teilbar sein!
### END SOLUTION

# Aufgabe 2: Rekursion (8 Punkte)
### a) (1 Punkt)
Gegeben sind die folgenden Funktionen `f` und `g` in Python:
```Python
def f(a):
    if a == 0:
        print('f')
    else:
        g(a)
        
def g(a):
    if a < 0:
        f(a+1)
    elif a > 0:
        print('g')
        f(a-1)
```

Geben Sie für den Aufruf `f(3)` die Folge der Funktionsaufrufe von `f` und `g` an sowie die Ausgabe des Programms.   

<!-- grade_id: rekursion_sol_1, points: 1, grade: true, solution: true -->
Funktionsaufrufe: $f(3) \rightarrow g(3) \rightarrow f(2) \rightarrow g(2) \rightarrow f(1) \rightarrow g(1) \rightarrow f(0)$   
Ausgabe:   
`g`   
`g`   
`g`   
`f`

### b) (1 Punkt)
Schreiben Sie eine *rekursive* Funktion `anz_ziffern(x)`, welche für eine übergebene natürliche Zahl `x` die Anzahl der Ziffern von `x` zurückgibt. Sehen Sie sich die gegebenen Tests an für Beispiele.   

<!-- grade_id: rekursion_sol_2, points: 0, grade: false, solution: true -->
```python
def anz_ziffern(x):
    # YOUR CODE HERE
    ### BEGIN SOLUTION
    """Gibt für die übergebene natürliche Zahl die Anzahl ihrer Ziffern zurück.
   
    Args: 
        x (int): eine natürliche Zahl
    Returns: 
        int: die Anzahl der Ziffern von x
    """
    
    if x // 10 == 0:
        return 1
    return anz_ziffern(x//10) + 1
    ### END SOLUTION
```

<!-- grade_id: rekursion_test_2, points: 1, grade: true, solution: false -->
```python
assert(anz_ziffern(2) == 1)
assert(anz_ziffern(42) == 2)
assert(anz_ziffern(873246187456216) == 15)
### HIDESTART
l = [10**x for x in range(10)]
for i in l:
    assert(anz_ziffern(i) == l.index(i)+1)
### HIDEEND
```

### c) (2 Punkte)
Gegeben sei die folgende Funktion `f`:
\begin{equation}
   f(a,b) =
   \begin{cases}
     a & \text{falls } b = 1 \\
     a \cdot f(a, (b-1)/2) \cdot f(a, (b-1)/2)& \text{falls } b > 1 \text{ und } b \text{ ungerade} \\
     f(a, b/2) \cdot f(a, b/2)& \text{sonst} \\
   \end{cases}
\end{equation}

Implementieren Sie diese Funktion in Python, indem Sie die vorgegebene Funktion `f(a,b)` vervollständigen. Sie können davon ausgehen, dass `b` immer eine natürliche Zahl ist.

<!-- grade_id: rekursion_sol_3, points: 0, grade: false, solution: true -->
```python
def f(a,b):
    # YOUR CODE HERE
    ### BEGIN SOLUTION
    """Berechnet rekursiv für die beiden Parameter a und b den Wert a**b.
   
    Args: 
        a (int): eine Zahl >= 0
        b (int): eine natürliche Zahl
    Returns: 
        int: a**b
    """

    if b == 1:
        return a
    if b % 2 == 1:
        return a*f(a,(b-1)//2)*f(a,(b-1)//2) 
    return f(a, b//2)*f(a, b//2)
    ### END SOLUTION
```

<!-- grade_id: rekursion_test_3_1, points: 1, grade: true, solution: false -->
```python
assert(f(13,12) == 23298085122481)
assert(f(17,7) == 410338673)
assert(f(9,21) == 109418989131512359209)
### HIDESTART
# Teste den ersten Fall:
for a in range(1,10):
    assert(f(a,1) == a)
    
# Teste den zweiten Fall:
for a in range(1,10):
    for b in [3,5,7,9,11,13,15,17,19]:
        assert(f(a,b) == a**b)
### HIDEEND
```

<!-- grade_id: rekursion_test_3_2, points: 1, grade: true, solution: false -->
```python
### HIDESTART
# Teste den dritten Fall:
for a in range(1,10):
    for b in [2,4,6,8,10,12,14,16,18,20]:
        assert(f(a,b) == a**b)
### HIDEEND
```

### d) (1 Punkt)
Wie kann die durch `f` berechnete Funktion einfach ausgedrückt werden?
Welche Vorteile hat dies im Vergleich zu einer (naiven) iterativen Berechnung?

<!-- grade_id: rekursion_sol_4, points: 1, grade: true, solution: true -->
### BEGIN SOLUTION
Die Funktion berechnet $a^b$ für ein $a \geq 0$ und eine natürliche Zahl $b$.
Die Berechnung wird schneller, da man nun anstelle von $b$ Multiplikationen lediglich *ungefähr* $\log b$ Multiplikationen (mathematisch präzise: $\mathcal O(\log b)$ Multiplikationen) benötigt. Für hohe Werte von $b$ kann dies durchaus einen Unterschied in der Praxis machen.
### END SOLUTION

### e) (3 Punkte)
Modifizieren Sie die aus der Vorlesung bekannte Methode `hanoi` (siehe Vorlesung Sektion 7.5.17) so, dass diese zurückgibt, wie oft eine Scheibe bewegt wurde.   
**Wichtig**: Nicht-rekursive Lösungen werden mit 0 Punkten bewertet, d.h. die folgende Lösung ist nicht legal: 

```Python
def hanoi(hoehe, von, nach, ablage):
    return 2**hoehe-1
```
Beachten Sie weiterhin, dass Sie den Methodenkopf der Funktion `hanoi` nicht verändern dürfen (siehe Tests).   

<!-- grade_id: hanoi_sol_5, points: 0, grade: false, solution: true -->
```python
def hanoi(hoehe, von, nach, ablage):
    # YOUR CODE HERE
    ### BEGIN SOLUTION
    """Modifikation der Methode hanoi aus der Vorlesung, welche die Anzahl der Scheibenbewegungen zurückgibt.
   
    Args: 
        hoehe (int): die initiale hoehe des Startstapels
        von (str): name des Startstapels
        nach (str): name des Zielstapels
        ablage (str): name des Hilfsstapels
    Returns: 
        int: Anzahl der Scheibenbewegungen um alle Scheiben vom Startstapel zum Zielstapel zu bewegen
    """
    if hoehe==1:
        #print("Lege Scheibe von", von, "nach", nach)
        return 1
    else:
        m = hanoi(hoehe-1, von, ablage, nach)
        #print("Lege Scheibe von", von, "nach", nach)
        n = hanoi(hoehe-1, ablage, nach, von)
        return m + n + 1
    ### END SOLUTION
```

<!-- grade_id: hanoi_test_5, points: 3, grade: true, solution: false -->
```python
assert(hanoi(5, 'Start', 'Ziel', 'Ablage') == 31)
assert(hanoi(8, 'Start', 'Ziel', 'Ablage') == 255)
### HIDESTART
for i in range(1,10):
    assert(hanoi(i, 'Start', 'Ziel', 'Ablage') == 2**i - 1)
### HIDEEND
```

# Aufgabe 3: Sichtbarkeit von Variablen (5 Punkte)
### a) (2 Punkte)
Gegeben ist das folgende Programm:
```Python
a_var = 'global value a'
b_var = 'global value b'

def func():
    a_var = 'local value a'
    print(a_var, '[ a_var inside func() ]')
    print(b_var, '[ b_var inside func() ]')

print(a_var, '[ a_var outside func() ]')
a_func()
print(a_var, '[ a_var outside func() ]')
```

Erläutern Sie anhand dieses Programms und dessen Ausgabe die Begriffe *global scope* und *local scope*.   

<!-- grade_id: scope_sol_1, points: 2, grade: true, solution: true -->
### BEGIN SOLUTION
- Global Scope: Variablen, an die außerhalb aller Funktionen zugewiesen wurde, sind im global scope. Dies betrifft hier die Variablen `a_var` und `b_var` aus Zeile 1 und 2.
- Local Scope: Variablen, an die innerhalb eines def zugewiesen wurde, sind im local scope dieser Funktion. Dies betrifft hier die Variable a_var an welche direkt nach Zeile `def func():` der Wert 'local value a' zugewiesen wird. Befindet sich eine zu lesende Variabe nicht im local scope einer Funktion, so wird im global scope nachgeschaut, ob sich dort solch eine Variable befindet. Daher kann die Funktion `func` im obigen Beispiel auch auf die Variable `b_var` lesend zugreifen.

### END SOLUTION

### b) (2 Punkte)
Gegeben sei das folgende Programm:
```Python
lst = [1,2,3,4]

def f1():
    lst.append(5)
    
f1()
print(lst)

def f2():
    lst += [6]

f2()
print(lst)
```
Erläutern Sie, warum der Aufruf der Funktion `f1` fehlerfrei gelingt, jedoch nicht der Aufruf der Funktion `f2`.   

<!-- grade_id: scope_sol_2, points: 2, grade: true, solution: true -->
### BEGIN SOLUTION
Da sich die Variable `lst` im global Scope befindet, ist diese von Funktionen aus lediglich lesbar, d.h. Zuweisungen an `lst` innerhalb von Funktionen sind verboten.
Der Befehl `lst.append(5)` ist keine Zuweisung an die Variable `lst`, der Befehl `lst += [6]` jedoch schon, weswegen letzterer einen Fehler wirft.
### END SOLUTION

### c) (1 Punkt)
Gegeben sei das folgende Programm:

```Python
a_var = 'global value'

def outer():
    a_var = 'enclosed value'

    def inner():
        print(a_var)
    
    inner()

outer()
```

Geben Sie die Ausgabe des Programms an und begründen Sie, wie diese zustande gekommen ist.   

<!-- grade_id: scope_sol_3, points: 1, grade: true, solution: true -->
### BEGIN SOLUTION
Der `print`-Befehl befindet sich in der Funktion `inner` und greift auf die Variable `a_var` zu, welche sich im local scope der Funktion `outer` befindet. Dies ist möglich, da `inner` nicht neben, sondern in `outer` definiert wurde.
### END SOLUTION

# Aufgabe 4: Flexible Funktionsparameter (3 Punkte)
Gegeben sei das folgende Programm:

```Python
def f(a, b, c):
    print("a:", a, end=", ")
    print("b:", b, end=", ")
    print("c:", c)

f(**{'a': 1, 'b': 2, 'c': 3})
f(*[1,2,3])
```

Offensichtlich bewirken beide Aufrufe `f(**{'a': 1, 'b': 2, 'c': 3})` und `f(*[1,2,3])` für dieses Programm dasselbe.
Dies heißt jedoch nicht, dass beide Aufrufe in ihrer Semantik äquivalent sind.
Erläutern Sie die semantischen Gemeinsamkeiten und Unterschiede der Aufrufe im Allgemeinen.

<!-- grade_id: funktionsparameter_sol, points: 3, grade: true, solution: true -->
### BEGIN SOLUTION
Der Befehl `f(**{'a': 1, 'b': 2, 'c': 3})` ordnet jedem Parameter der Funktion `f` einen Schlüssel aus dem Dictionary zu. Dabei müssen die Variablennamen exakt den Schlüsseln im Dictionary entsprechen, ansonsten wird ein Fehler geworfen.
Bei dem Befehl `f(*[1,2,3])` werden die Funktionsparameter ihrer Reihenfolge nach auf die Listenelemente zugewiesen.
Beide Befehle haben gemeinsam, dass die Anzahl der Elemente im Dictionary bzw. in der Liste exakt der Anzahl an Funktionsparametern entsprechen muss (sofern keine optionalen Parameter in der Funktion definiert sind). Ist dies nicht der Fall wird ein Fehler geworfen.
### END SOLUTION

<!-- grade_id: docstring_tests, points: 1, grade: true, solution: false -->
```python
assert(berechne_iterativ.__doc__ != None)
assert(anz_ziffern.__doc__ != None)
assert(f.__doc__ != None)
assert(hanoi.__doc__ != None)
```