<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->
Bitte hier eintragen (s.u)!

# Abgabehinweise 

Beachten Sie unbedingt diese Hinweise, sonst erhalten Sie keine Punkte aus dieser Abgabe! Für Details siehe z.B. [Folien der nullten Zentralübung](http://groups.uni-paderborn.de/fg-karl/lehre/ws1617/gp1/hue/gp1-zue0.pdf)

## Namen und Matrikelnummern 

Tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle **oberhalb** dieser Zelle ein.
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Zwischen Matrikelnummer und E-Mail-Adresse *muss genau ein* Leerzeichen stehen; sonst keine Leerzeichen, Tabs, Spiegelstriche, oder ähnliches in einer Zeile. Wir empfehlen *dringend* die Benutzung der Uni-Paderborn E-Mail-Adressen! 

<!--- Wenn Sie diese Ansicht verwenden (was Sie nicht tun sollten; klicken Sie am besten einmal auf "Run cell"), dann steht die Formatierung hier richtig:

Format:

MatrikelNummer E-Mail-Adresse
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
234567 testStudent2@mail.upb.de
345678 vorname.nachname@mail.upb.de

... und nicht hier: -->

Format:

MatrikelNummer E-Mail-Adresse
<br \>
MatrikelNummer2 E-Mail-Adresse2

z.B.:

123456 testStudent@mail.upb.de
<br \>
234567 testStudent2@mail.upb.de
<br \>
345678 vorname.nachname@mail.upb.de

Hinweis: Wenn Sie das so in die obere Zelle eingeben, wird in der ausgeführten Markdownzelle (Strg+Enter) alles in einer Zeile angezeigt. Das ist egal; wichtig ist Formatierung der Eingabe.


## Abgabe mit Gruppenaccount 

Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen (gp1\_16\_...), den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 

## Abgabe: Submit 

Es reicht nicht, nur das Übungsblatt zu verändern. Sie müssen unter *Assignments* auf *Submit* clicken (oder entsprechend über die Kommandozeile), für das entsprechende Übungsblatt! (Wenn Sie Blatt 9 abgeben wolllen hilft es nicht, bei Blatt 1 auf Submit zu drücken!) Für Details siehe Beschreibung Abgabeprozess. 

## Dateinamen

Geben Sie Ihre Lösung in der vom Server erhaltenen Datei hue09.ipynb ab (nicht umbenennen, keine Kopie erstellen, keine anderen Dateien in das Verzeichnis legen, auch nicht Dateiname hue9.ipynb !). Sonst kann eine Bewertung nicht stattfinden.

## Struktur des Notebooks 

Fügen Sie keine Zellen hinzu und löschen Sie keine Zellen.
Ändern Sie nicht den Typ einer Zelle.
Geben Sie Lösungen nur in den Lösungszellen ab. Änderungen in den anderen Zellen werden nach Abgabe automatisch rückgängig gemacht.

## Code-Zellen 

Entfernen Sie die Zeilen mit Inhalt "raise NotImplementedError()" aus den Zellen, in die Sie Ihre Lösungen schreiben. *Ersetzen* Sie diese Zeilen durch Ihre eigene Lösungen. 

## Nie ```input```

Benutzen Sie *niemals* die Funktion `input()`. Das verhindert die automatische Auswertung Ihrer Abgabe und führt zu 0 Punkten für das gesamte Blatt! 

## Kommentare 

Kommentieren Sie Code-Abgaben! Insbesondere *muss* jede Funktion (Klasse, Methode, ...) einen sinnvollen ```docstring``` haben! Bei fehlendem docstring werden automatisch Punkte abgezogen! 


# Abgabetermin: 15.1.2017 um 23:59 Uhr

# Aufgabe 1: Exceptions
## a)
Gegeben ist das (sinnfreie) Programm in der nachfolgenden Codezelle (die Kommentierung wurde aufgrund der Sinnfreiheit des Programms ausgelassen).
Begründen Sie für jede `print`-Anweisung (`print('1')` bis `print('8')`), warum diese ausgeführt wird, oder warum nicht.

```python
def func(x, y):
    try:
        r = x / y
    except ZeroDivisionError:
        print('1')
    else:
        try:
            print('2', result)
        except NameError:
            print('3')
    finally:
        print('4')
        raise NameError()


try:
    func(2, 1)
except ZeroDivisionError:
    print('5')
except NameError:
    print('6')
else:
    print('7', r)
finally:
    print('8')
```

<!-- grade_id: exceptions_sol_1, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## b)
Schreiben Sie Funktionen `f`, `g` und `h`, welche die jeweiligen Fehler werfen, die in den Tests gefangen werden sollen. Es ist ihnen nicht gestattet den Befehl `raise` zu verwenden.   
*Tipp*: Informieren Sie sich über die einzelnen Fehler in der Python-Dokumentation.

<!-- grade_id: exception_sol_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

# YOUR CODE HERE
raise NotImplementedError()

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: exception_test_2_1, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
try:
    f()
    assert(False) # Darf nicht ausgeführt werden!
except KeyError:
    assert(True) # Muss ausgeführt werden!
```

<!-- grade_id: exception_test_2_2, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
try:
    g()
    assert(False) # Darf nicht ausgeführt werden!
except TypeError:
    assert(True) # Muss ausgeführt werden!
```

<!-- grade_id: exception_test_2_3, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
try:
    h()
    assert(False) # Darf nicht ausgeführt werden!
except RecursionError:
    assert(True) # Muss ausgeführt werden!
```

## c)
Vereinfachen Sie das folgende Programm mit Hilfe von `try`/`except`.
Schreiben Sie dazu Ihren Code in eine Funktion namens `print_object_einfach(o)`.
Diese soll dieselbe Funktionalität wie `print_object(o)` haben.

```python
def print_object(o):
    """Gibt das Objekt auf der Konsole aus, falls es mit print ausgegeben werden kann, ansonsten wird 'unprintable object'
    auf der Konsole ausgegeben.
    
    Args:
        o: Das auszugebende Objekt
    """
    if isinstance(o, str):
        print(o)
    elif isinstance(o, dict):
        print(o)
    elif isinstance(o, list):
        print(o)
    # 97 elifs später...
    else:
        print("unprintable object")
```

<!-- grade_id: exception_sol_3, points: 2, grade: true, solution: true -->
```python
def print_object_einfach(o):
    # YOUR CODE HERE
    raise NotImplementedError()
```

## d)
Schreiben Sie eine Funktion `verarbeite_eingabe(zahl, von, bis)`, welche überprüft, ob die Zahl `zahl` sich im Interval `[von, bis]` befindet.
Ist dies der Fall, dann soll die Zahl auf der Konsole ausgegeben werden.
Sollte dies nicht der Fall sein, so soll ein `BereichsError` geworfen werden, welcher die Zahl `zahl`, sowie die Intervallgrenzen speichert.
Erstellen Sie dazu eine eigene Exception `BereichsError`, welche Sie dann mit `raise` werfen können.    

<!-- grade_id: exception_sol_4, points: 3, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# YOUR CODE HERE
raise NotImplementedError()
### END SOLUTION
```

<!-- grade_id: exception_test_4, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
try:
    verarbeite_eingabe(101, 1, 100)
    assert(False) # Darf nicht ausgeführt werden!
except BereichsError:
    assert(True) # Muss ausgeführt werden!
```

## e)
Rufen Sie nun die Funktion `verarbeite_eingabe` testweise zweimal auf: Der erste Aufruf soll keinen `BereichsError` werfen, der zweite jedoch schon.
Fangen Sie den `BereichsError` mithilfe eines `try`/`except`-Block ab.   

<!-- grade_id: exception_sol_5, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# Aufgabe 2: Uhr mit Properties
Schreiben Sie eine Klasse `Uhr`, welche über Properties `minuten` und `stunden` verfügen soll.
Implementieren Sie dazu die jeweiligen getter/setter-Methoden und verwenden Sie diese als Properties für die Attribute `minuten` bzw. `stunden`.
Bedenken Sie dazu, dass ein Setzen der Minuten (60 Minuten sind eine Stunde!) auch eventuell eine Erhöhung der Stunden implizieren kann.
Insbesondere sollen in der Property `minuten` nur Zahlen zwischen `0` und `59` gespeichert werden.
Sobald die Stunden `23` betragen und um eins erhöht werden, sollen sie wieder auf `0` zurückgesetzt werden.
Implementieren Sie außerdem eine Methode `tic(self)`, welche den Minutenzähler der Uhr um `1` erhöht.
Zu Beginn sollen sowohl Stunden- als auch Minutenzähler auf 0 gestellt werden.

<!-- grade_id: properties_sol_1, points: 0, grade: false, solution: true -->
```python
class Uhr:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: properties_test_1, points: 3, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
c = Uhr()

assert(c.minuten == 0 and c.stunden == 0) # Zu Beginn: Stunden- und Minutenzähler auf 0

for i in range(24):
    assert(c.stunden == i)
    for j in range(60):
        assert(c.stunden == i and c.minuten == j)
        c.tic()

assert(c.minuten == 0 and c.stunden == 0)
```

# Aufgabe 3: Schach
In dieser Aufgabe sollen Sie eine einfache Version eines Schachspiels implementieren.
Die Klassen `Feld` und `Figur` sind Ihnen dazu bereits vorgegeben.   

*Hinweis*: Die in dieses Notebook eingebundenen Bilder werden Ihnen nur angezeigt, wenn Sie mit dem Internet verbunden sind.

## a)
Schreiben Sie eine Klasse `Schachbrett`, welche die folgenden Anforderungen erfüllt:
 - Im Konstruktor soll ein Spielfeld aus 64 Feldern (8x8, 8 Zeilen und 8 Spalten, siehe Bild, siehe Tests) anlegt werden.
![Imgur](http://groups.uni-paderborn.de/fg-ti/teaching/gp1/brett.png)

 - `get_feld(x,y)`: Liefert das Objekt der Klasse `Feld`, welches die übergebenen Koordinaten hat. Falls `x` oder `y` keine korrekten Koordinaten sind (d.h. `not x in range(1,9)` oder `not y in range(1,9)`), dann soll ein `ValueError` geworfen werden.
 - `setze_figur(figur, x, y)`: Setzt die übergebene Figur auf das Feld mit den übergebenen Koordinaten. Sie können davon ausgehen, dass das Feld nicht besetzt ist und die Figur vorher auf keinem Feld gestanden hat. Beachten Sie, dass die Beziehung zwischen `Feld` und `Figur` bidirektional sein soll, d.h. sobald das Attribut `feld` einer Figur verändert wird, muss auch das Attribut `figur` des entsprechenden Feldes geändert werden und umgekehrt (siehe Tests).

```python
class Feld:
    """Ein einzelnes Feld eines Schachbretts"""
    
    def __init__(self, koordinaten):
        """Konstruktor"""
        self.koordinaten = koordinaten
        self.figur = None
    
    def __eq__(self, feld):
        """Zum Vergleich mit anderem Feld"""
        return feld.koordinaten[0] == self.koordinaten[0] and feld.koordinaten[1] == self.koordinaten[1]
    
class Figur:
    """Eine Figur, welche eine Farbe hat und auf einem Feld des Schachbrettes stehen kann."""
    
    def __init__(self, color):
        """Konstruktor"""
        self.feld = None
        self.color = color
        
    def berechne_zuege(self, brett):
        """Liefert eine Liste mit Feldern, welche alle Felder enthält, auf die sich die Figur ausgehend von ihrer 
        aktuellen Position in einem Zug bewegen kann. Muss von erbenden Klassen überschrieben werden!"""
        pass
```

<!-- grade_id: schach_sol_1, points: 0, grade: false, solution: true -->
```python
class Schachbrett:
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: schach_test_1, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
s = Schachbrett()
for i in range(1,9):
    for j in range(1,9):
        feld = s.get_feld(i,j)
        assert(feld.koordinaten[0] == i and feld.koordinaten[1] == j)

# Tests falsche Koordinaten in get_feld:
try:
    s.get_feld(0,1)
    assert(False)
except ValueError:
    assert(True)
    
k = Figur("white")
s.setze_figur(k, 4, 5)
assert(s.get_feld(4, 5).figur == k and k.feld == s.get_feld(4, 5))
```

## b)
Schreiben Sie eine Klasse `Koenig`, welche von der Klasse `Figur` erbt.
Diese Klasse soll die Methode `berechne_zuege` der Klasse `Figur` überschreiben.
Die Methode `berechne_zuege` soll eine Liste zurückgeben, welche alle Felder enthält, auf die sich der König ausgehend von seiner aktuellen Position in einem Zug bewegen kann.
Gehen Sie davon aus, dass alle anderen 63 Felder nicht mit anderen Figuren belegt sind.


Hinweis: Der König kann sich im Schach immer nur um ein Feld weiter bewegen, allerdings in jede beliebige Richtung: Nach oben, unten, links, rechts oder in die Diagonalen (wir ignorieren hier die Möglichkeit der Rochade!).   
Beispiel: Das nachfolgende Bild gibt die möglichen Züge des Königs an, wenn dieser sich auf dem Feld mit den Koordinaten (5,4) befindet.
 ![Imgur](http://groups.uni-paderborn.de/fg-ti/teaching/gp1/koenig.png)
Beachten Sie Randfälle.

<!-- grade_id: schach_sol_2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: schach_test_2, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
s = Schachbrett()
k = Koenig("white")
s.setze_figur(k, 5, 4)
felder = k.berechne_zuege(s)
assert(s.get_feld(5, 3) in felder)
assert(s.get_feld(5, 5) in felder)
assert(s.get_feld(6, 3) in felder)
assert(s.get_feld(6, 4) in felder)
assert(s.get_feld(6, 5) in felder)
assert(s.get_feld(4, 4) in felder)
assert(s.get_feld(4, 5) in felder)
assert(s.get_feld(4, 3) in felder)
assert(len(felder) == 8)
```

## c)
Schreiben Sie eine weitere Klasse `Koenig_Advanced`, welche von der Klasse `Koenig` erbt und in der Methode `berechne_zuege` nun alle Felder zurückgibt, auf die der König in einem Zug ziehen kann und welche zusätzlich NICHT von einer Figur mit derselben Farbe belegt sind.

<!-- grade_id: schach_sol_3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: schach_test_3, points: 2, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
s = Schachbrett()
k = Koenig_Advanced("white")
s.setze_figur(k, 5, 4)

# Belege Felder:
s.setze_figur(Figur("white"), 5, 5)
s.setze_figur(Figur("black"), 6, 4)
s.setze_figur(Figur("white"), 4, 5)
s.setze_figur(Figur("black"), 4, 3)

felder = k.berechne_zuege(s)
assert(s.get_feld(5, 3) in felder)
assert(not s.get_feld(5, 5) in felder)
assert(s.get_feld(6, 3) in felder)
assert(s.get_feld(6, 4) in felder)
assert(s.get_feld(6, 5) in felder)
assert(s.get_feld(4, 4) in felder)
assert(not s.get_feld(4, 5) in felder)
assert(s.get_feld(4, 3) in felder)
assert(len(felder) == 6)
```

<!-- grade_id: docu_test, points: 1, grade: true, solution: false -->
```python
# Diese Zelle bitte unverändert lassen!
# Alle Dokumentationen vorhanden? 

l = []
        
# Aufgabe 1:
l.append(("f", None))
l.append(("g", None))
l.append(("h", None))

l.append(("verarbeite_eingabe", None))
l.append(("BereichsError", None))
l.append(("print_object_einfach", None))

# Aufgabe 2:
l.append(("Uhr", None))
l.append(("Uhr", "__init__"))
l.append(("Uhr", "get_minuten"))
l.append(("Uhr", "get_stunden"))
l.append(("Uhr", "set_minuten"))
l.append(("Uhr", "set_stunden"))
l.append(("Uhr", "tic"))

# Aufgabe 3
l.append(("Schachbrett", None))
l.append(("Schachbrett", "__init__"))
l.append(("Schachbrett", "get_feld"))
l.append(("Schachbrett", "setze_figur"))
l.append(("Koenig", None))
l.append(("Koenig", "berechne_zuege"))
l.append(("Koenig_Advanced", None))
l.append(("Koenig_Advanced", "berechne_zuege"))

for a in l:
    try:
        attr = None
        if a[1] == None:
            attr = eval(a[0])
            assert(attr.__doc__ != None)
            print(attr.__doc__ != None, "-->", a[0])    
        else:
            attr = getattr(eval(a[0]), a[1])
            assert(attr.__doc__ != None)
            print(attr.__doc__ != None, "-->", a[0] + "." + a[1])    
    except AttributeError:
        print("Not Implemented", "-->", a[0] + "." + a[1])
    except AssertionError:
        print(attr.__doc__ != None, "-->", a[0] + "." + a[1])
```