<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->


Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 13.11.2016 um 23:59 Uhr

## Aufgabe 1: 
a) (1 Punkt) Gegeben ist folgendes Programm:
```Python
a = 5
b = 17
a = b + 2
print(a, b)
```
Was gibt das Programm aus? Erklären Sie kurz.

<!-- grade_id: 1, points: 1, grade: true, solution: true -->
    ### BEGIN SOLUTION 
Das Programm gibt 19 17 aus. 
Grund: der Wert von ```a``` wird von ```b + 2``` überschrieben, wobei ```b = 17```
    ### END SOLUTION

b) (1 Punkt) Sie ändern in der dritten Zeile $a$ auf $b$. Was gibt das Programm nun aus?  Erklären Sie kurz.

```Python
a = 5
b = 17
b = b + 2
print(a, b)
```

<!-- grade_id: 3, points: 1, grade: true, solution: true -->
    ### BEGIN SOLUTION 
Das Programm gibt 5 19 aus. 
Grund: ```a``` bleibt unverändert und der Wert von ```b``` wird von ```b + 2``` überschrieben, wobei ```b = 17```. 
    ### END SOLUTION

c) (1 Punkt) Sie vertauschen die erste mit der zweiten Zeile. Was gibt das Programm nun aus?  Erklären Sie kurz.

```Python
b = 17
a = 5
a = b + 2
print(a, b)
```

<!-- grade_id: 6, points: 1, grade: true, solution: true -->
    ### BEGIN SOLUTION 
Das Programm gibt 19 17 aus. 
Grund: ```b``` bleibt unverändert und der Wert von ```a``` wird von ```b + 2``` überschrieben, wobei ```b = 17```. 
    ### END SOLUTION

d) (1 Punkt) Sie ändern in der dritten Zeile $a$ auf $d$. Was gibt das Programm nun aus? Erklären Sie kurz.

```Python
b = 17
a = 5
d = b + 2
print(a, b)
```

<!-- grade_id: 8, points: 1, grade: true, solution: true -->
    ### BEGIN SOLUTION 
Das Programm gibt 5 17 aus. 
Grund: ```a``` und ```b``` bleiben unverändert.
    ### END SOLUTION

e) (1 Punkt) Sie ändern in der ersten UND dritten Zeile $a$ auf $d$. Was gibt das Programm nun aus?  Erklären Sie kurz.

```Python
d = 5
b = 17
d = b + 2
print(a, b)
```


<!-- grade_id: 10, points: 1, grade: true, solution: true -->
    ### BEGIN SOLUTION 
Das Programm ist fehlerhaft, da die Variable ```a``` noch nicht definiert ist.
    ### END SOLUTION

## Aufgabe 2: Funktionen

a) (2 Punkte) Schreiben Sie vier Funktionen namens ```add, multiply, divide``` und ```exponentiate```, die zwei als Parameter übergebene Werte miteinander addieren, multiplizieren, dividieren und exponieren. Das entsprechende Ergebnis soll von den Funktionen zurückgegeben werden.

<!-- grade_id: 21, points: 0, grade: false, solution: true -->
```python
def add(a, b):
    ### BEGIN SOLUTION 
    print('ADDING ',a,' + ',b)
    out = a + b
    return out
    ### END SOLUTION

def multiply(a, b):
    ### BEGIN SOLUTION 
    print('MULTIPLYING ',a,' * ',b)
    out = a * b
    return out
    ### END SOLUTION
    
def divide(a, b):
    ### BEGIN SOLUTION 
    print('DIVIDING ',a,' / ',b)
    out = a / b
    return out
    ### END SOLUTION
    
def exponentiate(a, b):
    ### BEGIN SOLUTION 
    print('EXPONENTIATING ',a,' ** ',b)
    out = a ** b
    return out
    ### END SOLUTION
```

<!-- grade_id: 41, points: 2, grade: true, solution: false -->
```python
assert(add(4,6)==10)
assert(multiply(3,3)==9)
assert(divide(6,3)==2)
assert(exponentiate(2,3)==8)
```

b) (2 Punkte) Schreiben Sie anschließend eine Funktion namens totalMoney, die drei Parameter k, n und p (in dieser Reihenfolge) übergeben bekommt und welche das Ergebnis der folgenden Rechnung zurückgibt.

$$K \left(1 + \frac{p}{100} \right)^n,$$

Verwenden Sie in dieser Aufgabe ausschließlich die in Teilaufgabe a) definierten Funktionen und keine Python-Rechenoperatoren.

<!-- grade_id: 12, points: 0, grade: false, solution: true -->
```python
def totalMoney(k,n,p):
    ### BEGIN SOLUTION 
    out = multiply(k,exponentiate(add(divide(p,100),1),n))
    return out
    ### END SOLUTION
```

<!-- grade_id: 81, points: 2, grade: true, solution: false -->
```python
assert(totalMoney(1000,3,5)==1157.6250000000002)
```

Output:
```
DIVIDING  5  /  100
ADDING  0.05  +  1
EXPONENTIATING  1.05  **  3
MULTIPLYING  1000  *  1.1576250000000001
```

## Aufgabe 3: Datenbearbeitung

a) (1 Punkt) Schreiben Sie ein Programm, das eine Temperatur von Grad Celsius nach Grad Fahrenheit umrechnen kann. Geben Sie das Ergebnis auf dem Bildschirm aus, z.B. "XX°C sind XX°F".

<!-- grade_id: 85, points: 1, grade: true, solution: true -->
```python
    ### BEGIN SOLUTION 
Celsiusdegree = 30
Fahrenheit = Celsiusdegree*1.8 + 32
print(Celsiusdegree, '°C sind', Fahrenheit, '°F')

Fahrenheit = 86
Celsiusdegree = (Fahrenheit - 32) / 1.8
print(Fahrenheit, '°F sind', Celsiusdegree, '°C')
    ### END SOLUTION
```

Output:
```
30 °C sind 86.0 °F
86 °F sind 30.0 °C
```

b) (1 Punkt) Schreiben Sie ein Programm, das den Umfang und die Fläche eines Kreises anhand seines Durchmessers berechnen kann. Geben Sie das Ergebnis auf dem Bildschirm aus, z.B. "Die Fläche beträgt XX mm2 und der Umfang beträgt XX mm".

<!-- grade_id: 87, points: 1, grade: true, solution: true -->
```python
    ### BEGIN SOLUTION 
diameter = 10
perimeter = 3.14 * diameter
surface = 3.14*(diameter/2)**2
print("Die Fläche beträgt", surface, "mm2 und der Umfang beträgt", perimeter, "mm")
    ### END SOLUTION
```

Output:
```
Die Fläche beträgt 78.5 mm2 und der Umfang beträgt 31.400000000000002 mm
```

c) (1 Punkt) Schreiben Sie Ihr aktuelles Alter, ein späteres Alter und wieviel Geld Sie pro Tag benötigen in drei verschiedene Variablen. 
Berechnen Sie anhand der Daten, wieviel Geld Sie von ihrem letzten Geburtstag an bis zum Erreichen des späteren Alters benötigen.
Geben Sie das Resultat aus, z.B. "Bis zum XX-ten Lebensjahr benötigen Sie XX Euro".

Hinweis: Sie dürfen dabei zur Vereinfachung annehmen, dass jedes Lebensjahr genau 365 Tage hat.

<!-- grade_id: 89, points: 2, grade: true, solution: true -->
```python
    ### BEGIN SOLUTION 
actualAge = 24
futureAge = 50
dailyExpense = 30
costs = (futureAge - actualAge) * dailyExpense * 365
print('Bis zum ', futureAge, '-ten Lebensjahr benötigen Sie ', costs, ' Euro.', sep='')
    ### END SOLUTION
```

Output:
```
Bis zum 50-ten Lebensjahr benötigen Sie 284700 Euro.
```

d) (2 Punkte) Schreiben Sie eine Funktion, die die Eckpunkte von einem Dreieck übergeben bekommt. Die Funktion soll dabei die Seitenlängen berechnen und den entsprechenden Umfang zurückgeben. 

<!-- grade_id: 82, points: 0, grade: false, solution: true -->
```python
def calcTriangle(x1, y1, x2, y2, x3, y3):
    ### BEGIN SOLUTION
    d1 = ((x1 - x2)**2 + (y1 - y2)**2)**0.5
    d2 = ((x1 - x3)**2 + (y1 - y3)**2)**0.5
    d3 = ((x2 - x3)**2 + (y2 - y3)**2)**0.5
    perimeter = d1 + d2 +d3
    ### END SOLUTION
    return perimeter

perimeter = calcTriangle(0,0,1,1,2,2)
print(perimeter)
```

Output:
```
5.656854249492381
```

<!-- grade_id: 83, points: 2, grade: true, solution: false -->
```python
assert(calcTriangle(0,0,1,1,2,2)==5.656854249492381)
```

## Aufgabe 4: Mehrfache Eingabe- und Rückgabewerte

b) (2 Punkte) Schreiben Sie eine Funktion, die eine Zeitdauer in Tagen (d), Stunden (h), Minuten (m) und Sekunden (s) übergeben bekommt und die Gesamtdauer in Sekunden zurückgibt.

<!-- grade_id: 74, points: 0, grade: false, solution: true -->
```python
def convDate(d, h, m, s):
    ###BEGIN SOLUTION###
    s += d*60*60*24 + h*60*60 + m *60
    ###END SOLUTION####
    return s

    
s = convDate(0, 1, 1, 10)    
print("seconds: ", s)   
```

Output:
```
seconds:  3670
```

<!-- grade_id: 75, points: 2, grade: true, solution: false -->
```python
s = convDate(0, 1, 1, 10)   
assert(s==3670)
```

a) (3 Punkte)  Schreiben Sie eine weitere Funktion, die die umgekehrte Funktionalität hat. 
Die Funktion soll dabei den Gesamtdauer in Sekunden übergeben bekommen und in Tagen (d), Stunden (h), Minuten (m) und Sekunden (s) umrechnen und zurückgeben. 

<!-- grade_id: 71, points: 0, grade: false, solution: true -->
```python
def convSec(s):
    ###BEGIN SOLUTION###
    d = s // (60 * 60 * 24);
    s -= d * (60 * 60 * 24);
    h = s // (60 * 60);
    s -= h * (60 * 60);
    m = s // 60;
    s -= m*60
    ###END SOLUTION####
    return d, h, m, s

    
d, h, m, s = convSec(3670)    
print("days: ", d, " hours: ", h," mins: ", m, "secs: ", s)   
    
```

Output:
```
days:  0  hours:  1  mins:  1 secs:  10
```

<!-- grade_id: 72, points: 3, grade: true, solution: false -->
```python
d, h, m, s = convSec(3670)   
assert(d==0)
assert(h==1)
assert(m==1)
assert(s==10)
```