<!-- grade_id: matrikel, points: 0, grade: true, solution: true -->


Bitte tragen Sie Ihre Matrikelnummern und E-Mail-Adressen *zeilenweise* in die Zelle *oberhalb* dieser Zelle ein.
Beachten Sie: Die Abgabe muss mit Ihrem *Gruppenaccount* erfolgen, den Sie in der Präsenzübung erhalten haben! Abgaben, die über Ihren Einzelaccount erfolgen, werden ignoriert. 
Achten Sie dabei streng auf die Formatierung, denn die Auswertung erfolgt automatisch. 
Benutzen Sie *keine* Leerzeilen.
Nur zwischen Matrikelnummer und E-Mail-Adresse darf ein Leerzeichen stehen.

Format:

MatrikelNummer E-Mail-Adresse

z.B.:

123456 testStudent@mail.upb.de

### Abgabetermin: 13.11.2016 um 8:00 Uhr

## Aufgabe 1: 
a) (1 Punkt) Gegeben ist folgendes Programm:
```Python
a = 5
b = 17
a = b + 2
print(a, b)
```
Was gibt das Programm aus? Erklären Sie kurz.

<!-- grade_id: 1, points: 1, grade: true, solution: true -->
```python

```

b) (1 Punkt) Sie ändern in der dritten Zeile $a$ auf $b$. Was gibt das Programm nun aus?  Erklären Sie kurz.

```Python
a = 5
b = 17
b = b + 2
print(a, b)
```

<!-- grade_id: 3, points: 1, grade: true, solution: true -->
```python

```

c) (1 Punkt) Sie ändern die Reihenfolge der Variablendarlegung. Was gibt das Programm nun aus?  Erklären Sie kurz.

```Python
b = 17
a = 5
a = b + 2
print(a, b)
```

<!-- grade_id: 6, points: 1, grade: true, solution: true -->
```python

```

d) (1 Punkt) Sie ändern in der dritten Zeile $a$ auf $d$. Was gibt das Programm nun aus? Erklären Sie kurz.

```Python
b = 17
a = 5
d = b + 2
print(a, b)
```

<!-- grade_id: 8, points: 1, grade: true, solution: true -->
```python

```

e) (1 Punkt) Sie ändern in der ersten UND dritten Zeile $a$ auf $d$. Was gibt das Programm nun aus?  Erklären Sie kurz.

```Python
d = 5
b = 17
d = b + 2
print(a, b)
```


<!-- grade_id: 10, points: 1, grade: true, solution: true -->
```python

```

## Aufgabe 2: Funktionen

a) (2 Punkte) Schreiben Sie fünf Funktionen namens ```add, multiply, divide``` und ```exponentiate```, die zwei als Parameter übergebene Werte miteinander addieren, subtrahieren, multiplizieren, dividieren und exponieren. Das entsprechende Ergebnis soll von den Funktionen zurückgegeben werden.

<!-- grade_id: 21, points: 0, grade: false, solution: true -->
```python
def add(a, b):
    ### BEGIN SOLUTION 
    print('ADDING ',a,' + ',b)
    out = a + b
    return out
    ### END SOLUTION

def multiply(a, b):
    ### BEGIN SOLUTION 
    print('MULTIPLYING ',a,' * ',b)
    out = a * b
    return out
    ### END SOLUTION
    
def divide(a, b):
    ### BEGIN SOLUTION 
    print('DIVIDING ',a,' / ',b)
    out = a / b
    return out
    ### END SOLUTION
    
def exponentiate(a, b):
    ### BEGIN SOLUTION 
    print('EXPONENTIATING ',a,' ** ',b)
    out = a ** b
    return out
    ### END SOLUTION
```

<!-- grade_id: 41, points: 2, grade: true, solution: false -->
```python
assert(add(4,6)==10)
assert(multiply(3,3)==9)
assert(divide(6,3)==2)
assert(exponentiate(2,3)==8)
```

Output:
```
ADDING  4  +  6
MULTIPLYING  3  *  3
DIVIDING  6  /  3
EXPONENTIATING  2  **  3
```

b) (2 Punkte) Schreiben Sie anschließend eine Funktion namens totalMoney, die drei Parameter k, n und p (in dieser Reihenfolge) übergeben bekommt und welche das Ergebnis der folgenden Rechnung zurückgibt.

$$K \left(1 + \frac{p}{100} \right)^n,$$

Verwenden Sie in dieser Aufgabe ausschließlich die in Teilaufgabe a) definierten Funktionen und keine Python-Rechenoperatoren.

<!-- grade_id: 12, points: 0, grade: false, solution: true -->
```python
def totalMoney(k,n,p):
    ### BEGIN SOLUTION 
    out = multiply(k,exponentiate(add(divide(p,100),1),n))
    return out
    ### END SOLUTION
```

<!-- grade_id: 81, points: 2, grade: true, solution: false -->
```python
assert(totalMoney(1000,3,5)==1157.6250000000002)
```

Output:
```
DIVIDING  5  /  100
ADDING  0.05  +  1
EXPONENTIATING  1.05  **  3
MULTIPLYING  1000  *  1.1576250000000001
```

## Aufgabe 3: Funktionen und Gültigkeitsbereich von Variablen

Das gegebene Programm ist fehlerhaft: 

```Python
def foo(x):
    x += 1
    print('x innerhalb foo = ', x)
    print('y innerhalb foo = ', y)
    y = 1
    print('y innerhalb foo = ', y)

foo(5)
y = 10
print('y außerhalb =', y)
```

a) (1 Punkt) Erklären Sie und kommentieren Sie die richtige Zeile aus damit das Programm Fehlerfrei läuft? 

<!-- grade_id: 52, points: 4, grade: false, solution: true -->
```python
def foo(x):
    x += 1
    print('x innerhalb foo = ', x)
#    print('y innerhalb foo = ', y)
    y = 1
    print('y innerhalb foo = ', y)
    return x

foo(5)
y = 10
print('y außerhalb =', y)
```

Output:
```
x innerhalb foo =  6
y innerhalb foo =  1
y außerhalb = 10
```

<!-- grade_id: 53, points: 1, grade: true, solution: false -->
```python
assert(foo(5)==6)
```

Output:
```
x innerhalb foo =  6
y innerhalb foo =  1
```

b) (1 Punkt) Gegeben ist die folgende Funktion: 

```Python
def f(a):
    return a+3
```

Welche Werte werden bei folgendem Aufruf ausgegeben:

```Python
a=2

print(f(a+3), f(f(a+3)))
```
Begründen Sie kurz.

<!-- grade_id: 57, points: 1, grade: true, solution: true -->
```python

```

c) (1 Punkt) Was ist der Rückgabewert des folgenden Funktionsaufrufs? Begründen Sie kurz.

```Python
def ausgabe(x):
    print(x,2*x)
    return 2*x
    print(x, 3*x)
    return 3*x
```

<!-- grade_id: 59, points: 1, grade: true, solution: true -->
```python

```

d) (2 Punkte) Gegeben ist folgendes Program. Welche Werte haben am Ende die in der drittletzten Zeile deklarierten Variablen $a,b,x,y$? Begründen Sie kurz.

```Python
def foo(x, y):
    a = 24
    x,y = y,x
    b = 4
    b = 16
    c = 10
    print(a,b,x,y)

a,b,x,y = 1,15,3,4
foo(17,4)
print(a,b,x,y)
```

<!-- grade_id: 61, points: 2, grade: true, solution: true -->
```python

```

e) (2 Punkte) Gegeben sind die folgenden beiden Funktionen:

```Python
def foo1(a):
     print(a + 1)
def foo2(a):
     return a + 1
```     
Einer der beiden folgenden Aufrufe ist fehlerhaft. Finden Sie, welcher es ist und begründen Sie kurz.
```Python
foo1(3) + 1
foo2(3) + 1
```

<!-- grade_id: 63, points: 2, grade: true, solution: true -->
```python

```

## Aufgabe 4: Multiple Eingabe- und Rückgabewerten

a) (2 Punkte) Schreiben Sie eine Funktion, die eine Zeitdauer in Sekunden als Parameter übergeben bekommt. Die Funktion soll dabei die eingegebene Parameter in Tagen (d), Stunden (h), Minuten (m) und Sekunden (s) konvertieren und zurückgeben. Drucken Sie anschließend die zurückgegebene Werten aus. 


<!-- grade_id: 71, points: 0, grade: false, solution: true -->
```python
def convSec(s):
    ###BEGIN SOLUTION###
    d = s // (3600)
    m = (s%3600) // 60
    s -= d * 3600 + m * 60
    
    return d, m, s

    
d, m, s = convSec(3670)    
print("days: ", d, " mins: ", m, "secs: ", s)   
    ###END SOLUTION####
```

Output:
```
days:  1  mins:  1 secs:  10
```

<!-- grade_id: 72, points: 2, grade: true, solution: false -->
```python
d, m, s = convSec(3670)   
assert(d==1)
assert(m==1)
assert(s==10)
```

b) (2 Punkte) Schreiben Sie eine weitere Funktion, die die umgekehrte Funktionalität hat. Eingabeparametern sind die Anzahl der Tagen, Stunden, Minuten und Sekunden und die Rückgabewert ist die vergangene Zeitdauer in Sekunden. Drucken Sie entsprechend die 

<!-- grade_id: 74, points: 0, grade: false, solution: true -->
```python
def convDate(d, m, s):
    ###BEGIN SOLUTION###
    s += d*3600 + m *60
    
    return s

    
s = convDate(1, 1, 10)    
print("seconds: ", s)   
    ###END SOLUTION####
```

Output:
```
seconds:  3670
```

<!-- grade_id: 75, points: 2, grade: true, solution: false -->
```python
s = convDate(1, 1, 10)   
assert(s==3670)
```